[1557772729] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11663736 kB) free.
[1557773220] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1557774109] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 307.74, 224.81, 105.92
[1557774229] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 53.61, 161.17, 97.54
[1557774349] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 7.47, 107.93, 85.74
[1557774420] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1557774540] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1557775934] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (170924 kB) free!
[1557776054] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (169064 kB) free!
[1557776083] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (177208 kB) free!
[1557776149] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.44, 0.58, 12.56
[1557776174] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (147360 kB) free!
[1557776203] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (167824 kB) free!
[1557776260] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 341 processes
[1557776323] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156308 kB) free!
[1557776329] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.0% (5436432 kB) free!
[1557776380] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 340 processes
[1557776449] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11650672 kB) free.
[1557776501] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 343 processes
[1557776940] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1557777060] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1557777101] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1557777349] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.27, 0.24, 3.60
[1557777649] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7244164 kB) free!
[1557777661] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1557777769] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.7% (5280912 kB) free!
[1557777781] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1557777890] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11660004 kB) free.
[1557778381] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557778489] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8346676 kB) free!
[1557778501] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1557778609] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11656628 kB) free.
[1557779101] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1557779209] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.8% (5313296 kB) free!
[1557779221] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1557779329] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11655668 kB) free.
[1557779821] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557779941] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557780133] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 287.00, 207.57, 97.81
[1557780253] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 287.56, 234.09, 120.84
[1557780373] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 287.90, 251.88, 141.09
[1557780529] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5951128 kB) free!
[1557780541] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1557780649] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11650328 kB) free.
[1557780662] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1557781249] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.5% (6151716 kB) free!
[1557781261] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1557781369] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11655524 kB) free.
[1557781381] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1557781969] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.6% (6208760 kB) free!
[1557781981] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557782089] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11658344 kB) free.
[1557782101] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1557782690] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6312708 kB) free!
[1557782701] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557782773] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.54, 1.75, 12.21
[1557782810] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11655932 kB) free.
[1557782821] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1557782821] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557783973] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.69, 1.55, 4.44
[1557784010] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7005544 kB) free!
[1557784022] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557784130] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11655444 kB) free.
[1557784142] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1557784730] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6913868 kB) free!
[1557784742] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557784850] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11644564 kB) free.
[1557784862] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1557785450] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6831568 kB) free!
[1557785462] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 353 processes
[1557785570] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1557785570] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11653028 kB) free.
[1557785582] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557786170] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6735120 kB) free!
[1557786182] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1557786290] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11656920 kB) free.
[1557786302] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1557786890] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6726500 kB) free!
[1557786902] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1557787011] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11654292 kB) free.
[1557787023] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1557787610] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6587136 kB) free!
[1557787622] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557787730] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11655400 kB) free.
[1557787742] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1557788330] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6428492 kB) free!
[1557788342] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557788450] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11643260 kB) free.
[1557788462] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1557789050] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6351916 kB) free!
[1557789062] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1557789170] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11654460 kB) free.
[1557789182] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557789770] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6768144 kB) free!
[1557789782] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557789890] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11653024 kB) free.
[1557789903] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1557790490] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6721152 kB) free!
[1557790502] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1557790610] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11650616 kB) free.
[1557790622] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557790742] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1557791210] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9867600 kB) free!
[1557791331] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8275356 kB) free!
[1557791451] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.2% (6512156 kB) free!
[1557792051] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11615208 kB) free.
[1557792196] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 127.43, 183.31, 97.70
[1557792316] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 17.34, 122.70, 85.87
[1557792436] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.85, 82.31, 75.54
[1557792542] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1557793142] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557793251] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8722164 kB) free!
[1557793263] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1557793371] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7062760 kB) free!
[1557793383] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557793491] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11604348 kB) free.
[1557793983] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1557793983] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557794236] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.83, 0.82, 11.33
[1557794583] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557794691] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9029388 kB) free!
[1557794703] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1557794811] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7527992 kB) free!
[1557794823] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557794931] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.1% (5490928 kB) free!
[1557795436] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.70, 0.52, 3.43
[1557795532] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1557795532] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10907260 kB) free.
[1557795949] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 229.21, 135.13, 56.93
[1557796069] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 231.74, 167.23, 78.18
[1557796132] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9295048 kB) free!
[1557796189] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 47.72, 125.54, 74.29
[1557796252] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.7% (7761376 kB) free!
[1557796372] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.6% (5737672 kB) free!
[1557796972] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10677780 kB) free.
[1557797572] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9061780 kB) free!
[1557797692] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.2% (7480900 kB) free!
[1557797812] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.9% (5357916 kB) free!
[1557797823] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1557797989] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.24, 0.55, 10.90
[1557798412] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10518092 kB) free.
[1557798423] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557798543] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557798663] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557799012] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9388488 kB) free!
[1557799132] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7837040 kB) free!
[1557799189] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.38, 0.32, 3.21
[1557799252] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.0% (5895460 kB) free!
[1557799853] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11134772 kB) free.
[1557800453] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9774400 kB) free!
[1557800573] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8248916 kB) free!
[1557800693] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.2% (6521212 kB) free!
[1557801293] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11531528 kB) free.
[1557802263] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1557802493] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8771684 kB) free!
[1557802613] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.6% (7202496 kB) free!
[1557802733] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11523572 kB) free.
[1557802863] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1557802983] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557803103] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1557803297] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.9% (14614740 kB) free.
[1557803342] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 83.1% (13510724 kB) free.
[1557803703] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1557803933] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8640908 kB) free!
[1557804053] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.2% (7018808 kB) free!
[1557804174] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11495896 kB) free.
[1557804304] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1557804424] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557804544] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1557804773] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9804976 kB) free!
[1557804893] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8160056 kB) free!
[1557805013] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.1% (6457484 kB) free!
[1557805144] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1557805613] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11458384 kB) free.
[1557805744] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1557805864] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557805984] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1557806584] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557806813] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8729956 kB) free!
[1557806933] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.6% (7203408 kB) free!
[1557807053] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.4% (5149536 kB) free!
[1557807185] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557807305] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1557807425] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1557807586] SERVICE ALERT: secondary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1557807653] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10668204 kB) free.
[1557807706] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1557808036] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 76.26, 17.25, 6.03
[1557808156] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.73, 11.68, 5.35
[1557808253] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8835820 kB) free!
[1557808276] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.79, 7.98, 4.76
[1557808373] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.6% (7171508 kB) free!
[1557808493] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4873584 kB) free!
[1557809094] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10508772 kB) free.
[1557809428] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 631 processes
[1557809489] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 636 processes
[1557809694] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8704224 kB) free!
[1557809814] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7069216 kB) free!
[1557809934] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4873352 kB) free!
[1557810041] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 47.0% (23172320 kB) free.
[1557810319] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 104999 MB (46% inode=99%):
[1557810425] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1557811026] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1557811146] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1557811266] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1557812466] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1557813066] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1557813186] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1557813306] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1557813328] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (4715576 kB) free!
[1557813337] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.8% (3374088 kB) free!
[1557813448] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.0% (4677508 kB) free!
[1557813457] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.5% (3300968 kB) free!
[1557813568] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (4646384 kB) free!
[1557813577] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.1% (3221036 kB) free!
[1557813907] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1557814177] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1675776 kB) free!
[1557814507] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1557814627] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557814747] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557815347] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1557815947] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1557816067] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1557816187] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557816444] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.10, 5.89, 5.39
[1557816564] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.09, 5.56, 5.33
[1557816568] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2396928 kB) free!
[1557816685] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.73, 5.32, 5.27
[1557816788] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1557817109] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 15.4% (2495152 kB) free!
[1557817229] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 14.1% (2287936 kB) free!
[1557817349] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 12.9% (2098656 kB) free!
[1557817388] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1557817508] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1557817628] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557817885] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.36, 4.19, 4.62
[1557817949] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1281620 kB) free!
[1557819988] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 643 processes
[1557820048] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 647 processes
[1557820249] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1557820747] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.9% (23172056 kB) free.
[1557820797] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.4% (22649608 kB) free.
[1557820840] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.5% (14716804 kB) free.
[1557821686] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4730288 kB) free!
[1557821806] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.0% (4665660 kB) free!
[1557821926] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (4610296 kB) free!
[1557826249] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9512980 kB) free!
[1557826301] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1557826369] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 32.8% (16171068 kB) free.
[1557826421] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1557827828] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1557827927] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2288996 kB) free!
[1557828148] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 661 processes
[1557828208] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 655 processes
[1557828268] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 652 processes
[1557828428] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1557828548] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1557828668] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1557828973] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 303.22, 252.84, 133.77
[1557829093] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 95.33, 206.47, 132.40
[1557829213] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14.44, 138.77, 116.57
[1557829249] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.4% (2821944 kB) free!
[1557829268] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1557829370] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.4% (2658548 kB) free!
[1557829439] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.7% (2882908 kB) free!
[1557829490] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.2% (2478564 kB) free!
[1557829559] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.0% (2759196 kB) free!
[1557829679] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.4% (2659584 kB) free!
[1557829759] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174128 kB) free!
[1557829808] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195152 kB) free!
[1557829869] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1557829879] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168276 kB) free!
[1557829928] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177292 kB) free!
[1557829989] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1557829999] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171068 kB) free!
[1557830049] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174924 kB) free!
[1557830109] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1557830690] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (812688 kB) free!
[1557830879] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1623852 kB) free!
[1557831013] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.90, 2.21, 18.65
[1557831769] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8643820 kB) free!
[1557831890] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.7% (15627072 kB) free.
[1557832813] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.82, 2.03, 4.40
[1557834110] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.0% (14796352 kB) free.
[1557834164] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.4% (22805696 kB) free.
[1557834214] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.7% (22226340 kB) free.
[1557834264] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 77.2% (12543644 kB) free.
[1557834289] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9065852 kB) free!
[1557834306] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.8% (22013420 kB) free.
[1557834309] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1557834409] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.2% (15398924 kB) free.
[1557834909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1557835029] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1557835149] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1557836989] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 285.75, 185.70, 82.16
[1557837109] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 160.98, 187.90, 96.53
[1557837229] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 22.10, 125.88, 84.88
[1557837676] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 273.05, 193.12, 89.23
[1557837796] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 85.18, 162.66, 91.74
[1557837916] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 11.89, 108.99, 80.68
[1557839029] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.19, 0.63, 12.55
[1557839716] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.67, 0.73, 11.96
[1557840829] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.32, 0.46, 3.30
[1557840916] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.29, 0.36, 3.51
[1557842347] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (169168 kB) free!
[1557842467] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (168952 kB) free!
[1557842587] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (170148 kB) free!
[1557845021] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 345 processes
[1557845141] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1557845349] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1557845618] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.6% (3815412 kB) free!
[1557845739] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.4% (3777320 kB) free!
[1557845810] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8700536 kB) free!
[1557845859] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.3% (3749224 kB) free!
[1557845930] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.3% (14939492 kB) free.
[1557845949] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557846069] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1557846189] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1557846459] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 62.3% (15266464 kB) free.
[1557847730] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9436600 kB) free!
[1557847850] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.9% (15229268 kB) free.
[1557852189] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557852789] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1557852909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1557852941] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 349 processes
[1557853029] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1557853061] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 348 processes
[1557853181] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 347 processes
[1557853471] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (159280 kB) free!
[1557853591] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (167360 kB) free!
[1557853711] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169788 kB) free!
[1557854381] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1557854911] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 68.3% (11096232 kB) free.
[1557854982] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 347 processes
[1557855102] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 346 processes
[1557855222] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 344 processes
[1557855511] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (183984 kB) free!
[1557855631] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 74.4% (12100196 kB) free.
[1557856231] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (150900 kB) free!
[1557856351] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (185812 kB) free!
[1557856422] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1557856422] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1557856471] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 88.6% (14394108 kB) free.