[1567190264] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1567190384] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1567191856] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9862244 kB) free!
[1567191976] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.0% (9343428 kB) free!
[1567192096] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.9% (8844712 kB) free!
[1567193625] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.0% (3664664 kB) free!
[1567193741] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 294 processes
[1567193746] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.9% (3649480 kB) free!
[1567193866] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.8% (3636272 kB) free!
[1567194095] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (255240 kB) free!
[1567194214] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (3371260 kB) free!
[1567194215] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (972984 kB) free!
[1567194284] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4911984 kB) free!
[1567194334] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8301024 kB) free!
[1567194335] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (946748 kB) free!
[1567194341] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 288 processes
[1567194404] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (4908084 kB) free!
[1567194454] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (6474868 kB) free!
[1567194524] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 20.0% (4901044 kB) free!
[1567195097] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (1733996 kB) free!
[1567195177] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.03, 6.33, 5.48
[1567195297] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.89, 6.49, 5.64
[1567195417] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.13, 6.32, 5.68
[1567196384] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1567196504] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1567196617] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.12, 2.51, 4.24
[1567197704] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1567197825] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1567198346] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (2484408 kB) free!
[1567198466] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (229364 kB) free!
[1567198586] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (230096 kB) free!
[1567199137] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 87951 MB (39% inode=99%):
[1567199186] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.0% (43344600 kB) free.
[1567201066] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195284 kB) free!
[1567201124] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170480 kB) free!
[1567201666] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 46.1% (11304056 kB) free.
[1567201724] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 50.1% (12273656 kB) free.
[1567211735] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 96.2% (33701676 kB) free.
[1567211854] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.8% (22082200 kB) free!
[1567213654] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 41.9% (55220712 kB) free.
[1567214223] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 51.8% (12787292 kB) free.
[1567215341] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 355 processes
[1567215462] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 367 processes
[1567216623] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183960 kB) free!
[1567216666] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (4214044 kB) free!
[1567216743] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186500 kB) free!
[1567216786] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.1% (4197624 kB) free!
[1567216864] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180940 kB) free!
[1567216906] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (4172596 kB) free!
[1567219124] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4812524 kB) free!
[1567219244] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (4802824 kB) free!
[1567219364] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (4793120 kB) free!
[1567220624] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1567220744] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1567222306] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2270192 kB) free!
[1567224335] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 306 processes
[1567224455] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1567226880] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.6% (40708556 kB) free.
[1567226923] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.4% (22405552 kB) free.
[1567226962] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.8% (23141064 kB) free.
[1567228364] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2352636 kB) free!
[1567230567] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187032 kB) free!
[1567230687] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192060 kB) free!
[1567230807] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179764 kB) free!
[1567232886] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9499744 kB) free!
[1567233006] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (9499380 kB) free!
[1567233126] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.3% (9499412 kB) free!
[1567234544] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1567234664] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1567235526] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (1811496 kB) free!
[1567236224] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.45, 5.31, 5.10
[1567236344] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.57, 5.09, 5.05
[1567236464] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.34, 4.96, 5.01
[1567237064] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.88, 4.99, 4.99
[1567240417] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.77, 5.58, 5.29
[1567240537] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.27, 4.89, 5.07
[1567240586] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8816792 kB) free!
[1567240657] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.42, 4.44, 4.88
[1567240664] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1567240706] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.3% (8510512 kB) free!
[1567240784] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1567240826] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.7% (8213344 kB) free!
[1567241257] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.23, 5.78, 5.16
[1567241377] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.89, 4.85, 4.91
[1567241923] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.9% (4139948 kB) free!
[1567242043] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.6% (3826808 kB) free!
[1567242164] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.3% (3737976 kB) free!
[1567242626] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4643112 kB) free!
[1567242764] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1887120 kB) free!
[1567244048] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.3% (20558724 kB) free.
[1567244086] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.8% (22006784 kB) free.
[1567244126] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.7% (40778608 kB) free.
[1567244158] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 30.8% (15158728 kB) free.
[1567244192] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 74.9% (18348680 kB) free.
[1567244656] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185808 kB) free!
[1567244776] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171896 kB) free!
[1567244896] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168580 kB) free!
[1567248297] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.0% (3179628 kB) free!
[1567248333] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.9% (8314568 kB) free!
[1567248417] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.8% (2645344 kB) free!
[1567248453] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.6% (8169324 kB) free!
[1567248537] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1567248537] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2130296 kB) free!
[1567248573] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.2% (7965952 kB) free!
[1567249773] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (1752048 kB) free!
[1567250195] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.9% (3167720 kB) free!
[1567250316] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.4% (2800424 kB) free!
[1567250436] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2446640 kB) free!
[1567252177] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.55, 6.68, 5.44
[1567252297] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.37, 6.25, 5.44
[1567252417] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.01, 5.21, 5.17
[1567253017] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.97, 3.17, 3.96
[1567254217] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.77, 6.94, 5.31
[1567254337] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.76, 5.34, 4.93
[1567254537] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1567256264] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.2% (22355752 kB) free.
[1567256305] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 75.7% (37292368 kB) free.
[1567256392] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1567256395] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1567256418] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 72.5% (17759256 kB) free.
[1567256516] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 351 processes
[1567256516] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 7.38, 5.63, 3.93
[1567256538] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.5% (21350260 kB) free.
[1567257029] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192704 kB) free!
[1567257116] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.67, 6.07, 5.12
[1567257145] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (2315080 kB) free!
[1567257151] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193544 kB) free!
[1567257236] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.24, 5.79, 5.14
[1567257265] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (1594932 kB) free!
[1567257271] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190292 kB) free!
[1567257356] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.69, 5.22, 5.02
[1567257385] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (438224 kB) free!
[1567257476] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184120 kB) free!
[1567257596] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178264 kB) free!
[1567257716] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191840 kB) free!
[1567257956] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.53, 2.48, 3.69
[1567259737] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.86, 7.87, 5.47
[1567259857] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.07, 7.61, 5.67
[1567259977] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.57, 6.12, 5.36
[1567260516] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 12.6% (6224268 kB) free!
[1567260577] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1567260577] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.53, 3.22, 4.11
[1567260636] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 11.2% (5506416 kB) free!
[1567260756] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4751364 kB) free!
[1567265133] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.2% (22254396 kB) free.
[1567265167] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.7% (22466780 kB) free.
[1567265203] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 71.3% (35115280 kB) free.
[1567265242] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.2% (22365200 kB) free.
[1567265734] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182276 kB) free!
[1567265854] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189676 kB) free!
[1567265974] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1567265974] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182452 kB) free!
[1567268165] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (230404 kB) free!
[1567268285] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (214280 kB) free!
[1567268405] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (236660 kB) free!
[1567268447] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.1% (23222960 kB) free.
[1567268481] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 31.9% (15731784 kB) free.
[1567269584] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1567269704] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1567269713] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 908 processes
[1567269833] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 396 processes
[1567269956] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 133.46, 121.14, 66.30
[1567270076] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 18.71, 81.34, 58.37
[1567270196] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.58, 55.10, 51.53
[1567270425] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 327 processes
[1567270545] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1567271396] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.04, 2.07, 15.07
[1567271804] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.5% (7660760 kB) free!
[1567271924] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.5% (7658536 kB) free!
[1567271977] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567272045] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.5% (7659444 kB) free!
[1567272596] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.84, 0.91, 4.80
[1567275045] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4031364 kB) free!