[1591987507] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1591987507] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 3.82, 15.78, 25.64
[1591987591] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1591987717] SERVICE ALERT: smscapp2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=8.72% system=28.05% iowait=0.25% idle=62.98%
[1591988707] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 42.52, 34.97, 32.36
[1591988827] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 28.41, 32.59, 31.86
[1591988947] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 26.95, 31.26, 31.51
[1591990254] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 1.4% (350388 kB) free!
[1591991454] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 86.0% (21077244 kB) free.
[1591992116] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (1556124 kB) free!
[1591992240] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (230880 kB) free!
[1591992360] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235328 kB) free!
[1591993747] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 17.01, 23.93, 28.45
[1591994160] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 58.4% (28791464 kB) free.
[1591994347] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 41.71, 30.09, 28.84
[1591994467] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 36.30, 30.73, 29.11
[1591995067] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 39.26, 31.37, 30.60
[1591995187] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 27.72, 29.29, 29.93
[1591996862] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 87.1% (42917672 kB) free.
[1591997446] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(188) < 103.16.101.72 [13/Jun/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1591997547] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.7% (43723832 kB) free.
[1591997549] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1264) < 172.105.59.173 [13/Jun/2020:03:01:32 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1591997566] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591997669] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591997694] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(238) < 2020/06/13 03:01:22 [error] 12475#0: *94239540 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.62.98, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1591997814] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591998787] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 52.24, 41.29, 33.74
[1591998907] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 27.57, 34.48, 32.08
[1591999027] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 32.58, 36.04, 33.12
[1591999211] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.2% (4218596 kB) free!
[1591999812] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (252300 kB) free!
[1592000227] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 26.11, 24.75, 27.56
[1592003227] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 34.94, 35.45, 32.73
[1592003347] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 18.63, 30.68, 31.41
[1592003467] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 25.31, 29.01, 30.69
[1592004012] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.6% (10442452 kB) free.
[1592004667] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 36.47, 25.86, 26.98
[1592005212] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (235300 kB) free!
[1592005332] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (260648 kB) free!
[1592005452] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1592005452] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (253660 kB) free!
[1592005657] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 76.0% (18630584 kB) free.
[1592006309] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.7% (3393016 kB) free!
[1592006467] SERVICE ALERT: mrq4;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 57.64, 39.32, 31.07
[1592006587] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 29.60, 34.50, 30.28
[1592006707] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 43.35, 36.94, 31.63
[1592006909] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166700 kB) free!
[1592007308] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 20.30, 25.50, 29.31
[1592007908] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 32.11, 33.25, 31.78
[1592008028] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 45.45, 35.27, 32.51
[1592008148] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 30.13, 31.55, 31.41
[1592008457] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592008582] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1592008721] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592008850] SERVICE ALERT: smscapp2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=9.09% system=27.96% iowait=0.04% idle=62.91%
[1592009234] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 65.0% (15936804 kB) free.
[1592010548] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 109.17, 78.05, 57.89
[1592010729] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 29.5% (38865020 kB) free.
[1592010954] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (1828144 kB) free!
[1592011014] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (1831840 kB) free!
[1592011074] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 7.5% (1830728 kB) free!
[1592011134] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 7.4% (1825376 kB) free!
[1592011851] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592011976] SERVICE ALERT: smscapp2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=8.90% system=27.92% iowait=0.32% idle=62.86%
[1592014254] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1592014374] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1592015756] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.18, 6.96, 5.18
[1592015876] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.84, 6.43, 5.23
[1592015996] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.68, 5.56, 5.07
[1592016252] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1592017196] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.03, 4.26, 4.70
[1592019181] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592019308] SERVICE ALERT: smscapp2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=12.76% system=27.42% iowait=4.67% idle=55.15%
[1592019596] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.71, 7.62, 6.14
[1592019716] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.80, 7.03, 6.12
[1592019836] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.87, 7.01, 6.22
[1592019911] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592020041] SERVICE ALERT: smscapp2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1592020148] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1592020167] SERVICE ALERT: smscapp2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=11.46% system=27.66% iowait=1.74% idle=59.13%
[1592023634] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.8% (3863608 kB) free!
[1592023752] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 844.43, 225.18, 77.62
[1592023754] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.6% (3567572 kB) free!
[1592023872] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 135.26, 169.70, 76.14
[1592023874] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.6% (3334672 kB) free!
[1592024070] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(18) < 203.212.70.170 [13/Jun/2020:10:24:20 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1592024190] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592024474] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2213656 kB) free!
[1592024790] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6) < 203.212.70.170 [13/Jun/2020:10:32:17 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1592024910] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(12) < 203.212.70.170 [13/Jun/2020:10:38:13 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1592025030] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592025341] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1592025461] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1592025581] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1592026169] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4272436 kB) free!
[1592026230] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 203.212.70.170 [13/Jun/2020:10:53:38 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1592026274] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.5% (2817452 kB) free!
[1592026289] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.9% (3784136 kB) free!
[1592026350] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592026409] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.5% (3262740 kB) free!
[1592026781] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 368 processes
[1592026874] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1592026874] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (762156 kB) free!
[1592026950] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(9) < 203.212.70.170 [13/Jun/2020:11:08:03 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1592027009] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (1338600 kB) free!
[1592027070] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592029057] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184640 kB) free!
[1592029177] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186740 kB) free!
[1592029181] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1592029297] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188700 kB) free!
[1592029301] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1592029304] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1592029421] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1592029424] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1592029544] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1592030009] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.5% (15840320 kB) free.
[1592030021] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1592030021] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 917 processes
[1592030099] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 46.5% (11389724 kB) free.
[1592030144] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1592030179] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 913 processes
[1592030299] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 911 processes
[1592030419] SERVICE ALERT: sservercluster2;Total Processes;OK;HARD;1;PROCS OK: 344 processes
[1592030670] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(16) < 172.105.59.173 [13/Jun/2020:12:04:56 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1592030790] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(27) < 172.105.59.173 [13/Jun/2020:12:16:25 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1592030911] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592031510] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 172.105.57.57 [13/Jun/2020:12:27:10 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1592031630] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592032230] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(15) < 203.212.70.170 [13/Jun/2020:12:31:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1592032350] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592033547] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9528720 kB) free!
[1592033667] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.8% (9276504 kB) free!
[1592033770] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592033787] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.3% (9033208 kB) free!
[1592033890] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 82.1% (40464956 kB) free.
[1592034150] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 172.105.57.57 [13/Jun/2020:13:12:06 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1592034270] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592034899] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.1% (4669804 kB) free!
[1592035019] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.6% (4563632 kB) free!
[1592035139] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.1% (4425416 kB) free!
[1592036187] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4733932 kB) free!
[1592036939] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2456928 kB) free!
[1592037747] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 20.0% (4907028 kB) free!
[1592037867] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.9% (4872204 kB) free!
[1592037987] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.7% (4836124 kB) free!
[1592038274] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1592038409] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.5% (3255152 kB) free!
[1592038529] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.0% (2916936 kB) free!
[1592038649] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.6% (2610976 kB) free!
[1592039250] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1592568 kB) free!
[1592040870] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(174) < 185.255.8.58 [13/Jun/2020:15:00:12 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1592040991] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1592041050] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1592041050] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.9% (15927924 kB) free.
[1592041100] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592041193] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592041193] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592041223] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 83.3% (41030624 kB) free.
[1592041316] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38770 MB (75% inode=98%):
[1592041316] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2304 processes
[1592041421] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1592044137] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 19.1% (6250796 kB) free!
[1592044257] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 22.0% (7210848 kB) free.
[1592045457] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 19.2% (6292192 kB) free!
[1592045577] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 18.9% (6172332 kB) free!
[1592045697] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 18.5% (6052640 kB) free!
[1592048462] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9565832 kB) free!
[1592048582] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (9464892 kB) free!
[1592048703] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (9337836 kB) free!
[1592050650] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.3% (3657948 kB) free!
[1592050770] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.4% (3468312 kB) free!
[1592050890] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.6% (3271020 kB) free!
[1592051491] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.5% (15839920 kB) free.
[1592051523] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592051523] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592051646] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38769 MB (75% inode=98%):
[1592051647] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2308 processes
[1592052253] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592052253] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592052381] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38769 MB (75% inode=98%):
[1592052382] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2308 processes
[1592053497] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (3196772 kB) free!
[1592055303] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4607944 kB) free!
[1592056436] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.08, 4.22, 4.94
[1592057036] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.88, 6.08, 5.37
[1592057156] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.55, 5.54, 5.26
[1592057276] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.37, 5.45, 5.25
[1592060491] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.7% (3977496 kB) free!
[1592060611] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.5% (3715704 kB) free!
[1592060731] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.6% (3500160 kB) free!
[1592061129] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.8% (24745108 kB) free!
[1592061249] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.0% (19708796 kB) free!
[1592061369] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.2% (16092456 kB) free!
[1592062498] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 62.7% (20533552 kB) free.
[1592062531] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.5% (15826656 kB) free.
[1592062583] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592062583] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1592062676] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.94, 4.76, 4.89
[1592062703] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38768 MB (75% inode=98%):
[1592062704] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2305 processes
[1592064548] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 17.80, 33.06, 42.75
[1592065148] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 76.83, 49.54, 45.14
[1592065374] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1592065494] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 334 processes
[1592065614] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 337 processes
[1592065748] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 36.75, 40.32, 42.65
[1592066214] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1592071614] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1592071734] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1592071748] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1592071748] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 55.74, 42.60, 40.79
[1592072348] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 37.41, 37.93, 40.73
[1592072731] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)