[1566153464] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8306104 kB) free!
[1566153502] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566153584] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.4% (6580212 kB) free!
[1566153622] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1566153704] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11612812 kB) free.
[1566154223] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1566154304] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7079448 kB) free!
[1566154343] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566154424] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11613124 kB) free.
[1566154463] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566154753] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190952 kB) free!
[1566154873] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167444 kB) free!
[1566154993] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187868 kB) free!
[1566155025] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7294988 kB) free!
[1566155063] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566155145] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11605612 kB) free.
[1566155183] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1566155615] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (3154368 kB) free!
[1566155735] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.4% (3151476 kB) free!
[1566155745] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7772004 kB) free!
[1566155783] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1566155855] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 19.4% (3155376 kB) free!
[1566155865] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11614080 kB) free.
[1566155903] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1566156193] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 59.5% (14690684 kB) free.
[1566156465] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8220404 kB) free!
[1566156503] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566156585] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.4% (6590924 kB) free!
[1566156624] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1566156683] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4743484 kB) free!
[1566156705] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11611788 kB) free.
[1566156803] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4746860 kB) free!
[1566156923] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (4743352 kB) free!
[1566157223] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1566157305] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6681168 kB) free!
[1566157343] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1566157425] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11614004 kB) free.
[1566157943] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566158025] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6791948 kB) free!
[1566158063] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566158145] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11606220 kB) free.
[1566158483] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9306752 kB) free!
[1566158603] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.9% (8806584 kB) free!
[1566158663] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566158723] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.9% (8803984 kB) free!
[1566158745] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7280416 kB) free!
[1566158783] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1566158856] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1563848 kB) free!
[1566158866] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11613804 kB) free.
[1566158903] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566159465] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7840480 kB) free!
[1566159503] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566159585] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11614880 kB) free.
[1566159624] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1566160185] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8060784 kB) free!
[1566160223] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1566160305] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.0% (6923584 kB) free!
[1566160343] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1566160425] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11613036 kB) free.
[1566160943] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566161025] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6575448 kB) free!
[1566161063] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566161123] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 21.5% (5280964 kB) free.
[1566161145] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11603888 kB) free.
[1566161593] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (208224 kB) free!
[1566161663] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566161713] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (200108 kB) free!
[1566161745] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7278984 kB) free!
[1566161783] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1566161833] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (209192 kB) free!
[1566161865] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11611704 kB) free.
[1566161903] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566162465] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7562736 kB) free!
[1566162504] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566162585] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11613292 kB) free.
[1566162624] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1566163185] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8071832 kB) free!
[1566163224] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1566163306] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.3% (6547092 kB) free!
[1566163344] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1566163426] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11612044 kB) free.
[1566163944] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1566164026] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7194040 kB) free!
[1566164064] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566164123] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165148 kB) free!
[1566164146] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11602276 kB) free.
[1566164184] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1566164243] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (200140 kB) free!
[1566164364] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166272 kB) free!
[1566164746] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7726532 kB) free!
[1566164784] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1566164866] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11611168 kB) free.
[1566164905] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1566164964] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1566164964] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 49.8% (12215152 kB) free.
[1566165466] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8144224 kB) free!
[1566165504] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566165586] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.9% (6380204 kB) free!
[1566165624] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1566165706] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11611344 kB) free.
[1566166224] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1566166306] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6933728 kB) free!
[1566166344] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566166426] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11612640 kB) free.
[1566166464] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566167026] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7622148 kB) free!
[1566167064] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566167147] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11602544 kB) free.
[1566167184] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1566167746] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7914924 kB) free!
[1566167785] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1566167866] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11612316 kB) free.
[1566167905] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566168466] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8179420 kB) free!
[1566168505] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566168586] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6392496 kB) free!
[1566168625] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1566168706] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11607932 kB) free.
[1566169225] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1566169306] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7114212 kB) free!
[1566169345] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566169426] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11610172 kB) free.
[1566169465] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566170026] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7348780 kB) free!
[1566170065] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566170146] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11604616 kB) free.
[1566170185] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1566170746] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7842788 kB) free!
[1566170786] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1566170866] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11608896 kB) free.
[1566170906] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1566171467] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8256428 kB) free!
[1566171506] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566171587] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.4% (6587920 kB) free!
[1566171626] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1566171707] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11610128 kB) free.
[1566172226] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1566172307] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7087884 kB) free!
[1566172346] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566172427] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11610824 kB) free.
[1566172466] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566173027] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7344244 kB) free!
[1566173067] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566173147] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11601768 kB) free.
[1566173187] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1566173747] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7761700 kB) free!
[1566173787] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1566173867] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11610648 kB) free.
[1566173907] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566174467] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7987772 kB) free!
[1566174507] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566174587] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6399168 kB) free!
[1566174627] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1566174707] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11610324 kB) free.
[1566175227] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1566175307] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7613840 kB) free!
[1566175347] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566175428] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11503360 kB) free.
[1566175467] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566175751] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 586 processes
[1566175811] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 582 processes
[1566176027] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8284608 kB) free!
[1566176068] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566176147] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.2% (6527184 kB) free!
[1566176188] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1566176267] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11609948 kB) free.
[1566176364] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1566176723] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3759372 kB) free!
[1566176788] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1566176867] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7073968 kB) free!
[1566176908] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566176987] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11610132 kB) free.
[1566177028] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1566177433] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.7% (19904456 kB) free.
[1566177533] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1566177587] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7630084 kB) free!
[1566177628] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1566177653] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1566177707] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11608696 kB) free.
[1566177748] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566177773] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1566178307] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7982696 kB) free!
[1566178348] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566178427] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11607728 kB) free.
[1566178469] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566179027] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8111348 kB) free!
[1566179068] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566179147] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.7% (6281140 kB) free!
[1566179188] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1566179268] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11606036 kB) free.
[1566179788] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1566179867] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6485708 kB) free!
[1566179908] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566179933] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 365 processes
[1566179987] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11609056 kB) free.
[1566180053] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1566180508] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566180587] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6646120 kB) free!
[1566180628] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1566180707] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11606064 kB) free.
[1566181228] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1566181307] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6963420 kB) free!
[1566181349] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566181427] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11607068 kB) free.
[1566181469] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566182027] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7256140 kB) free!
[1566182069] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566182147] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11598884 kB) free.
[1566182189] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1566182747] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8091256 kB) free!
[1566182789] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1566182867] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.0% (6874972 kB) free!
[1566182909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566182987] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11605876 kB) free.
[1566183029] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1566183588] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7901188 kB) free!
[1566183629] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1566183702] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 621.06, 486.20, 241.23
[1566183708] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11504808 kB) free.
[1566183749] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1566183822] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 99.72, 339.75, 217.99
[1566183942] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 13.74, 227.44, 191.60
[1566184308] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8306788 kB) free!
[1566184350] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566184428] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.7% (6773980 kB) free!
[1566184470] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566184548] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11599052 kB) free.
[1566184590] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1566185148] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7194348 kB) free!
[1566185190] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1566185268] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11605528 kB) free.
[1566185310] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566185868] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7722716 kB) free!
[1566185910] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566185988] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.7% (6257276 kB) free!
[1566186030] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1566186108] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11602756 kB) free.
[1566186342] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.69, 0.93, 15.20
[1566186630] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1566186708] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7151148 kB) free!
[1566186751] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566186828] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11604736 kB) free.
[1566186871] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1566187252] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.6% (14570220 kB) free.
[1566187316] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.6% (40713168 kB) free.
[1566187428] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7780684 kB) free!
[1566187471] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566187542] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.80, 0.77, 4.78
[1566187549] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.8% (6319336 kB) free!
[1566187591] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1566187669] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11601572 kB) free.
[1566187794] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1566188191] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1566188269] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6917248 kB) free!
[1566188311] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566188389] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11605972 kB) free.
[1566188431] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1566188989] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7784768 kB) free!
[1566189031] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1566189109] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.9% (6377964 kB) free!
[1566189152] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1566189229] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11602004 kB) free.
[1566189751] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566189829] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6860612 kB) free!
[1566189871] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566189949] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11594604 kB) free.
[1566189991] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1566190183] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4893916 kB) free!
[1566190266] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (240996 kB) free!
[1566190303] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (4832836 kB) free!
[1566190386] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (244888 kB) free!
[1566190423] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.5% (4775896 kB) free!
[1566190507] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (195116 kB) free!
[1566190549] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7715156 kB) free!
[1566190591] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1566190669] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.8% (6303028 kB) free!
[1566190711] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1566190790] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11601832 kB) free.
[1566191311] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1566191389] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7605064 kB) free!
[1566191431] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1566191509] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11600784 kB) free.
[1566191551] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566191964] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4905600 kB) free!
[1566192084] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (4833360 kB) free!
[1566192109] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8830156 kB) free!
[1566192152] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566192204] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (4746608 kB) free!
[1566192229] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7765268 kB) free!
[1566192272] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566192349] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.9% (6338136 kB) free!
[1566192392] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1566192992] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1566193112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566193232] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1566193549] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11591580 kB) free.
[1566194665] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1566194786] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566195204] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1748052 kB) free!
[1566196416] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 63 users currently logged in
[1566196476] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 63 users currently logged in
[1566196536] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 63 users currently logged in
[1566196596] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 63 users currently logged in
[1566196799] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1566196921] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9572032 kB) free!
[1566196973] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2436972 kB) free!
[1566197042] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (9334568 kB) free!
[1566197162] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (9256260 kB) free!
[1566197965] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1566198085] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1566198205] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1566198824] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (201400 kB) free!
[1566198962] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (2897320 kB) free!
[1566199033] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (162952 kB) free!
[1566199154] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194232 kB) free!
[1566199274] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169072 kB) free!
[1566200796] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1566201883] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.7% (14412108 kB) free.
[1566201895] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 370 processes
[1566201942] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.6% (40685400 kB) free.
[1566202015] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1566202025] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.6% (21960956 kB) free.
[1566202074] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.9% (22428784 kB) free.
[1566202129] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.6% (21704288 kB) free.
[1566202181] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.1% (21340232 kB) free.
[1566202194] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1566204294] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (560040 kB) free!
[1566204414] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (503920 kB) free!
[1566204534] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (423804 kB) free!
[1566204696] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 64 users currently logged in
[1566205000] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.87, 5.80, 5.08
[1566205120] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.04, 4.88, 4.84
[1566205664] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.4% (14527048 kB) free.
[1566206796] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1566207011] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 663 processes
[1566207071] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 660 processes
[1566208073] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.9% (2587612 kB) free!
[1566208193] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (301520 kB) free!
[1566208194] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1566208313] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1566208313] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (299856 kB) free!
[1566209754] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.1% (7941736 kB) free!
[1566209796] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 65 users currently logged in
[1566209815] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 371 processes
[1566209874] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.0% (7384332 kB) free!
[1566209935] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 370 processes
[1566209994] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.7% (7220816 kB) free!
[1566210096] SERVICE FLAPPING ALERT: webserver2;Current Users;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1566210096] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1566210485] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175644 kB) free!
[1566210605] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164088 kB) free!
[1566210725] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164084 kB) free!
[1566211405] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1566211794] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1566212106] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.2% (14335000 kB) free.
[1566212146] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.8% (22662600 kB) free.
[1566212320] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.38, 6.09, 5.50
[1566212394] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240728 kB) free!
[1566212440] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.04, 5.49, 5.34
[1566212560] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.26, 5.82, 5.47
[1566212714] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (162808 kB) free!
[1566212755] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178648 kB) free!
[1566212834] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (165316 kB) free!
[1566212875] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184260 kB) free!
[1566212954] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169488 kB) free!
[1566212995] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180384 kB) free!
[1566213160] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 6.03, 4.86, 4.95
[1566213237] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 406.74, 125.33, 45.54
[1566213357] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 564.60, 278.84, 111.47
[1566213477] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 585.76, 381.47, 169.37
[1566213942] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.84, 6.09, 5.39
[1566214062] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.16, 5.71, 5.33
[1566214083] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9271460 kB) free!
[1566214182] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.77, 4.78, 5.05
[1566214203] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9414148 kB) free!
[1566214323] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.7% (9202052 kB) free!
[1566214782] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.33, 1.97, 3.41
[1566214960] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.24, 6.06, 5.25
[1566215080] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.02, 6.29, 5.42
[1566215200] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 8.29, 6.70, 5.66
[1566215605] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 7 users currently logged in
[1566215796] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1566216477] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.17, 3.28, 15.45
[1566216696] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 59 users currently logged in
[1566217000] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1566217000] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.45, 3.60, 4.30
[1566217296] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 62 users currently logged in
[1566217323] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4398740 kB) free!
[1566217356] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 62 users currently logged in
[1566217416] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 62 users currently logged in
[1566217476] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 62 users currently logged in
[1566217585] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 405 processes
[1566217600] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1566217705] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1566217720] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1566217840] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1566217872] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 672 processes
[1566217932] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 673 processes
[1566218277] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.91, 2.97, 4.83
[1566220120] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1566220776] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 55 users currently logged in
[1566220840] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1566220973] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1566223154] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1566223776] SERVICE FLAPPING ALERT: webserver2;Current Users;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1566226782] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.73, 9.42, 7.19
[1566226902] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.89, 8.82, 7.24
[1566227022] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.16, 8.67, 7.38
[1566227200] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1566227305] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1566227425] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1566227478] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 69.6% (17185596 kB) free.
[1566227516] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 47.0% (23152836 kB) free.
[1566227556] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 37.4% (18447220 kB) free.
[1566227597] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.0% (14636920 kB) free.
[1566227830] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4589440 kB) free!
[1566227950] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4509404 kB) free!
[1566228070] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (4426908 kB) free!
[1566228084] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (201608 kB) free!
[1566228204] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182452 kB) free!
[1566228324] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197064 kB) free!
[1566228822] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.90, 4.27, 4.96
[1566230488] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 69.48, 30.35, 11.84
[1566230605] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.2% (2630448 kB) free!
[1566230609] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.67, 20.43, 10.45
[1566230622] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.71, 19.51, 10.91
[1566230725] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.6% (1721208 kB) free!
[1566230729] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.66, 13.78, 9.22
[1566230742] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.16, 13.11, 9.61
[1566230832] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 591 processes
[1566230845] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1043264 kB) free!
[1566230862] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.61, 8.93, 8.50
[1566230892] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 596 processes
[1566231462] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1566231462] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.35, 1.52, 4.64
[1566231671] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 20.7% (5078144 kB) free.
[1566231929] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.70, 0.83, 2.93
[1566231949] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8884948 kB) free!
[1566232070] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.4% (7595724 kB) free!
[1566232190] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11590084 kB) free.
[1566232790] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7787992 kB) free!
[1566232833] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1566232871] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4834232 kB) free!
[1566232910] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11587104 kB) free.
[1566232953] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566232991] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (4818848 kB) free!
[1566233111] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (4781240 kB) free!
[1566234110] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8009124 kB) free!
[1566234153] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566234230] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11490236 kB) free.
[1566234273] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566234911] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1566234911] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 21.4% (5243844 kB) free.
[1566235430] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8285964 kB) free!
[1566235473] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566235550] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.4% (6589092 kB) free!
[1566235593] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1566235670] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11586924 kB) free.
[1566237440] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7057716 kB) free!
[1566237560] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 39.4% (19419600 kB) free.
[1566238511] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4837656 kB) free!
[1566238631] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 20.2% (4949480 kB) free.
[1566238677] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 17.24, 24.65, 12.73
[1566238797] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.50, 17.35, 11.50
[1566238917] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.76, 12.44, 10.41
[1566239193] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1566239231] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4904528 kB) free!
[1566239270] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7738016 kB) free!
[1566239313] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566239351] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4896612 kB) free!
[1566239390] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11588404 kB) free.