[1588185205] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (16991112 kB) free!
[1588185325] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37126184 kB) free.
[1588185925] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15996972 kB) free!
[1588186045] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36384684 kB) free.
[1588186946] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (1709204 kB) free!
[1588187006] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183232 kB) free!
[1588187066] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (184192 kB) free!
[1588187126] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (187248 kB) free!
[1588187245] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16732432 kB) free!
[1588187366] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36395824 kB) free.
[1588187426] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 70.5% (17272216 kB) free.
[1588187656] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1531500 kB) free!
[1588188565] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17440780 kB) free!
[1588188685] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37026464 kB) free.
[1588188856] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.2% (2637552 kB) free!
[1588189285] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16883584 kB) free!
[1588189405] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33908996 kB) free.
[1588190294] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (232392 kB) free!
[1588190414] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (224316 kB) free!
[1588190454] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (404332 kB) free!
[1588190534] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (212884 kB) free!
[1588190574] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1228504 kB) free!
[1588190605] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14226096 kB) free!
[1588190694] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1570220 kB) free!
[1588190704] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 222221 MB (20% inode=99%):
[1588190725] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (33004280 kB) free.
[1588190824] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 222221 MB (20% inode=99%):
[1588190944] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 222221 MB (20% inode=99%):
[1588192334] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 11.0% (5424268 kB) free!
[1588193170] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.9% (3652452 kB) free!
[1588193770] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183176 kB) free!
[1588193889] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 77.9% (38382536 kB) free.
[1588194305] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.2% (43945752 kB) free.
[1588194325] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8576432 kB) free!
[1588194446] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33239460 kB) free.
[1588195876] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(300) < 182.18.144.246 [30/Apr/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1588195935] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.5% (44598208 kB) free.
[1588195996] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588196003] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(418) < 2020/04/30 03:01:32 [error] 22121#0: *86309554 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1588196049] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2350) < 203.212.70.170 [30/Apr/2020:03:02:32 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1588196123] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588196170] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588197256] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.8% (4358860 kB) free.
[1588198570] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.5% (10161680 kB) free.
[1588199771] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (202920 kB) free!
[1588199891] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197596 kB) free!
[1588200011] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1588200011] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (248056 kB) free!
[1588201211] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 37.7% (9245632 kB) free.
[1588203256] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1588203611] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (256308 kB) free!
[1588203731] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182988 kB) free!
[1588203745] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 78.3% (19196548 kB) free.
[1588203851] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182072 kB) free!
[1588203960] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 60.6% (14851440 kB) free.
[1588204432] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 52.9% (13063000 kB) free.
[1588208032] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.5% (3335820 kB) free!
[1588208152] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.2% (3266872 kB) free!
[1588208272] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.2% (3253720 kB) free!
[1588208694] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 46.2% (60889064 kB) free.
[1588208944] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 277292 MB (25% inode=99%):
[1588209445] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26330484 kB) free!
[1588209565] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38869016 kB) free.
[1588211872] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2252104 kB) free!
[1588213451] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1588216165] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26400472 kB) free!
[1588216285] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.1% (22487480 kB) free!
[1588216405] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35901100 kB) free.
[1588217769] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(680) < 103.255.146.164 [30/Apr/2020:09:01:41 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588217889] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588217989] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.26, 5.46, 5.14
[1588218109] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.75, 6.94, 5.80
[1588218205] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26359244 kB) free!
[1588218229] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.77, 6.09, 5.62
[1588218325] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.8% (22133448 kB) free!
[1588218445] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35920056 kB) free.
[1588220845] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (25647152 kB) free!
[1588220845] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1588220965] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.3% (21428084 kB) free!
[1588220965] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1588221085] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16990284 kB) free!
[1588221685] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32760872 kB) free.
[1588221960] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4780164 kB) free!
[1588222080] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (4682184 kB) free!
[1588222200] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (4589240 kB) free!
[1588223485] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (24650020 kB) free!
[1588223605] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.0% (19738040 kB) free!
[1588223725] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.3% (14941488 kB) free!
[1588224000] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1951744 kB) free!
[1588224325] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1588224325] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32312888 kB) free.
[1588225964] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.1% (3379744 kB) free!
[1588226056] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.3% (2162672 kB) free!
[1588226084] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.7% (3069476 kB) free!
[1588226125] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26386420 kB) free!
[1588226176] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3642440 kB) free.
[1588226204] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.9% (2678852 kB) free!
[1588226245] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.2% (43700768 kB) free.
[1588226804] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (763792 kB) free!
[1588227376] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (3000040 kB) free!
[1588227496] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3537524 kB) free.
[1588228604] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.8% (15900160 kB) free.
[1588228690] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(230) < 103.255.146.164 [30/Apr/2020:12:02:18 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588228696] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (2888408 kB) free!
[1588228810] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588228816] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3649108 kB) free.
[1588230565] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1588230616] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.4% (2829016 kB) free!
[1588230685] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1588230736] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (3056812 kB) free!
[1588230805] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1588230857] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (3024312 kB) free!
[1588231457] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3423040 kB) free.
[1588231645] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17664436 kB) free!
[1588231765] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34713844 kB) free.
[1588232057] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3237720 kB) free!
[1588232177] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.3% (2492480 kB) free!
[1588232297] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3170956 kB) free!
[1588232793] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 340 processes
[1588232913] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 340 processes
[1588232965] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5507812 kB) free!
[1588233033] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1588233085] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33643504 kB) free.
[1588234285] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6656768 kB) free!
[1588234406] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34748016 kB) free.
[1588235005] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5510920 kB) free!
[1588235125] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30877908 kB) free.
[1588236325] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1588237505] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.3% (9011464 kB) free!
[1588237525] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17320540 kB) free!
[1588237604] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.6% (3723840 kB) free!
[1588237625] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.9% (8812020 kB) free!
[1588237645] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.1% (42296384 kB) free.
[1588237724] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.6% (3492064 kB) free!
[1588237745] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.5% (8616084 kB) free!
[1588237844] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.6% (3270636 kB) free!
[1588238245] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16330852 kB) free!
[1588238365] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39385352 kB) free.
[1588238410] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(305) < 103.255.146.164 [30/Apr/2020:14:47:50 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588238445] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1588238445] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2042924 kB) free!
[1588238530] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588238898] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1588238898] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1231208 kB) free!
[1588239045] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.4% (15817948 kB) free.
[1588239130] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(135) < 185.255.8.58 [30/Apr/2020:15:00:04 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1588239145] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4735236 kB) free!
[1588239251] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588239265] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4505644 kB) free!
[1588239385] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184084 kB) free!
[1588239403] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1588239498] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.9% (2742632 kB) free!
[1588239523] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 350 processes
[1588239565] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15637272 kB) free!
[1588239633] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 353 processes
[1588239643] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1588239685] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39208744 kB) free.
[1588240233] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1588240243] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1588240285] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23987144 kB) free!
[1588240340] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1588240405] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37037176 kB) free.
[1588240430] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.58, 3.18, 4.86
[1588240450] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(128) < 103.255.146.164 [30/Apr/2020:15:21:51 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588240463] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1588240570] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588240745] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4824784 kB) free!
[1588240900] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 83.8% (20540616 kB) free.
[1588241030] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.27, 5.72, 5.19
[1588241150] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.95, 5.72, 5.24
[1588241170] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(401) < 103.255.146.164 [30/Apr/2020:15:34:52 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588241270] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.31, 6.49, 5.57
[1588241290] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588241298] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.0% (3415224 kB) free.
[1588241605] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12261004 kB) free!
[1588241725] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36848592 kB) free.
[1588241898] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (3088564 kB) free!
[1588242018] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.5% (2841096 kB) free!
[1588242138] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.8% (2569152 kB) free!
[1588242325] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10869556 kB) free!
[1588242445] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35158120 kB) free.
[1588242489] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9666516 kB) free!
[1588242609] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (9606148 kB) free!
[1588242729] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (9447664 kB) free!
[1588242738] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1314556 kB) free!
[1588242805] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1588242925] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1588243338] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.4% (2506764 kB) free!
[1588243645] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11598432 kB) free!
[1588243765] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34867048 kB) free.
[1588243766] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(4) < Channel error on connection <0.13767.5343> (10.147.212.181:60996 -> 10.147.212.160:5672, vhost: '/', user: 'guest'), channel 1:
[1588243886] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1588244365] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10702988 kB) free!
[1588244485] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34016276 kB) free.
[1588244890] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(521) < 103.255.146.164 [30/Apr/2020:16:31:25 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588245010] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588245085] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9812440 kB) free!
[1588245205] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33367208 kB) free.
[1588245805] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7605580 kB) free!
[1588245925] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32171664 kB) free.
[1588246525] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4513676 kB) free!
[1588246645] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31130052 kB) free.
[1588246826] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(Return code of 255 for service 'rq2-rabbitmqerror' on host 'rq-2' was out of bounds)
[1588246826] SERVICE ALERT: rq-2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'rq-2' was out of bounds)
[1588246829] HOST ALERT: rq-2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.204)
[1588246877] SERVICE ALERT: rq-2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'rq-2' was out of bounds)
[1588246949] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;HARD;1;(Return code of 255 for service 'rq2-rabbitmqerror' on host 'rq-2' was out of bounds)
[1588246949] SERVICE ALERT: rq-2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'rq-2' was out of bounds)
[1588246976] SERVICE ALERT: rq-2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'rq-2' was out of bounds)
[1588246997] SERVICE ALERT: rq-2;Total Processes;OK;HARD;1;PROCS OK: 349 processes
[1588247001] HOST ALERT: rq-2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1588247078] SERVICE ALERT: sigqscluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1588247096] SERVICE ALERT: rq-2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1588247123] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(13) < 2020/04/30 17:10:30 [error] 22131#0: *86398292 upstream timed out (110: Connection timed out) while connecting to upstream, client: 183.82.105.122, server: rq-2-1gbps.server, request: "GET /rabbitmq/api/queues/?page=1&page_size=100&name=BL5
[1588247208] SERVICE ALERT: sigqscluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1588247222] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 80.24, 63.88, 33.73
[1588247232] SERVICE ALERT: sigqscluster1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1588247243] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588247246] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3655016 kB) free!
[1588247272] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 717.25, 552.09, 289.55
[1588247338] SERVICE ALERT: sigqscluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1588247343] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 78.75, 68.68, 39.15
[1588247362] SERVICE ALERT: sigqscluster1;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1588247366] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28188980 kB) free.
[1588247410] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(815) < 103.255.146.164 [30/Apr/2020:17:19:31 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588247445] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.7% (3739464 kB) free!
[1588247464] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 81.67, 73.21, 44.39
[1588247492] SERVICE ALERT: sigqscluster1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1588247530] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588247546] SERVICE ALERT: rq-2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1588247547] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1588247565] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.8% (3539828 kB) free!
[1588247685] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.9% (3351452 kB) free!
[1588247966] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (2576208 kB) free!
[1588248086] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29427260 kB) free.
[1588248686] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (826700 kB) free!
[1588248686] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.31347.5305> (10.147.212.53:49116 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1588248693] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 1846 processes
[1588248806] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (380316 kB) free!
[1588248806] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1588248886] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (334692 kB) free!
[1588248926] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31267100 kB) free.
[1588249140] SERVICE ALERT: sigqscluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1588249330] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(734) < 103.255.146.164 [30/Apr/2020:17:46:55 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588249407] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.12840.5316> (10.147.212.181:33302 -> 10.147.212.160:5672, vhost: '/', user: 'guest'), channel 1:
[1588249450] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588249526] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (1770516 kB) free!
[1588249527] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1588249646] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38119684 kB) free.
[1588249929] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4129848 kB) free!
[1588249947] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.11720.0> (10.147.212.158:54952 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1588250067] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1588250086] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.3% (15780328 kB) free.
[1588250246] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4143196 kB) free!
[1588250366] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (396276 kB) free!
[1588250486] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39415428 kB) free.
[1588251003] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 725 processes
[1588251064] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.76, 2.16, 19.64
[1588251086] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (7040968 kB) free!
[1588251100] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.3% (4235620 kB) free!
[1588251124] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 796 processes
[1588251138] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (280512 kB) free!
[1588251206] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2684520 kB) free!
[1588251220] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.1% (3956600 kB) free!
[1588251222] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 20.0% (4904888 kB) free!
[1588251244] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 845 processes
[1588251267] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.4367.7> (10.147.212.158:60910 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1588251326] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15509628 kB) free!
[1588251340] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 15.0% (3666992 kB) free!
[1588251342] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 20.0% (4897564 kB) free!
[1588251388] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1588251462] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.9% (4888232 kB) free!
[1588251738] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.2% (1817576 kB) free!
[1588251781] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 26.7% (13144444 kB) free.
[1588251927] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (8403036 kB) free!
[1588251940] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2167208 kB) free!
[1588252382] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (1169412 kB) free!
[1588252502] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (1143320 kB) free!
[1588252527] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21487472 kB) free!
[1588252623] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (579656 kB) free!
[1588252864] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 1.16, 1.01, 3.74
[1588252938] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.5% (3492900 kB) free.
[1588253127] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1588253127] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39192140 kB) free.
[1588253538] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (953324 kB) free!
[1588253658] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (1256576 kB) free!
[1588253727] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8184500 kB) free!
[1588253778] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1604612 kB) free!
[1588253847] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3413176 kB) free!
[1588253967] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387836 kB) free!
[1588254033] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1588254978] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.2% (2638664 kB) free!
[1588255167] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23250372 kB) free!
[1588255233] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 354 processes
[1588255444] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 692 processes
[1588255495] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 11.6% (15277348 kB) free!
[1588255578] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1324252 kB) free!
[1588255615] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.6% (27158080 kB) free.
[1588255767] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33138724 kB) free.
[1588255833] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1588256086] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.1% (3383404 kB) free!
[1588256125] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1588256178] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.5% (3488436 kB) free.
[1588256206] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.8% (3314920 kB) free!
[1588256245] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1588256326] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.5% (3246168 kB) free!
[1588256367] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8234004 kB) free!
[1588256487] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35277400 kB) free.
[1588257033] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1588257033] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 356 processes
[1588257087] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (12048688 kB) free!
[1588257207] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (8139304 kB) free!
[1588257327] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34494416 kB) free.
[1588257378] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (3026544 kB) free!
[1588257387] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.7001.9> (10.147.212.158:38176 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1588257472] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.71, 10.38, 16.03
[1588257498] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3355268 kB) free.
[1588257507] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1588257927] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (11985724 kB) free!
[1588258048] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34901156 kB) free.
[1588258098] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.4% (2672432 kB) free!
[1588258126] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2140652 kB) free!
[1588258218] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3293520 kB) free.
[1588258233] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 355 processes
[1588258615] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.9% (21006988 kB) free!
[1588258647] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7104376 kB) free!
[1588258735] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.0% (27648444 kB) free.
[1588258767] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1172864 kB) free!
[1588258818] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.3% (1840356 kB) free!
[1588258833] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 356 processes
[1588258887] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35535044 kB) free.
[1588258938] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.1% (2624292 kB) free!
[1588259058] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (3113004 kB) free!
[1588259335] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.7% (18107724 kB) free!
[1588259455] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 27.3% (36022768 kB) free.
[1588259487] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (4993804 kB) free!
[1588259607] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (382392 kB) free!
[1588259727] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32054000 kB) free.
[1588260033] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1588260327] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (1818348 kB) free!
[1588260447] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (16851048 kB) free!
[1588260567] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31443140 kB) free.
[1588260858] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.4% (4126284 kB) free.
[1588261126] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.2% (15982900 kB) free.
[1588261167] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7451560 kB) free!
[1588261287] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33281196 kB) free.
[1588261887] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8426540 kB) free!
[1588262007] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4233420 kB) free!
[1588262127] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31758860 kB) free.
[1588262658] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3250704 kB) free!
[1588262728] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6799780 kB) free!
[1588262779] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (3190640 kB) free!
[1588262848] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2654336 kB) free!
[1588262899] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1028864 kB) free!
[1588262968] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30244280 kB) free.
[1588263568] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6546404 kB) free!
[1588263688] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31843696 kB) free.
[1588264288] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11693068 kB) free!
[1588264408] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7809308 kB) free!
[1588264528] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31618628 kB) free.
[1588265128] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10343908 kB) free!
[1588265248] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6368156 kB) free!
[1588265368] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30569824 kB) free.
[1588265968] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9284368 kB) free!
[1588266088] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (5584868 kB) free!
[1588266208] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29571600 kB) free.
[1588266499] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (2814912 kB) free!
[1588266808] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8830032 kB) free!
[1588266929] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.9% (22306468 kB) free!
[1588267049] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28845372 kB) free.
[1588267649] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8073916 kB) free!
[1588267672] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.30, 2.65, 4.04
[1588267769] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31068172 kB) free.
[1588268369] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11063196 kB) free!
[1588268489] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7612592 kB) free!
[1588268609] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30897768 kB) free.
[1588269209] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10510460 kB) free!
[1588269329] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (6922008 kB) free!
[1588269445] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1588269449] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29865584 kB) free.
[1588269499] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (664696 kB) free!
[1588269565] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1588270049] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9478928 kB) free!
[1588270099] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.0% (2602792 kB) free!
[1588270127] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1588270169] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.5% (23011864 kB) free!
[1588270289] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28239076 kB) free.
[1588270833] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1588270889] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8156672 kB) free!
[1588271009] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1588271009] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30695948 kB) free.