[1566412401] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1566412442] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.27, 3.99, 4.64
[1566412521] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 123335 MB (54% inode=99%):
[1566412526] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 79.5% (39179992 kB) free.
[1566412800] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.28, 24.98, 13.72
[1566412920] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.99, 17.61, 12.39
[1566413040] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.54, 12.54, 11.16
[1566414240] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1566414240] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.94, 2.80, 4.90
[1566415412] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.80, 7.27, 7.75
[1566415532] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.01, 5.19, 6.92
[1566415652] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.50, 3.61, 6.14
[1566415876] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 615 processes
[1566415914] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 41.5% (10240308 kB) free.
[1566415936] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 615 processes
[1566416252] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.76, 1.00, 3.50
[1566416640] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 232.11, 178.75, 88.81
[1566416744] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6717308 kB) free!
[1566416760] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 232.47, 196.51, 106.22
[1566416864] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (398960 kB) free!
[1566416880] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 74.38, 160.50, 104.88
[1566416984] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (632436 kB) free!
[1566417725] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.1% (2465276 kB) free!
[1566418325] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2423428 kB) free!
[1566418652] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.54, 7.14, 6.75
[1566418680] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.97, 3.59, 19.94
[1566418773] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.84, 5.08, 6.04
[1566418893] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.40, 3.53, 5.36
[1566419493] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.50, 1.00, 3.10
[1566420714] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195344 kB) free!
[1566420834] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196076 kB) free!
[1566420954] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198168 kB) free!
[1566421080] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.83, 2.89, 4.14
[1566421726] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.9% (46288520 kB) free.
[1566421893] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.65, 7.62, 6.79
[1566422013] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.41, 5.22, 6.01
[1566422133] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.93, 3.76, 5.38
[1566422733] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.76, 1.20, 3.20
[1566423725] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 46.7% (11440848 kB) free.
[1566423865] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 43.7% (10714364 kB) free.
[1566424794] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1566425133] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.44, 6.17, 5.57
[1566425253] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.38, 4.24, 4.93
[1566427680] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.51, 9.38, 9.38
[1566427800] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.61, 7.18, 8.57
[1566427920] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.02, 5.82, 7.90
[1566428253] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.23, 10.54, 6.65
[1566428374] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.69, 7.21, 5.91
[1566428494] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.38, 4.97, 5.25
[1566429094] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.91, 1.24, 3.04
[1566429120] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.60, 2.72, 4.15
[1566431494] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.43, 15.49, 8.53
[1566431614] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.97, 10.45, 7.52
[1566431734] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.84, 7.24, 6.70
[1566432120] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 34.01, 24.16, 11.26
[1566432240] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.63, 16.96, 10.19
[1566432334] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.73, 1.56, 3.84
[1566432360] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.25, 12.31, 9.30
[1566433560] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.41, 2.74, 4.41
[1566435584] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 43.1% (56756964 kB) free.
[1566435954] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 47.3% (11663984 kB) free.
[1566436325] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.7% (4097584 kB) free!
[1566436336] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 603 processes
[1566436396] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 601 processes
[1566436445] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178396 kB) free!
[1566436465] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195844 kB) free!
[1566436554] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185344 kB) free!
[1566436566] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176264 kB) free!
[1566436585] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189720 kB) free!
[1566436674] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181060 kB) free!
[1566436696] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 604 processes
[1566436705] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195636 kB) free!
[1566436756] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 609 processes
[1566436794] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (889312 kB) free!
[1566436817] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 603 processes
[1566437160] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 11.93, 29.83, 16.89
[1566437282] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.52, 21.07, 15.25
[1566437402] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.60, 15.26, 13.83
[1566438335] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.85, 5.53, 5.99
[1566438455] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.53, 3.88, 5.33
[1566438575] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.71, 2.89, 4.80
[1566438602] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 212.51, 179.18, 99.28
[1566441602] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.60, 2.78, 13.27
[1566443975] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1566444002] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.53, 2.76, 4.43
[1566445586] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1566445706] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566447002] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.81, 13.08, 9.67
[1566447122] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.68, 9.63, 8.82
[1566447242] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.93, 7.42, 8.11
[1566447575] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 93.41, 27.60, 10.12
[1566447695] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 13.18, 18.71, 8.99
[1566447816] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.09, 12.66, 7.95
[1566448416] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.32, 2.63, 4.66
[1566448442] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.63, 2.82, 4.23
[1566449766] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.4% (22415172 kB) free.
[1566449905] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.4% (21180204 kB) free.
[1566449961] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.1% (40443940 kB) free.
[1566449994] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 75.6% (18667668 kB) free.
[1566450161] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 52.8% (8584868 kB) free.
[1566450174] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.3% (21635240 kB) free.
[1566450505] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (320808 kB) free!
[1566450594] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (1349464 kB) free!
[1566450625] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (282700 kB) free!
[1566450714] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (1648304 kB) free!
[1566450745] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (221928 kB) free!
[1566450816] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 94.17, 28.23, 10.80
[1566450834] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.1% (3234276 kB) free!
[1566450936] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 13.50, 19.13, 9.58
[1566450967] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (705476 kB) free!
[1566451056] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.46, 13.04, 8.51
[1566451087] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (659752 kB) free!
[1566451124] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168880 kB) free!
[1566451207] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (577648 kB) free!
[1566451246] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.7% (21257540 kB) free.
[1566451253] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1566451262] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.1% (22337408 kB) free.
[1566451341] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1566451341] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.8% (20187328 kB) free.
[1566451373] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1566451603] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1566451656] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1566451656] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.05, 2.63, 4.95
[1566451723] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1566451794] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1566451843] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1566451941] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.5% (3813276 kB) free!
[1566451961] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (3102728 kB) free!
[1566452061] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.1% (3719884 kB) free!
[1566452081] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (3103600 kB) free!
[1566452181] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.1% (3237820 kB) free!
[1566452201] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 19.0% (3092128 kB) free!
[1566452394] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1566452573] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1566452693] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1566452801] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (149556 kB) free!
[1566452813] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1566453381] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168844 kB) free!
[1566453535] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 357 processes
[1566453562] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (865620 kB) free!
[1566453655] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 358 processes
[1566453682] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236072 kB) free!
[1566453775] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1566453802] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (222892 kB) free!
[1566454056] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 55.60, 13.05, 5.38
[1566454176] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 21.50, 20.97, 9.80
[1566454296] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.52, 14.22, 8.68
[1566454536] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 62 users currently logged in
[1566454596] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 62 users currently logged in
[1566454656] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 62 users currently logged in
[1566454716] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 62 users currently logged in
[1566454896] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.50, 2.42, 4.84
[1566455451] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (216772 kB) free!
[1566455471] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192388 kB) free!
[1566455571] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194224 kB) free!
[1566455591] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191872 kB) free!
[1566455691] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1566455691] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (207668 kB) free!
[1566455711] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1566455711] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178808 kB) free!
[1566456842] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.79, 10.31, 10.71
[1566456962] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.20, 7.62, 9.67
[1566457082] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.37, 5.89, 8.79
[1566457296] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 101.01, 24.39, 9.42
[1566457416] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 22.93, 24.53, 11.70
[1566457536] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.80, 16.68, 10.38
[1566458262] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 210.44, 81.40, 30.61
[1566458282] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.91, 3.17, 4.69
[1566458382] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 238.18, 134.73, 56.31
[1566458502] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 242.20, 170.50, 78.93
[1566458736] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.85, 1.20, 3.56
[1566458781] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.4% (2555612 kB) free!
[1566459381] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2133544 kB) free!
[1566459393] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1566459514] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566459526] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.89, 5.38, 5.03
[1566459646] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.88, 4.91, 4.91
[1566459710] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9241704 kB) free!
[1566459830] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8030436 kB) free!
[1566459951] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.0% (6413400 kB) free!
[1566460113] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566460233] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1566460247] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.45, 6.00, 5.59
[1566460302] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.60, 1.86, 17.28
[1566460353] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566460367] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.92, 5.20, 5.34
[1566460487] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.13, 4.76, 5.16
[1566460551] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11454472 kB) free.
[1566461087] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.41, 3.81, 4.50
[1566461136] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.80, 6.02, 7.32
[1566461256] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.59, 4.22, 6.50
[1566461282] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 16.35, 24.01, 12.76
[1566461316] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 47 users currently logged in
[1566461376] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.65, 3.07, 5.81
[1566461402] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.68, 17.01, 11.56
[1566461443] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1566461522] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.85, 12.60, 10.60
[1566461976] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.23, 1.49, 3.62
[1566462102] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.68, 0.76, 3.19
[1566462311] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1566462506] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 405 processes
[1566462626] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1566462753] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1566462891] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1566462951] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9302880 kB) free!
[1566463071] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8075176 kB) free!
[1566463191] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.1% (6446416 kB) free!
[1566463211] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 327 processes
[1566463322] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.16, 3.09, 4.13
[1566463331] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 320 processes
[1566463353] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566463474] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566463594] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1566463791] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1566463791] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11419288 kB) free.
[1566464376] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.18, 5.42, 6.48
[1566464391] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7453596 kB) free!
[1566464497] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.66, 3.82, 5.77
[1566464511] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11417292 kB) free.
[1566464550] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.2% (22496112 kB) free.
[1566464591] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.8% (22512652 kB) free.
[1566464617] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.82, 2.86, 5.18
[1566464629] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 53.7% (26473912 kB) free.
[1566464691] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.1% (21605796 kB) free.
[1566464733] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.9% (14609232 kB) free.
[1566465111] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7773348 kB) free!
[1566465217] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.12, 1.08, 3.07
[1566465232] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11417736 kB) free.
[1566465394] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1566465394] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1566466594] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1566466714] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1566466835] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1566467141] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (1263408 kB) free!
[1566467261] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (255584 kB) free!
[1566467381] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (179456 kB) free!
[1566467434] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1566467554] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566467617] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.26, 7.38, 7.82
[1566467674] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566467737] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.77, 5.20, 6.97
[1566467857] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.79, 3.73, 6.22
[1566467920] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1566468231] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7457416 kB) free!
[1566468274] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566468351] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11406696 kB) free.
[1566468394] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1566468457] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.04, 1.27, 3.66
[1566468775] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 358 processes
[1566468895] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1566468951] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7817888 kB) free!
[1566468994] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1566469071] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11310692 kB) free.
[1566469114] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1566469671] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8111272 kB) free!
[1566469715] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566469791] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6417252 kB) free!
[1566469835] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1566469912] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11416000 kB) free.
[1566470555] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1566470634] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.6% (8158120 kB) free!
[1566470695] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 357 processes
[1566470754] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.9% (8351748 kB) free!
[1566470784] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 639 processes
[1566470815] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1566470857] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.48, 7.23, 7.03
[1566470874] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.7% (8254416 kB) free!
[1566470904] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 637 processes
[1566470977] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.83, 5.10, 6.27
[1566471024] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 636 processes
[1566471097] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.88, 3.75, 5.64
[1566471122] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.39, 10.77, 9.06
[1566471242] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.13, 8.23, 8.34
[1566471362] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.62, 6.36, 7.64
[1566471624] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 393 processes
[1566471697] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.40, 1.33, 3.38
[1566472074] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3766324 kB) free!
[1566472355] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185604 kB) free!
[1566472475] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170232 kB) free!
[1566472498] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182760 kB) free!
[1566472562] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.34, 3.71, 4.52
[1566472596] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191984 kB) free!
[1566472618] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186644 kB) free!
[1566472726] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9466552 kB) free!
[1566472738] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171004 kB) free!
[1566472846] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.9% (9337580 kB) free!
[1566472966] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.7% (9221452 kB) free!
[1566472991] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173844 kB) free!
[1566473111] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185548 kB) free!
[1566473231] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182716 kB) free!
[1566473435] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1566473556] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1566474097] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.84, 7.30, 6.98
[1566474217] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.94, 5.17, 6.23
[1566474337] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.77, 3.65, 5.54
[1566474937] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.91, 1.15, 3.29
[1566475015] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 357 processes
[1566475135] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 360 processes
[1566475255] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1566475855] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 355 processes
[1566475911] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1566475947] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.6% (14557656 kB) free.
[1566475975] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1566475993] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.0% (21076968 kB) free.
[1566476033] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.9% (40829856 kB) free.
[1566476099] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.8% (22018828 kB) free.
[1566476127] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.1% (21992484 kB) free.
[1566476162] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.82, 7.30, 8.07
[1566476164] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;1;CRITICAL - 0.8% (199888 kB) free!
[1566476282] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.39, 6.02, 7.51
[1566476284] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;2;CRITICAL - 0.8% (201156 kB) free!
[1566476402] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.26, 5.12, 7.00
[1566476404] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 0.8% (198344 kB) free!
[1566476555] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1566476813] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1566477166] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4713484 kB) free!
[1566477337] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.42, 10.60, 7.86
[1566477457] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.18, 7.42, 7.03
[1566477577] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.03, 5.34, 6.31
[1566477603] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.74, 2.75, 3.97
[1566477626] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1566477746] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1566478177] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.62, 1.42, 3.73
[1566478347] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (231828 kB) free!
[1566478356] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1566478467] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (227604 kB) free!
[1566478476] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566478587] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174432 kB) free!
[1566478596] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1566479796] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1566479813] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1566479933] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1566480053] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1566480111] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7919348 kB) free!
[1566480146] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1566480231] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.2% (6007768 kB) free!
[1566480266] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566480351] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11403372 kB) free.
[1566480520] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1566480603] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.01, 18.19, 10.55
[1566480702] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 61.14, 136.09, 84.28
[1566480723] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.74, 13.10, 9.61
[1566480822] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 8.81, 91.27, 74.13
[1566480843] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.89, 9.64, 8.76
[1566480942] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.20, 61.44, 65.28
[1566480951] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.6% (6200136 kB) free!
[1566481071] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11413688 kB) free.
[1566481253] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1566481596] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1566481671] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6774888 kB) free!
[1566481716] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566481791] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11412876 kB) free.
[1566481837] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1566482043] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 249.14, 161.57, 73.87
[1566482142] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.02, 2.14, 18.68
[1566482378] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.25, 5.78, 6.33
[1566482392] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7773744 kB) free!
[1566482436] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1566482498] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.93, 4.16, 5.67
[1566482512] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11415600 kB) free.
[1566482556] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1566482618] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.67, 3.06, 5.08
[1566483112] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8126192 kB) free!
[1566483156] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566483218] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.93, 1.02, 2.97
[1566483232] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.3% (6550736 kB) free!
[1566483241] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 13.1% (6479416 kB) free!
[1566483276] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566483352] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11402076 kB) free.
[1566483361] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 11.6% (5711520 kB) free!
[1566483481] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (5083372 kB) free!
[1566483942] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.95, 1.02, 3.54
[1566484081] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3264212 kB) free!
[1566484443] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.43, 3.22, 13.54
[1566485894] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1566485894] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.9% (40843572 kB) free.
[1566485932] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 38.5% (18994832 kB) free.
[1566485972] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.8% (14601024 kB) free.
[1566486818] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 52.34, 24.79, 9.95
[1566486938] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.48, 16.75, 8.81
[1566487058] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.94, 11.48, 7.83
[1566487443] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.47, 2.59, 3.95
[1566487658] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.62, 2.16, 4.44
[1566488066] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1566488186] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 402 processes
[1566488306] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 401 processes
[1566488676] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566488753] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8625152 kB) free!
[1566488796] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1566488873] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7375816 kB) free!
[1566488916] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566488983] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (155768 kB) free!
[1566488993] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11413904 kB) free.
[1566489103] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (179064 kB) free!
[1566489223] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (181676 kB) free!
[1566489506] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 405 processes
[1566489593] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7402044 kB) free!
[1566489713] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11411560 kB) free.
[1566490313] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7703980 kB) free!
[1566490433] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11414240 kB) free.
[1566490443] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.79, 14.10, 9.66
[1566490563] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.88, 10.32, 8.81
[1566490683] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.56, 7.70, 8.03
[1566490716] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566491033] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8023224 kB) free!
[1566491153] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.7% (6270828 kB) free!
[1566491273] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11411264 kB) free.
[1566491859] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.71, 15.18, 8.56
[1566491873] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6848644 kB) free!
[1566491883] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.70, 2.72, 4.08
[1566491916] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566491979] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.98, 10.31, 7.58
[1566491993] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11413396 kB) free.
[1566492036] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1566492099] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.70, 7.12, 6.74
[1566492503] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.7% (8731796 kB) free!
[1566492623] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.7% (8732344 kB) free!
[1566492699] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.46, 1.53, 3.87
[1566492743] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.7% (8731928 kB) free!
[1566493943] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4223652 kB) free!
[1566496899] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.48, 7.58, 6.71
[1566497019] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.53, 5.28, 5.98
[1566497139] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.76, 3.79, 5.34
[1566497739] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.47, 1.03, 3.11