[1588530957] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11471356 kB) free!
[1588531077] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7848964 kB) free!
[1588531110] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.3% (4928736 kB) free.
[1588531197] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33354220 kB) free.
[1588531377] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1588531797] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10855228 kB) free!
[1588531917] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7303108 kB) free!
[1588532037] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32648784 kB) free.
[1588532637] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11187884 kB) free!
[1588532686] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196956 kB) free!
[1588532746] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171348 kB) free!
[1588532757] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7734176 kB) free!
[1588532806] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (190892 kB) free!
[1588532866] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (189484 kB) free!
[1588532878] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33833388 kB) free.
[1588533467] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 79.3% (19445660 kB) free.
[1588533478] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11277204 kB) free!
[1588533510] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (1371608 kB) free!
[1588533599] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7789716 kB) free!
[1588533630] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (1219128 kB) free!
[1588533719] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32866792 kB) free.
[1588533750] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (648256 kB) free!
[1588534319] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10689632 kB) free!
[1588534439] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7210000 kB) free!
[1588534559] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31698780 kB) free.
[1588535159] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10057860 kB) free!
[1588535279] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34008536 kB) free.
[1588535879] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12744840 kB) free!
[1588535999] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9139100 kB) free!
[1588536119] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34003300 kB) free.
[1588536184] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (407848 kB) free!
[1588536257] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (229216 kB) free!
[1588536304] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (1462732 kB) free!
[1588536377] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (220468 kB) free!
[1588536425] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2187852 kB) free!
[1588536497] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (218380 kB) free!
[1588536719] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11257236 kB) free!
[1588536839] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7893388 kB) free!
[1588536959] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31793352 kB) free.
[1588537559] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9834132 kB) free!
[1588537679] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33914660 kB) free.
[1588538280] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11849544 kB) free!
[1588538297] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 12.3% (6073548 kB) free!
[1588538400] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8487500 kB) free!
[1588538520] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33617952 kB) free.
[1588539120] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11104884 kB) free!
[1588539240] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33886608 kB) free.
[1588539840] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13266124 kB) free!
[1588539850] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 90.2% (44466952 kB) free.
[1588539960] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9750932 kB) free!
[1588540080] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33376820 kB) free.
[1588540265] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.6% (44644840 kB) free.
[1588540407] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1588540527] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1588540572] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 31.5% (7730372 kB) free.
[1588540680] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11351252 kB) free!
[1588540800] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33868188 kB) free.
[1588541400] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12569880 kB) free!
[1588541520] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9151692 kB) free!
[1588541550] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.2% (2467088 kB) free!
[1588541555] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1588541640] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31215668 kB) free.
[1588541772] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194116 kB) free!
[1588541848] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(470) < 2020/05/04 03:01:32 [error] 22130#0: *86733173 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1588541892] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169776 kB) free!
[1588541898] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.2% (44439964 kB) free.
[1588541960] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(372) < 182.18.144.246 [04/May/2020:03:00:37 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1588541968] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588542012] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (257388 kB) free!
[1588542080] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588542150] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1227264 kB) free!
[1588542240] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9800924 kB) free!
[1588542360] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1588542360] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32740440 kB) free.
[1588542750] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3508132 kB) free.
[1588542960] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11272764 kB) free!
[1588543081] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33708684 kB) free.
[1588543680] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6344264 kB) free!
[1588543800] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (1490468 kB) free!
[1588543920] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30045572 kB) free.
[1588544520] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4443856 kB) free!
[1588544640] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33247104 kB) free.
[1588545013] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 38.4% (9407392 kB) free.
[1588545240] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6614932 kB) free!
[1588545360] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33262656 kB) free.
[1588545613] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (219232 kB) free!
[1588545733] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (207168 kB) free!
[1588545853] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1588545853] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (209736 kB) free!
[1588545960] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8773672 kB) free!
[1588546080] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4055376 kB) free!
[1588546200] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30313516 kB) free.
[1588546800] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6015064 kB) free!
[1588546920] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31280884 kB) free.
[1588547053] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 38.6% (9459008 kB) free.
[1588547520] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9363636 kB) free!
[1588547640] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34251812 kB) free.
[1588548240] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14257476 kB) free!
[1588548361] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33789948 kB) free.
[1588548853] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (146940 kB) free!
[1588548960] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14064212 kB) free!
[1588548973] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179636 kB) free!
[1588549080] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33050920 kB) free.
[1588549094] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (290036 kB) free!
[1588549706] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 76.6% (18784408 kB) free.
[1588549792] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 56.1% (13841228 kB) free.
[1588550162] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 61.9% (15176504 kB) free.
[1588550280] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19796836 kB) free!
[1588550400] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37499508 kB) free.
[1588551000] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18174912 kB) free!
[1588551120] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35700052 kB) free.
[1588552320] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18121492 kB) free!
[1588552440] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37230872 kB) free.
[1588553040] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16853820 kB) free!
[1588553160] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35392064 kB) free.
[1588553992] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.2% (3515192 kB) free!
[1588554112] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (1567884 kB) free!
[1588554150] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1588554232] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1556352 kB) free!
[1588554360] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19226268 kB) free!
[1588554481] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36518512 kB) free.
[1588555025] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 60.9% (80313540 kB) free.
[1588555680] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18899248 kB) free!
[1588555800] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36408880 kB) free.
[1588556400] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16519244 kB) free!
[1588556520] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34944588 kB) free.
[1588557720] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (16946580 kB) free!
[1588557840] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34535148 kB) free.
[1588559040] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16168848 kB) free!
[1588559160] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34241136 kB) free.
[1588559294] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1588560360] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13953140 kB) free!
[1588560480] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32843824 kB) free.
[1588561681] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12815860 kB) free!
[1588561801] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32842604 kB) free.
[1588562037] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.64, 7.62, 6.07
[1588562157] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.32, 6.67, 5.92
[1588562277] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.47, 6.50, 5.93
[1588562401] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10575676 kB) free!
[1588562521] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (31006628 kB) free.
[1588563362] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.1% (3954252 kB) free!
[1588563482] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.2% (3233760 kB) free!
[1588563602] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.1% (2966492 kB) free!
[1588563721] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7706440 kB) free!
[1588563755] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1588563841] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30529960 kB) free.
[1588563875] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1588563995] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1588564202] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1884880 kB) free!
[1588564441] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7763516 kB) free!
[1588564527] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1588564561] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29414100 kB) free.
[1588564647] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1588565161] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (7008100 kB) free!
[1588565281] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28435988 kB) free.
[1588565881] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6477584 kB) free!
[1588566001] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27703404 kB) free.
[1588566601] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4194956 kB) free!
[1588566721] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.5% (24382808 kB) free!
[1588566841] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19658512 kB) free!
[1588567442] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39703448 kB) free.
[1588567647] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1588567768] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 330 processes
[1588567888] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1588568642] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19808032 kB) free!
[1588568762] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.7% (15396548 kB) free!
[1588568882] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38649004 kB) free.
[1588569164] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1588569284] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1588569395] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 347 processes
[1588569404] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1588569482] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14368352 kB) free!
[1588569506] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178900 kB) free!
[1588569602] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39024884 kB) free.
[1588569626] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (203536 kB) free!
[1588569746] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189072 kB) free!
[1588569995] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1588570004] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1588570202] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15298360 kB) free!
[1588570322] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38808192 kB) free.
[1588570350] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2993188 kB) free!
[1588570384] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.1% (4067100 kB) free!
[1588570470] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3693608 kB) free.
[1588570504] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.7% (3742988 kB) free!
[1588570624] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.1% (3398192 kB) free!
[1588570660] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 64.2% (15741820 kB) free.
[1588570702] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(519) < 103.255.146.164 [04/May/2020:11:01:59 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588570822] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588570922] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15133344 kB) free!
[1588571042] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (10811912 kB) free!
[1588571162] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35387756 kB) free.
[1588571196] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1588571196] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 349 processes
[1588571224] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1438100 kB) free!
[1588571762] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (12039880 kB) free!
[1588571882] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36536184 kB) free.
[1588572482] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11516268 kB) free!
[1588572603] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (6980240 kB) free!
[1588572723] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27465908 kB) free.
[1588572870] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.4% (2503788 kB) free!
[1588572990] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (3100740 kB) free!
[1588573110] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3634720 kB) free.
[1588573323] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (395216 kB) free!
[1588573443] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36340636 kB) free.
[1588573663] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.3% (4738960 kB) free!
[1588573710] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3197736 kB) free!
[1588573783] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.3% (4735964 kB) free!
[1588573830] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (3179656 kB) free!
[1588573903] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.3% (4731236 kB) free!
[1588573950] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3354580 kB) free.
[1588574043] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6884548 kB) free!
[1588574163] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2436092 kB) free!
[1588574224] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.8% (16128744 kB) free.
[1588574283] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36480008 kB) free.
[1588574550] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (3145764 kB) free!
[1588574670] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.6% (2532616 kB) free!
[1588574790] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (3061984 kB) free!
[1588574796] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1588574883] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7481456 kB) free!
[1588575003] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2776144 kB) free!
[1588575123] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35882916 kB) free.
[1588575723] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (7011044 kB) free!
[1588575843] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2619420 kB) free!
[1588575963] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34619668 kB) free.
[1588575990] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3328180 kB) free.
[1588576563] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7731024 kB) free!
[1588576590] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3047328 kB) free!
[1588576683] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3232992 kB) free!
[1588576710] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (2966620 kB) free!
[1588576803] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31861404 kB) free.
[1588576830] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.0% (2933712 kB) free!
[1588576946] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 37.4% (9175908 kB) free.
[1588577404] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5718512 kB) free!
[1588577524] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32727452 kB) free.
[1588577860] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.7% (4344524 kB) free!
[1588577980] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.7% (4099088 kB) free!
[1588578100] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 15.8% (3860452 kB) free!
[1588578124] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5756920 kB) free!
[1588578244] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1174000 kB) free!
[1588578364] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386584 kB) free!
[1588578396] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1588578665] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.2% (8500240 kB) free!
[1588578700] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2040012 kB) free!
[1588578785] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.4% (8086408 kB) free!
[1588578905] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.6% (7683728 kB) free!
[1588579564] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1588579564] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19255580 kB) free!
[1588579822] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(237) < 103.255.146.164 [04/May/2020:13:31:40 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588579830] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1588579830] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3367992 kB) free.
[1588579942] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588580105] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4393580 kB) free!
[1588580164] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39022952 kB) free.
[1588580224] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (2120524 kB) free!
[1588580344] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (1668516 kB) free!
[1588580430] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.9% (1937108 kB) free!
[1588580464] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (1255132 kB) free!
[1588580546] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (4602868 kB) free!
[1588580550] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.5% (1713032 kB) free!
[1588580666] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4506544 kB) free!
[1588580670] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 20.0% (3246840 kB) free!
[1588580764] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10499132 kB) free!
[1588580786] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (4426352 kB) free!
[1588580884] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (5961908 kB) free!
[1588581004] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32710216 kB) free.
[1588581604] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5132376 kB) free!
[1588581724] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (718000 kB) free!
[1588581844] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33913584 kB) free.
[1588582342] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(421) < 103.255.146.164 [04/May/2020:14:16:34 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588582444] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6768044 kB) free!
[1588582462] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588582565] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2308536 kB) free!
[1588582685] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34065464 kB) free.
[1588583187] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2278696 kB) free!
[1588583285] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6302640 kB) free!
[1588583405] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1721768 kB) free!
[1588583525] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34961656 kB) free.
[1588584125] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9699464 kB) free!
[1588584245] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (4995028 kB) free!
[1588584365] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (461776 kB) free!
[1588584664] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.4% (16260772 kB) free.
[1588584863] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(753) < 103.255.146.164 [04/May/2020:15:01:17 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588584983] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588585565] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24630216 kB) free!
[1588586165] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (438280 kB) free!
[1588586670] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (886564 kB) free!
[1588586783] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(760) < 103.255.146.164 [04/May/2020:15:26:45 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588586796] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1588586877] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 30.49, 19.64, 12.77
[1588586903] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588587250] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9435548 kB) free!
[1588587270] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.1% (1800860 kB) free!
[1588587365] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29927828 kB) free.
[1588587371] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (9157864 kB) free!
[1588587477] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 17.31, 23.31, 18.93
[1588587491] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.0% (8892052 kB) free!
[1588587965] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4051152 kB) free!
[1588588085] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (441172 kB) free!
[1588588206] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38040292 kB) free.
[1588588805] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9840280 kB) free!
[1588588925] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5355340 kB) free!
[1588589045] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1232232 kB) free!
[1588589645] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14093324 kB) free!
[1588590245] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27048192 kB) free.
[1588590845] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (3506216 kB) free!
[1588590870] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1531284 kB) free!
[1588590965] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35186352 kB) free.
[1588591091] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4831236 kB) free!
[1588591470] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (2586976 kB) free!
[1588591565] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7067868 kB) free!
[1588591685] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2252496 kB) free!
[1588591805] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (375828 kB) free!
[1588592368] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/04 17:06:26 [error] 22134#0: *86809022 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1588592489] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588593005] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29214872 kB) free.
[1588593064] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.6% (3941660 kB) free!
[1588593088] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/04 17:13:23 [error] 22133#0: *86809835 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1588593104] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.7% (2384052 kB) free!
[1588593184] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.6% (3725544 kB) free!
[1588593208] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588593270] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (901060 kB) free!
[1588593304] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.6% (3503700 kB) free!
[1588593605] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3701044 kB) free!
[1588593725] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (378580 kB) free!
[1588593846] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29001640 kB) free.
[1588593870] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.2% (2631996 kB) free!
[1588593904] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1533248 kB) free!
[1588594445] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (3025068 kB) free!
[1588594565] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (381200 kB) free!
[1588594685] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35113476 kB) free.
[1588595285] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6655084 kB) free!
[1588595405] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (1973384 kB) free!
[1588595525] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389492 kB) free!
[1588595704] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1588595704] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.6% (16292736 kB) free.
[1588596422] SERVICE ALERT: rq-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 354 processes
[1588596542] SERVICE ALERT: rq-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 352 processes
[1588596662] SERVICE ALERT: rq-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 354 processes
[1588596725] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (31022752 kB) free.
[1588597325] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4567220 kB) free!
[1588597445] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (381740 kB) free!
[1588597565] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39726188 kB) free.
[1588597837] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 11.7% (1904516 kB) free!
[1588598165] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13918680 kB) free!
[1588598285] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (9377392 kB) free!
[1588598405] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (5106836 kB) free!
[1588599006] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17676084 kB) free!
[1588599606] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33314088 kB) free.
[1588599688] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1588599808] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1588600077] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.31, 4.66, 4.95
[1588600206] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6698668 kB) free!
[1588600326] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2359404 kB) free!
[1588600446] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38552468 kB) free.
[1588600677] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.49, 4.77, 5.16
[1588600797] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.55, 4.32, 4.95
[1588601046] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11475776 kB) free!
[1588601071] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.6% (4001240 kB) free.
[1588601166] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6131740 kB) free!
[1588601286] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (1714596 kB) free!
[1588601397] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.74, 6.57, 6.04
[1588601517] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.47, 6.26, 6.00
[1588601637] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.77, 6.22, 5.99
[1588601671] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (1062440 kB) free!
[1588601791] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.7% (1899016 kB) free!
[1588601886] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14740596 kB) free!
[1588601911] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (3093852 kB) free!
[1588602486] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34459600 kB) free.
[1588602511] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3282920 kB) free.
[1588603086] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8639736 kB) free!
[1588603206] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4583940 kB) free!
[1588603326] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39108812 kB) free.
[1588603437] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.39, 3.29, 4.39
[1588603926] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13470664 kB) free!
[1588604046] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9258344 kB) free!
[1588604166] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (4983484 kB) free!
[1588604767] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17414188 kB) free!
[1588605367] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33388052 kB) free.
[1588605511] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3167124 kB) free!
[1588605631] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3286964 kB) free.
[1588605904] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1588605967] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8689764 kB) free!
[1588606087] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4544920 kB) free!
[1588606207] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37810096 kB) free.
[1588606231] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3208404 kB) free!
[1588606351] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (3217460 kB) free!
[1588606471] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (3195552 kB) free!
[1588606807] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13400308 kB) free!
[1588606927] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9126788 kB) free!
[1588607047] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (5059140 kB) free!
[1588607647] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16309784 kB) free!
[1588608247] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32726352 kB) free.
[1588608847] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9104232 kB) free!
[1588608967] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (5064960 kB) free!
[1588609087] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (36999780 kB) free.
[1588609687] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11819112 kB) free!
[1588609807] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7393892 kB) free!
[1588609928] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (2902928 kB) free!
[1588610008] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1588610071] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (1493388 kB) free!
[1588610128] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1588610528] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17195552 kB) free!
[1588610671] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.2% (2799916 kB) free!
[1588611128] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30388588 kB) free.
[1588611728] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6826460 kB) free!
[1588611848] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2850484 kB) free!
[1588611871] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 32.1% (5225732 kB) free.
[1588611968] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35119460 kB) free.
[1588612471] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3162372 kB) free!
[1588612568] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10506452 kB) free!
[1588612591] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3390268 kB) free.
[1588612688] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6630424 kB) free!
[1588612808] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (2854528 kB) free!
[1588614008] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.3% (25461748 kB) free!
[1588614608] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (2940624 kB) free!
[1588615208] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14068512 kB) free!
[1588616408] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (3158868 kB) free!