[1555439750] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 324 processes
[1555439871] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1555439902] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.2% (8463680 kB) free!
[1555439991] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 320 processes
[1555440022] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 29.6% (14584276 kB) free.
[1555440318] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52901 MB (23% inode=99%):
[1555440622] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9424688 kB) free!
[1555440742] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 29.6% (14594352 kB) free.
[1555440918] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41123 MB (18% inode=99%):
[1555441038] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41123 MB (18% inode=99%):
[1555441158] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41123 MB (18% inode=99%):
[1555441342] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9744192 kB) free!
[1555441462] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 29.4% (14495136 kB) free.
[1555441852] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 507.86, 379.63, 181.79
[1555441972] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 240.40, 351.34, 197.41
[1555442092] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 32.83, 235.25, 173.54
[1555442201] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2733404 kB) free!
[1555442391] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1555442475] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2448764 kB) free!
[1555442594] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 489.54, 275.45, 113.88
[1555442636] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 502.76, 395.47, 197.63
[1555442714] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 496.86, 349.10, 160.46
[1555442756] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 425.42, 414.34, 229.21
[1555442834] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 498.41, 398.54, 201.44
[1555442876] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 58.68, 277.76, 201.61
[1555442922] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 707 processes
[1555442991] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1555443042] SERVICE ALERT: sservercluster2;Total Processes;OK;HARD;1;PROCS OK: 461 processes
[1555443111] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1555443231] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 322 processes
[1555443558] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52387 MB (23% inode=99%):
[1555444158] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43427 MB (19% inode=99%):
[1555444278] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41368 MB (18% inode=99%):
[1555444398] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41370 MB (18% inode=99%):
[1555444492] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.84, 0.81, 13.69
[1555444601] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2471356 kB) free!
[1555445234] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.44, 0.58, 16.02
[1555445276] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.37, 1.47, 16.52
[1555445598] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53145 MB (23% inode=99%):
[1555445692] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.53, 0.60, 4.17
[1555446198] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42895 MB (19% inode=99%):
[1555446318] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42902 MB (19% inode=99%):
[1555446434] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.45, 0.48, 4.73
[1555446439] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 40852 MB (18% inode=99%):
[1555447076] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.64, 1.63, 3.77
[1555447639] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53653 MB (23% inode=99%):
[1555448239] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45461 MB (20% inode=99%):
[1555448359] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43925 MB (19% inode=99%):
[1555448479] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41877 MB (18% inode=99%):
[1555448632] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1555449232] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555449352] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555449472] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1555449679] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53644 MB (23% inode=99%):
[1555450279] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42051 MB (18% inode=99%):
[1555450399] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42051 MB (18% inode=99%):
[1555450519] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42058 MB (18% inode=99%):
[1555450673] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 319 processes
[1555451273] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1555451393] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555451513] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1555451720] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53644 MB (23% inode=99%):
[1555452320] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45676 MB (20% inode=99%):
[1555452440] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41571 MB (18% inode=99%):
[1555452560] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41578 MB (18% inode=99%):
[1555454960] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52616 MB (23% inode=99%):
[1555455560] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43592 MB (19% inode=99%):
[1555455680] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39550 MB (17% inode=99%):
[1555455800] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 39556 MB (17% inode=99%):
[1555457000] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51583 MB (22% inode=99%):
[1555457601] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41558 MB (18% inode=99%):
[1555457721] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41565 MB (18% inode=99%):
[1555457841] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 40030 MB (17% inode=99%):
[1555459041] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51582 MB (22% inode=99%):
[1555459641] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42229 MB (18% inode=99%):
[1555459761] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42236 MB (18% inode=99%):
[1555459881] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 39162 MB (17% inode=99%):
[1555459913] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1555460514] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1555460634] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1555460754] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1555461081] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52604 MB (23% inode=99%):
[1555461681] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45434 MB (20% inode=99%):
[1555461780] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57390 MB (25% inode=99%):
[1555461801] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43900 MB (19% inode=99%):
[1555461921] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41851 MB (18% inode=99%):
[1555461964] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.1% (14480756 kB) free.
[1555462492] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 429.10, 157.66, 57.61
[1555462612] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 499.47, 274.37, 112.51
[1555462732] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 508.53, 352.33, 160.71
[1555463201] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.6% (2866912 kB) free!
[1555463801] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2455808 kB) free!
[1555465132] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.47, 0.75, 16.58
[1555465522] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53427 MB (23% inode=99%):
[1555465551] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11679840 kB) free.
[1555465672] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.8% (22251984 kB) free.
[1555466122] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45234 MB (20% inode=99%):
[1555466156] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (2847648 kB) free!
[1555466242] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43699 MB (19% inode=99%):
[1555466276] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (1679916 kB) free!
[1555466332] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.44, 0.49, 4.89
[1555466362] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41651 MB (18% inode=99%):
[1555466399] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (225700 kB) free!
[1555466962] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48359 MB (21% inode=99%):
[1555467562] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40677 MB (18% inode=99%):
[1555467682] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40679 MB (18% inode=99%):
[1555467802] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52897 MB (23% inode=99%):
[1555468402] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45216 MB (20% inode=99%):
[1555468522] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44449 MB (19% inode=99%):
[1555468642] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43425 MB (19% inode=99%):
[1555468798] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11376924 kB) free.
[1555468881] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.3% (20808692 kB) free.
[1555468974] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 88.4% (21661196 kB) free.
[1555469242] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49312 MB (21% inode=99%):
[1555469403] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (2725516 kB) free!
[1555469499] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46089 MB (20% inode=99%):
[1555469523] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (1013856 kB) free!
[1555469619] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46081 MB (20% inode=99%):
[1555469644] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1555469644] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (216616 kB) free!
[1555469740] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46081 MB (20% inode=99%):
[1555469842] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41630 MB (18% inode=99%):
[1555469962] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41120 MB (18% inode=99%):
[1555470082] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53277 MB (23% inode=99%):
[1555470682] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45212 MB (20% inode=99%):
[1555470802] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44957 MB (20% inode=99%):
[1555470922] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43421 MB (19% inode=99%):
[1555471523] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49307 MB (21% inode=99%):
[1555472123] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43675 MB (19% inode=99%):
[1555472127] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11560520 kB) free.
[1555472155] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1555472243] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43675 MB (19% inode=99%):
[1555472276] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 324.70, 97.26, 34.85
[1555472363] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41626 MB (18% inode=99%):
[1555472396] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 467.17, 227.03, 89.97
[1555472516] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 486.98, 314.05, 138.48
[1555472731] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6918172 kB) free!
[1555472851] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (4862884 kB) free!
[1555472963] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47288 MB (21% inode=99%):
[1555472971] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (2721224 kB) free!
[1555473328] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 465.16, 349.48, 167.90
[1555473448] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 74.38, 244.53, 152.01
[1555473563] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42936 MB (19% inode=99%):
[1555473568] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 10.20, 163.68, 133.60
[1555473683] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39350 MB (17% inode=99%):
[1555473803] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53330 MB (23% inode=99%):
[1555474034] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 117.37, 296.28, 184.41
[1555474154] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16.16, 198.38, 162.10
[1555474274] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.57, 132.91, 142.51
[1555474403] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43497 MB (19% inode=99%):
[1555474523] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43504 MB (19% inode=99%):
[1555474643] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43506 MB (19% inode=99%):
[1555475243] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49294 MB (21% inode=99%):
[1555475368] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.31, 0.79, 19.53
[1555475391] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 15 users currently logged in
[1555475516] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.94, 1.90, 14.65
[1555475722] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11438744 kB) free.
[1555475843] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43166 MB (19% inode=99%):
[1555475963] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41118 MB (18% inode=99%):
[1555476083] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 37015 MB (16% inode=99%):
[1555476327] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.5% (5677604 kB) free!
[1555476349] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.29, 0.55, 15.65
[1555476447] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (3682600 kB) free!
[1555476567] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (1609000 kB) free!
[1555476601] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555476721] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555476841] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555477168] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.35, 0.34, 3.09
[1555477316] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.02, 1.88, 3.71
[1555477549] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.35, 0.52, 4.66
[1555477883] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49240 MB (21% inode=99%):
[1555478483] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41807 MB (18% inode=99%):
[1555478579] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169868 kB) free!
[1555478604] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41814 MB (18% inode=99%):
[1555478699] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (1519368 kB) free!
[1555478724] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 67388 MB (29% inode=99%):
[1555478819] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190380 kB) free!
[1555478876] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193512 kB) free!
[1555478996] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177592 kB) free!
[1555479116] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175844 kB) free!
[1555479340] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 88805 MB (39% inode=99%):
[1555479419] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555479419] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 85.8% (21030152 kB) free.
[1555479539] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555479659] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555481764] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.5% (3004708 kB) free!
[1555481884] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.8% (2895352 kB) free!
[1555482004] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.9% (2752944 kB) free!
[1555482349] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 414.39, 178.63, 67.97
[1555482469] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 446.36, 268.84, 114.44
[1555482589] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1555482589] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 140.13, 237.13, 124.16
[1555482937] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11307788 kB) free.
[1555483286] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.5% (4075184 kB) free!
[1555483406] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165344 kB) free!
[1555483526] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165420 kB) free!
[1555483541] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.2% (5040452 kB) free!
[1555483661] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (4115340 kB) free!
[1555483768] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 50.99, 13.20, 4.86
[1555483781] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (2898404 kB) free!
[1555483804] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1070652 kB) free!
[1555483888] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 8.60, 9.46, 4.50
[1555483980] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46839 MB (20% inode=99%):
[1555484100] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46418 MB (20% inode=99%):
[1555484220] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46383 MB (20% inode=99%):
[1555484389] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.23, 1.61, 19.35
[1555485926] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 40.4% (9960216 kB) free.
[1555486189] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.69, 0.83, 3.41
[1555486691] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10769024 kB) free.
[1555487297] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (3906432 kB) free!
[1555487418] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (1790004 kB) free!
[1555487538] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237320 kB) free!
[1555488324] SERVICE FLAPPING ALERT: db2;HDD_Home status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1555488863] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9190312 kB) free!
[1555488983] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 30.6% (15063164 kB) free.
[1555489441] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555489583] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9319432 kB) free!
[1555489703] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 31.0% (15268308 kB) free.
[1555489859] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555490303] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9386796 kB) free!
[1555490423] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.2% (13912936 kB) free.
[1555492259] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555492379] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555492441] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555492499] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555492561] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555492681] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555494716] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.38, 12.84, 11.28
[1555494836] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.64, 9.46, 10.23
[1555494957] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.34, 7.12, 9.28
[1555495279] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10084468 kB) free.
[1555495387] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.9% (15102408 kB) free.
[1555495452] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.4% (22655740 kB) free.
[1555495526] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4851816 kB) free!
[1555495646] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (4841060 kB) free!
[1555495681] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555495766] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (4468064 kB) free!
[1555495886] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4220952 kB) free!
[1555496006] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (2128404 kB) free!
[1555496126] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (229444 kB) free!
[1555496157] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.10, 2.18, 4.11
[1555496389] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1555496967] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176456 kB) free!
[1555497023] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.8% (6323600 kB) free!
[1555497143] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.8% (13692296 kB) free.
[1555497743] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8759940 kB) free!
[1555497835] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4893744 kB) free!
[1555497863] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.6% (5703368 kB) free!
[1555497954] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.6% (21111224 kB) free.
[1555497955] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (4835332 kB) free!
[1555497983] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.7% (13178452 kB) free.
[1555498075] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (4781260 kB) free!
[1555498460] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185732 kB) free!
[1555498580] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180968 kB) free!
[1555498583] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.6% (8201368 kB) free!
[1555498700] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179632 kB) free!
[1555498703] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.6% (5228020 kB) free!
[1555498823] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.9% (12759340 kB) free.
[1555499157] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 480.76, 384.73, 195.20
[1555499277] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 407.30, 400.80, 224.72
[1555499397] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 56.41, 268.77, 197.69
[1555499423] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.7% (6240192 kB) free!
[1555499543] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (3263236 kB) free!
[1555499663] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.1% (11892880 kB) free.
[1555500194] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 20.0% (3257380 kB) free!
[1555500264] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.8% (7284048 kB) free!
[1555500314] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.1% (3098408 kB) free!
[1555500384] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (4036364 kB) free!
[1555500419] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (4440508 kB) free!
[1555500434] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.2% (2955400 kB) free!
[1555500504] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.5% (11071652 kB) free.
[1555500539] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4417528 kB) free!
[1555500659] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 24.8% (6080324 kB) free.
[1555501104] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.2% (6993712 kB) free!
[1555501224] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (3779572 kB) free!
[1555501344] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (702828 kB) free!
[1555501554] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1555501797] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1555501797] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.37, 1.99, 16.60
[1555501859] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.9% (4146848 kB) free!
[1555501944] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.1% (6959440 kB) free!
[1555501979] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.7% (4097608 kB) free!
[1555502099] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.5% (4054724 kB) free!
[1555502154] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1555502235] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1287516 kB) free!
[1555502274] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 320 processes
[1555502394] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1555503144] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.6% (13589692 kB) free.
[1555503597] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.29, 1.50, 3.80
[1555503744] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9092836 kB) free!
[1555503864] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.8% (5823984 kB) free!
[1555503899] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2460052 kB) free!
[1555503984] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 29.0% (14294064 kB) free.
[1555505184] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.6% (5697468 kB) free!
[1555505305] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.1% (13330680 kB) free.
[1555506504] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.3% (6040100 kB) free!
[1555506624] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (2733172 kB) free!
[1555506744] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.6% (12124956 kB) free.
[1555506926] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1555507075] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2283536 kB) free!
[1555507344] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.5% (7637632 kB) free!
[1555507464] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (4274116 kB) free!
[1555507584] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1555507584] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.9% (9300152 kB) free!
[1555508184] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.2% (11923724 kB) free.
[1555508784] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.3% (8016268 kB) free!
[1555508904] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (4583708 kB) free!
[1555509024] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (1400348 kB) free!
[1555509624] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.4% (10524588 kB) free.
[1555510191] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1555510825] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.0% (6913764 kB) free!
[1555510945] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (3749452 kB) free!
[1555511065] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.1% (13373416 kB) free.
[1555512265] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.2% (8004008 kB) free!
[1555512363] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178400 kB) free!
[1555512385] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (4685072 kB) free!
[1555512483] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173888 kB) free!
[1555512505] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.1% (13864388 kB) free.
[1555512603] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172744 kB) free!
[1555513705] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.0% (6921992 kB) free!
[1555513797] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1555513825] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (3749524 kB) free!
[1555513945] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.6% (13120496 kB) free.
[1555514545] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9234960 kB) free!
[1555514665] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.2% (5994712 kB) free!
[1555514786] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (2864792 kB) free!
[1555515386] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.6% (11614924 kB) free.
[1555515986] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.7% (7743564 kB) free!
[1555516106] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (4563132 kB) free!
[1555516226] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.9% (13240476 kB) free.
[1555516826] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9452296 kB) free!
[1555516946] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.6% (6211208 kB) free!
[1555517066] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (3108844 kB) free!
[1555517394] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1555517666] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.2% (11450132 kB) free.
[1555517994] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1555518115] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1555518235] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555518266] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.5% (8109256 kB) free!
[1555518386] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.3% (5073208 kB) free!
[1555518506] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.6% (13097648 kB) free.
[1555519706] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.6% (6231560 kB) free!
[1555519826] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (3187528 kB) free!
[1555519946] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.8% (11732876 kB) free.
[1555520546] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.6% (7687512 kB) free!
[1555520666] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (4915992 kB) free!
[1555520786] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.4% (13024180 kB) free.
[1555521386] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9128688 kB) free!
[1555521506] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.7% (6269900 kB) free!
[1555521626] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (3292776 kB) free!
[1555522226] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.0% (10334648 kB) free.
[1555522826] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.2% (6006232 kB) free!
[1555522946] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.5% (13079728 kB) free.
[1555523547] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9613816 kB) free!
[1555523667] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.2% (6492708 kB) free!
[1555523787] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.6% (13100892 kB) free.
[1555524987] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.8% (5314180 kB) free!
[1555525107] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.4% (13025752 kB) free.
[1555525707] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.0% (8876924 kB) free!