[1544331342] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 289 processes
[1544331462] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544331582] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1544332844] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.1% (6459328 kB) free!
[1544332860] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544332965] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10392868 kB) free.
[1544332980] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544344782] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544344902] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544345022] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1544347964] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.6% (6713520 kB) free!
[1544347980] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1544348084] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.0% (10325176 kB) free.
[1544348101] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1544353599] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1544354860] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 8.38, 8.58, 9.43
[1544356035] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.4% (3149028 kB) free!
[1544356155] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.1% (3109376 kB) free!
[1544356275] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.8% (3054444 kB) free!
[1544357022] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544357085] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4583236 kB) free!
[1544357143] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544357205] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 17.1% (8421384 kB) free!
[1544357263] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1544357325] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.4% (10038664 kB) free.
[1544357700] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544357820] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1544357925] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9836152 kB) free!
[1544357940] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544358045] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 15.7% (7737224 kB) free!
[1544358165] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.9% (9823496 kB) free!
[1544358541] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1544358661] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544358765] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4125312 kB) free!
[1544358781] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544359365] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.5% (9598640 kB) free!
[1544365875] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1633452 kB) free!
[1544366475] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 10.1% (1634520 kB) free!
[1544367075] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1544367075] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1632628 kB) free!
[1544367676] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 10.1% (1634836 kB) free!
[1544368276] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1633428 kB) free!
[1544368982] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544369102] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544369222] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544369263] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544369383] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544369398] SERVICE ALERT: sig1;Memory;WARNING;SOFT;1;WARNING - 19.5% (3169384 kB) free!
[1544369503] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544369518] SERVICE ALERT: sig1;Memory;WARNING;SOFT;2;WARNING - 19.4% (3149952 kB) free!
[1544369638] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 19.3% (3129392 kB) free!
[1544370103] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 288 processes
[1544371903] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544371966] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4265400 kB) free!
[1544372023] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544372143] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544372222] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544372565] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.5% (6143692 kB) free!
[1544372823] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544372943] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1544373063] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1544373663] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544373783] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1544373903] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1544374544] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1544374544] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1544375703] SERVICE FLAPPING ALERT: db1;CPU_Procs;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1544375703] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1544379676] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)