[1546541556] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546541676] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546541796] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1546542396] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546543046] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1546543166] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1546545996] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 289 processes
[1546546116] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546546236] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546546836] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1546549401] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 320.84, 137.02, 50.48
[1546549521] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 43.74, 91.85, 44.42
[1546549641] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 6.58, 61.67, 39.11
[1546550841] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.54, 1.59, 11.09
[1546551036] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546551156] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546551276] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1546552041] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.39, 0.46, 3.38
[1546552641] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.35, 4.59, 6.61
[1546552761] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.48, 3.24, 5.87
[1546552881] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1546552881] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.47, 2.35, 5.22
[1546553481] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.37, 0.79, 3.01
[1546553966] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1546554086] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1546555476] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546555596] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546555716] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546558716] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1546559316] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546559436] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546559556] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546563081] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.57, 4.51, 5.21
[1546563156] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546563201] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.41, 3.13, 4.62
[1546563276] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546563396] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546563996] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546565001] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1546566686] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1546566806] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1546567596] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546567716] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546567836] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546568436] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1546568436] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546570401] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 2.79, 51.69, 37.19
[1546570521] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.71, 34.73, 32.73
[1546570642] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.39, 23.37, 28.82
[1546571242] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.64, 3.68, 15.37
[1546571606] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1546571726] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1546572036] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546572156] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546572276] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1546572442] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.75, 0.65, 4.61
[1546576476] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546576596] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546576717] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546578442] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 0.68, 32.54, 28.89
[1546578562] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.75, 22.01, 25.47
[1546578682] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1546578682] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.79, 14.98, 22.47
[1546578872] SERVICE ALERT: sig5;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1546578874] SERVICE ALERT: sig5;Memory;CRITICAL;SOFT;1;(Return code of 255 for service 'Memory' on host 'sig5' was out of bounds)
[1546578874] HOST ALERT: sig5;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.57)
[1546578874] SERVICE ALERT: sig5;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig5' was out of bounds)
[1546578992] SERVICE ALERT: sig5;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 2 processes with command name 'ntpd'
[1546578993] HOST ALERT: sig5;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1546578994] SERVICE ALERT: sig5;Sensors Status;OK;HARD;1;SENSORS OK
[1546578994] SERVICE ALERT: sig5;Memory;OK;HARD;1;OK - 94.7% (15390740 kB) free.
[1546579282] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.56, 2.65, 12.13
[1546579919] HOST ALERT: sig5;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1546579950] SERVICE ALERT: sig5;CPU_Procs;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1546580015] HOST ALERT: sig5;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1546580070] SERVICE ALERT: sig5;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1546580316] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1546580482] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.24, 1.07, 3.99
[1546580916] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 296 processes
[1546581036] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1546581157] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1546581683] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.66, 5.41, 7.00
[1546581756] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1546581803] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.72, 3.85, 6.23
[1546581876] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1546581923] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.40, 2.71, 5.52
[1546581996] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1546582523] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.71, 0.95, 3.21
[1546582596] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1546584270] SERVICE ALERT: sig5;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546584399] SERVICE ALERT: sig5;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1546584400] SERVICE ALERT: sig5;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1546584402] SERVICE ALERT: sig5;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1546584403] SERVICE ALERT: sig5;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1546584404] SERVICE ALERT: sig5;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1546584404] SERVICE ALERT: sig5;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1546584429] HOST ALERT: sig5;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1546584450] SERVICE ALERT: sig5;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig5' was out of bounds)
[1546584457] SERVICE ALERT: sig5;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig5' was out of bounds)
[1546584519] SERVICE ALERT: sig5;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 221454 MB (98% inode=99%):
[1546584520] SERVICE ALERT: sig5;CPU_Procs;CRITICAL;HARD;2;CPU CRITICAL: 2 crit, 0 warn out of 306 processes
[1546584522] SERVICE ALERT: sig5;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 2 processes with command name 'ntpd'
[1546584523] SERVICE ALERT: sig5;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38807 MB (75% inode=98%):
[1546584523] HOST ALERT: sig5;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1546584524] SERVICE ALERT: sig5;Sensors Status;OK;HARD;1;SENSORS OK
[1546584524] SERVICE ALERT: sig5;Memory;OK;HARD;1;OK - 94.3% (15332968 kB) free.
[1546584570] SERVICE ALERT: sig5;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1546584577] SERVICE ALERT: sig5;Total Processes;OK;HARD;1;PROCS OK: 312 processes
[1546585120] SERVICE FLAPPING ALERT: sig5;CPU_Procs;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1546585120] SERVICE ALERT: sig5;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1546586723] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 21.33, 56.38, 28.68
[1546586843] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.37, 37.95, 25.28
[1546586963] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.14, 25.67, 22.32
[1546587563] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.77, 4.08, 12.04
[1546588421] SERVICE ALERT: sig5;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1546588423] HOST ALERT: sig5;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.57)
[1546588541] SERVICE ALERT: sig5;CPU Load;OK;HARD;1;OK - load average: 0.46, 0.16, 0.06
[1546588545] HOST ALERT: sig5;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1546588763] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.88, 0.83, 3.88
[1546589276] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1546589396] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1546589516] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1546591164] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.93, 11.27, 11.40
[1546591284] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.77, 7.74, 10.09
[1546591404] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.47, 5.34, 8.93
[1546592004] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.58, 1.29, 5.00
[1546594596] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1546594716] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1546594837] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1546596520] SERVICE FLAPPING ALERT: sig5;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1546597237] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546600837] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 289 processes
[1546600957] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546601077] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546602877] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1546602877] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546603404] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1546606404] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 0.91, 33.98, 35.57
[1546606524] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.10, 23.11, 31.39
[1546606644] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.11, 15.82, 27.72
[1546607244] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.89, 2.86, 14.92
[1546608444] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.64, 0.73, 4.62
[1546611277] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546611397] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546611517] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546612044] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 85.82, 79.06, 33.57
[1546612165] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 12.07, 53.08, 29.56
[1546612285] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1546612285] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.18, 35.70, 26.04
[1546612885] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.29, 7.84, 15.01
[1546614509] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 6 warn out of 371 processes
[1546614517] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546614685] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.54, 1.08, 3.39
[1546614926] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1546615046] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1546616309] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 5 warn out of 375 processes
[1546617116] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1546617685] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 21.05, 5.30, 3.11
[1546617805] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.28, 3.73, 2.81
[1546618117] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546618237] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546618357] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546623757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546623877] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1546623877] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546624109] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 6 warn out of 363 processes
[1546624709] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1546624709] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 5 warn out of 360 processes
[1546625005] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)