[1565980645] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192436 kB) free!
[1565980655] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189388 kB) free!
[1565980665] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565980757] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6682852 kB) free!
[1565980765] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (218932 kB) free!
[1565980780] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190556 kB) free!
[1565980785] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565980877] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11813648 kB) free.
[1565980885] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (212288 kB) free!
[1565980900] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186696 kB) free!
[1565980942] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187524 kB) free!
[1565981062] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187152 kB) free!
[1565981182] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199568 kB) free!
[1565981385] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565981477] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6757236 kB) free!
[1565981505] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565981597] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11821144 kB) free.
[1565981754] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1565982105] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565982197] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6821792 kB) free!
[1565982225] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565982317] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11812208 kB) free.
[1565982826] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565982917] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6645000 kB) free!
[1565982946] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565982993] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1565983038] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11820236 kB) free.
[1565983546] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565983637] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6690716 kB) free!
[1565983666] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565983757] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11823280 kB) free.
[1565983907] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173544 kB) free!
[1565983967] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.1% (4936576 kB) free.
[1565984266] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565984357] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7052836 kB) free!
[1565984386] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565984477] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11814264 kB) free.
[1565984485] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 44.4% (10945000 kB) free.
[1565984506] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565985077] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7609896 kB) free!
[1565985106] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565985197] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11821976 kB) free.
[1565985226] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1565985797] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7940048 kB) free!
[1565985827] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565985917] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11810644 kB) free.
[1565985947] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565986517] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8236648 kB) free!
[1565986547] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1565986637] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.1% (6446712 kB) free!
[1565986667] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565986757] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11821700 kB) free.
[1565987267] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565987358] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7152684 kB) free!
[1565987387] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565987478] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11813764 kB) free.
[1565987507] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1565988078] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7853268 kB) free!
[1565988107] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565988198] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11824392 kB) free.
[1565988228] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565988685] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177144 kB) free!
[1565988798] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7967028 kB) free!
[1565988805] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192684 kB) free!
[1565988827] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565988918] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11810304 kB) free.
[1565988925] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182064 kB) free!
[1565988947] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565989121] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (250372 kB) free!
[1565989241] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (223752 kB) free!
[1565989361] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (251176 kB) free!
[1565989518] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7953476 kB) free!
[1565989547] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565989638] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11822072 kB) free.
[1565989667] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565989961] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.6% (44635232 kB) free.
[1565990238] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7962552 kB) free!
[1565990267] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565990358] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11821428 kB) free.
[1565990387] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565990958] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7959660 kB) free!
[1565990987] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565991079] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11815944 kB) free.
[1565991108] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1565991678] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8028820 kB) free!
[1565991700] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 63.8% (15628460 kB) free.
[1565991707] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565991798] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11820488 kB) free.
[1565991827] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565991982] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 45.1% (11065016 kB) free.
[1565992398] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8176448 kB) free!
[1565992427] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565992518] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11813700 kB) free.
[1565992547] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1565993118] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8393496 kB) free!
[1565993147] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565993238] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.5% (6631224 kB) free!
[1565993267] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565993358] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11812968 kB) free.
[1565993867] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565993958] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6757512 kB) free!
[1565993988] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565994078] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11804224 kB) free.
[1565994587] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565994678] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6866840 kB) free!
[1565994707] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565994798] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11809064 kB) free.
[1565994827] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565995399] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7081280 kB) free!
[1565995427] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565995519] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11801380 kB) free.
[1565995547] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565996119] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7103872 kB) free!
[1565996147] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565996239] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11810020 kB) free.
[1565996267] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565996839] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7597312 kB) free!
[1565996867] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565996959] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11811728 kB) free.
[1565996988] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565997559] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7929284 kB) free!
[1565997587] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565997679] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11804072 kB) free.
[1565997707] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565998279] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8203136 kB) free!
[1565998307] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565998399] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.1% (6435204 kB) free!
[1565998427] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565998519] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11802992 kB) free.
[1565999027] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565999119] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7075932 kB) free!
[1565999147] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1565999239] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11807688 kB) free.
[1565999267] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565999840] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7476288 kB) free!
[1565999868] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565999960] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11809380 kB) free.
[1565999988] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1566000560] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7649308 kB) free!
[1566000588] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566000680] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11801168 kB) free.
[1566000708] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1566001280] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7750448 kB) free!
[1566001308] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566001400] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11809892 kB) free.
[1566001428] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566002000] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7821824 kB) free!
[1566002028] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566002120] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11798344 kB) free.
[1566002148] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1566002720] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7873856 kB) free!
[1566002749] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1566002840] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11805352 kB) free.
[1566002867] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4913000 kB) free!
[1566002869] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566002927] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.1% (4917960 kB) free.
[1566002945] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 590 processes
[1566003005] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 593 processes
[1566003227] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4822544 kB) free!
[1566003287] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 19.7% (4823000 kB) free!
[1566003347] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 19.7% (4823096 kB) free!
[1566003407] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 19.7% (4821048 kB) free!
[1566003440] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7896360 kB) free!
[1566003469] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1566003560] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11807284 kB) free.
[1566003589] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1566003925] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 57.6% (14202364 kB) free.
[1566004161] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7860652 kB) free!
[1566004189] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566004281] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11656156 kB) free.
[1566004309] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566004881] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7994808 kB) free!
[1566004909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1566005001] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11663056 kB) free.
[1566005029] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566005125] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.1% (4467456 kB) free!
[1566005182] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1566005182] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.3% (3248692 kB) free!
[1566005245] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192832 kB) free!
[1566005302] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1566005303] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.2% (3237008 kB) free!
[1566005365] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196304 kB) free!
[1566005423] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 13.1% (3210860 kB) free!
[1566005453] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1566005500] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164888 kB) free!
[1566005573] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1566005601] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8273632 kB) free!
[1566005623] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191488 kB) free!
[1566005630] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566005693] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1566005721] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.4% (6606860 kB) free!
[1566005746] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189724 kB) free!
[1566005750] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1566005843] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11657996 kB) free.
[1566005870] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1566006290] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 356 processes
[1566006410] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 369 processes
[1566006443] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7140984 kB) free!
[1566006470] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566006563] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11661172 kB) free.
[1566006590] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566006637] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 42.83, 17.97, 7.64
[1566006717] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 549.74, 282.81, 114.62
[1566006757] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 47.30, 27.98, 12.57
[1566006837] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 565.27, 377.10, 169.56
[1566006878] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 47.17, 34.26, 16.74
[1566006957] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 567.58, 440.25, 217.87
[1566007163] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7495940 kB) free!
[1566007190] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566007283] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11654528 kB) free.
[1566007310] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566007758] SERVICE ALERT: rq-1;CPU Load;OK;HARD;1;OK - load average: 0.76, 1.25, 4.76
[1566008090] SERVICE ALERT: rq-1;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1566008753] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 349 processes
[1566010318] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 354 processes
[1566010490] SERVICE ALERT: rq-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 368 processes
[1566010554] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 358 processes
[1566010610] SERVICE ALERT: rq-1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1566010697] SERVICE ALERT: smscapp1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 474 processes
[1566010702] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1566010730] SERVICE ALERT: rq-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1566010758] SERVICE ALERT: rq-1;CPU Load;WARNING;SOFT;1;WARNING - load average: 21.23, 14.38, 6.62
[1566010817] SERVICE ALERT: smscapp1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 476 processes
[1566010822] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1566010834] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1566010878] SERVICE ALERT: rq-1;CPU Load;WARNING;SOFT;2;WARNING - load average: 20.59, 16.49, 8.33
[1566010883] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8219940 kB) free!
[1566010894] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1566010910] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1566010918] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 355 processes
[1566010938] SERVICE ALERT: smscapp1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1566010954] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1566010998] SERVICE ALERT: rq-1;CPU Load;WARNING;HARD;3;WARNING - load average: 19.37, 17.31, 9.62
[1566011003] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.4% (6590344 kB) free!
[1566011014] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 62 users currently logged in
[1566011030] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566011124] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11305884 kB) free.
[1566011150] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1566011750] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1566011757] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.96, 5.32, 14.14
[1566012514] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1566012624] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189384 kB) free!
[1566012923] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6647080 kB) free!
[1566013043] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11652628 kB) free.
[1566013224] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 36.7% (8997412 kB) free.
[1566013414] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 65 users currently logged in
[1566013478] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.05, 8.98, 17.85
[1566013546] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 46.4% (11370932 kB) free.
[1566013550] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1566013643] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6943544 kB) free!
[1566013670] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566013763] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11652832 kB) free.
[1566013790] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566014157] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.83, 2.82, 3.97
[1566014363] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7376400 kB) free!
[1566014390] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566014483] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11644280 kB) free.
[1566014510] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566014963] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.9% (40838768 kB) free.
[1566014994] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1566015019] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.1% (14648704 kB) free.
[1566015083] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7858824 kB) free!
[1566015087] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.5% (22825768 kB) free.
[1566015110] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566015203] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11653388 kB) free.
[1566015230] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1566015692] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (2113652 kB) free!
[1566015803] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8219404 kB) free!
[1566015812] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188740 kB) free!
[1566015831] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566015878] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.99, 3.15, 4.53
[1566015923] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.4% (6617492 kB) free!
[1566015932] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192516 kB) free!
[1566015951] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1566016044] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11647700 kB) free.
[1566016071] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566016114] SERVICE FLAPPING ALERT: webserver2;Current Users;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1566016114] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 57 users currently logged in
[1566016532] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.4% (20584488 kB) free.
[1566016643] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7802716 kB) free!
[1566016671] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566016763] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6401840 kB) free!
[1566016791] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1566016883] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11644340 kB) free.
[1566017078] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.23, 5.41, 5.12
[1566017198] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.91, 5.43, 5.15
[1566017318] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1566017318] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 8.26, 6.37, 5.52
[1566017391] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566017483] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7462312 kB) free!
[1566017511] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566017603] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11649176 kB) free.
[1566017631] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566018203] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8993548 kB) free!
[1566018232] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566018323] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7939976 kB) free!
[1566018352] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1566018443] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.4% (6625368 kB) free!
[1566018472] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566018518] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.26, 4.01, 4.80
[1566019414] SERVICE FLAPPING ALERT: webserver2;Current Users;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1566019643] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11645440 kB) free.
[1566019672] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1566019672] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566020843] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8745904 kB) free!
[1566020872] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566020919] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.39, 5.64, 5.28
[1566020963] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7710324 kB) free!
[1566020992] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566021024] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (4607556 kB) free!
[1566021024] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.5% (2358500 kB) free!
[1566021039] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.99, 5.45, 5.25
[1566021083] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11484452 kB) free.
[1566021112] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1566021144] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.5% (4534716 kB) free!
[1566021144] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.5% (2362060 kB) free!
[1566021159] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.40, 5.46, 5.28
[1566021264] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.2% (4451304 kB) free!
[1566021264] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 14.5% (2362004 kB) free!
[1566021622] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1566021683] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7661768 kB) free!
[1566021713] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566021742] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1566021803] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11635860 kB) free.
[1566021833] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566021865] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (171868 kB) free!
[1566022403] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7919572 kB) free!
[1566022433] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566022523] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.7% (6261652 kB) free!
[1566022553] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1566022643] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11632252 kB) free.
[1566023153] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1566023243] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6939748 kB) free!
[1566023273] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566023363] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11634784 kB) free.
[1566023374] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (227628 kB) free!
[1566023393] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1566023494] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (247928 kB) free!
[1566023614] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (249172 kB) free!
[1566023963] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7439852 kB) free!
[1566023993] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566024083] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11627884 kB) free.
[1566024114] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566024264] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2376124 kB) free!
[1566024346] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (4446012 kB) free!
[1566024466] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (4310392 kB) free!
[1566024586] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (4159448 kB) free!
[1566024684] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7799696 kB) free!
[1566024713] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566024804] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11633668 kB) free.
[1566024833] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566025404] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8084044 kB) free!
[1566025433] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566025524] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6232600 kB) free!
[1566025553] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1566025644] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11632008 kB) free.
[1566026153] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1566026244] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6678776 kB) free!
[1566026273] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566026364] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11631204 kB) free.
[1566026386] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2183164 kB) free!
[1566026393] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1566026964] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7070844 kB) free!
[1566026993] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566027084] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11624064 kB) free.
[1566027114] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566027684] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7531624 kB) free!
[1566027713] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566027804] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11627324 kB) free.
[1566027833] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566028404] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7755964 kB) free!
[1566028433] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566028525] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11615196 kB) free.
[1566028553] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1566028942] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1566029062] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566029124] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8013132 kB) free!
[1566029153] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1566029244] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11522748 kB) free.
[1566029273] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566029844] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8224256 kB) free!
[1566029873] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566029964] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6411512 kB) free!
[1566029994] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1566030084] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11617396 kB) free.
[1566030593] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566030684] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6900752 kB) free!
[1566030713] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566030759] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1566030804] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11625640 kB) free.
[1566030833] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566031404] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7589796 kB) free!
[1566031433] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566031524] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11614628 kB) free.
[1566031553] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1566032124] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7761236 kB) free!
[1566032153] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1566032166] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.0% (22307556 kB) free.
[1566032208] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.1% (22327836 kB) free.
[1566032244] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11622796 kB) free.
[1566032263] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.9% (40848768 kB) free.
[1566032273] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1566032352] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.1% (14652800 kB) free.
[1566032844] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7992216 kB) free!
[1566032874] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566032965] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11516484 kB) free.
[1566032994] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1566032994] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1566033564] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8392960 kB) free!
[1566033594] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566033684] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6848144 kB) free!
[1566033714] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566033804] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11622848 kB) free.
[1566033834] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566034404] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7326656 kB) free!
[1566034434] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566034524] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11608208 kB) free.
[1566034554] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1566034557] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 508.75, 204.26, 77.33
[1566034677] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 565.10, 326.56, 137.54
[1566034761] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (292148 kB) free!
[1566034797] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 572.60, 408.34, 190.42
[1566034881] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (161776 kB) free!
[1566034961] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9051492 kB) free!
[1566035002] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (152884 kB) free!
[1566035081] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.8% (8787460 kB) free!
[1566035124] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7600956 kB) free!
[1566035154] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1566035201] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.3% (8503388 kB) free!
[1566035244] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11620236 kB) free.
[1566035275] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566035844] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7851340 kB) free!
[1566035874] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566035964] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11617960 kB) free.
[1566035994] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1566036293] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1566036564] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8073172 kB) free!
[1566036594] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566036685] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.8% (6291064 kB) free!
[1566036714] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566036805] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11614876 kB) free.
[1566037010] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1566037072] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8810544 kB) free!
[1566037130] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1566037193] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.7% (7714040 kB) free!
[1566037250] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1566037313] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.5% (6659128 kB) free!
[1566037314] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566037405] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6532028 kB) free!
[1566037434] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1566037525] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11604760 kB) free.
[1566037601] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4405692 kB) free!
[1566037797] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.70, 2.55, 14.35
[1566038034] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566038125] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6894524 kB) free!
[1566038155] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1566038203] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.0% (4409660 kB) free!
[1566038245] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11611244 kB) free.
[1566038275] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566038323] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (2290424 kB) free!
[1566038443] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (1140940 kB) free!
[1566038450] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1566038513] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4810936 kB) free!
[1566038559] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.23, 4.57, 4.79
[1566038845] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7170392 kB) free!
[1566038875] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566038965] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11614736 kB) free.
[1566038995] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566039113] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1566039113] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.2% (5006828 kB) free!
[1566039262] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1566039383] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566039565] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7574340 kB) free!
[1566039595] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566039597] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.37, 2.67, 4.37
[1566039685] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11608768 kB) free.
[1566039713] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (3120280 kB) free!
[1566039715] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566039760] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.22, 5.65, 5.23
[1566039880] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.58, 5.69, 5.30
[1566040000] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.60, 5.91, 5.44
[1566040243] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 31.0% (7593900 kB) free.
[1566040250] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1566040285] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7793760 kB) free!
[1566040315] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566040370] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1566040405] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11614136 kB) free.
[1566040435] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566040490] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1566040600] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.95, 4.51, 4.99
[1566041006] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8015872 kB) free!
[1566041036] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566041126] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11606620 kB) free.
[1566041156] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1566041726] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8148104 kB) free!
[1566041756] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1566041846] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6390552 kB) free!
[1566041876] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566041966] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11611272 kB) free.
[1566042043] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197152 kB) free!
[1566042163] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188884 kB) free!
[1566042199] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9461548 kB) free!
[1566042283] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184292 kB) free!
[1566042290] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1566042290] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1566042319] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.9% (9319976 kB) free!
[1566042439] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.6% (9182888 kB) free!
[1566042476] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566042566] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7185376 kB) free!
[1566042596] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566042686] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11602404 kB) free.
[1566042705] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.1% (14646464 kB) free.
[1566042716] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1566042763] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1566042763] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.2% (21622144 kB) free.
[1566042802] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.9% (40853028 kB) free.
[1566042841] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 34.4% (16973412 kB) free.
[1566043286] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7726148 kB) free!
[1566043316] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1566043406] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11610008 kB) free.
[1566043436] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566044006] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7978388 kB) free!
[1566044037] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566044126] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11503500 kB) free.
[1566044157] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1566044691] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1566044726] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8172232 kB) free!
[1566044757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1566044811] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1566044847] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.1% (6431352 kB) free!
[1566044877] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566044967] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11608752 kB) free.
[1566045440] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 68.8% (33904664 kB) free.
[1566045477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566045567] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6652700 kB) free!
[1566045597] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566045687] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11599032 kB) free.
[1566045714] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (1125836 kB) free!
[1566045717] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566045834] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (156008 kB) free!
[1566045954] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156720 kB) free!
[1566046011] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1566046131] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1566046251] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1566046287] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7499784 kB) free!
[1566046317] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566046407] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11611456 kB) free.
[1566046438] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566046851] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1566047007] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7888212 kB) free!
[1566047037] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566047127] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11601192 kB) free.
[1566047157] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1566047452] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1566047572] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1566047572] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.3% (4740404 kB) free!
[1566047692] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 351 processes
[1566047692] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 14.6% (3574820 kB) free!
[1566047727] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8182220 kB) free!
[1566047757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1566047813] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.6% (3568876 kB) free!
[1566047847] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6408744 kB) free!
[1566047877] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566047967] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11608660 kB) free.
[1566048292] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1566048477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566048539] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 381 processes
[1566048568] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6627116 kB) free!
[1566048597] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1566048659] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 383 processes
[1566048688] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11601252 kB) free.
[1566048779] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 383 processes
[1566048892] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1566049012] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1566049132] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1566049197] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566049288] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6733040 kB) free!
[1566049317] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566049408] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11606176 kB) free.
[1566049414] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1566049437] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566049613] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2301584 kB) free!
[1566050008] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7035348 kB) free!
[1566050038] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1566050128] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11596848 kB) free.
[1566050158] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1566050213] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 40.1% (9833348 kB) free.
[1566050613] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9655120 kB) free!
[1566050728] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7740676 kB) free!
[1566050733] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.8% (8761004 kB) free!
[1566050758] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1566050813] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 11.6% (2831560 kB) free!
[1566050848] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11605272 kB) free.
[1566050853] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.9% (7847004 kB) free!
[1566050878] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566050933] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (1278204 kB) free!
[1566051053] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 32.8% (8040808 kB) free.
[1566051179] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 380 processes
[1566051241] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.0% (14624180 kB) free.
[1566051299] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 378 processes
[1566051448] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7895924 kB) free!
[1566051478] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566051532] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1566051568] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11608188 kB) free.
[1566051598] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1566051652] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1566051653] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177628 kB) free!
[1566051773] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169672 kB) free!
[1566051893] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189920 kB) free!
[1566052168] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8024296 kB) free!
[1566052198] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566052252] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1566052288] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.5% (6183360 kB) free!
[1566052319] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566052372] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1566052409] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11604872 kB) free.
[1566052492] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1566052918] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566053008] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6488708 kB) free!
[1566053038] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566053128] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11596440 kB) free.
[1566053638] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566053692] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1566053728] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6628140 kB) free!
[1566053758] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1566053848] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11601988 kB) free.
[1566053853] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4630104 kB) free!
[1566054292] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1566054293] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 15.4% (3766180 kB) free!
[1566054358] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1566054412] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1566054448] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6717588 kB) free!
[1566054478] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566054532] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1566054568] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11603260 kB) free.
[1566054598] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566054845] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.5% (2520344 kB) free!
[1566054965] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.5% (2357812 kB) free!
[1566054982] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1566055085] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 14.5% (2361260 kB) free!
[1566055102] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566055168] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6891956 kB) free!
[1566055199] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566055225] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;1;CRITICAL - 5.1% (1257820 kB) free!
[1566055288] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11596832 kB) free.
[1566055319] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1566055345] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;2;CRITICAL - 7.3% (1798808 kB) free!
[1566055465] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 11.8% (2902604 kB) free!
[1566055493] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197556 kB) free!
[1566055888] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7084720 kB) free!
[1566055919] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1566056008] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11604736 kB) free.
[1566056039] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566056065] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 7.4% (1809336 kB) free!
[1566056093] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 27.2% (6657700 kB) free.
[1566056608] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7229964 kB) free!
[1566056639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1566056693] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 12.2% (2996904 kB) free!
[1566056729] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11594288 kB) free.
[1566056759] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1566056813] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 12.2% (2993380 kB) free!
[1566056885] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1566056885] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (631304 kB) free!
[1566056933] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1566056933] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 12.2% (2990164 kB) free!
[1566057328] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7361972 kB) free!
[1566057359] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1566057448] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11600916 kB) free.
[1566057479] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1566057899] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 382 processes
[1566057932] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (4703672 kB) free!
[1566058019] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 380 processes
[1566058048] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7434604 kB) free!
[1566058052] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.8% (4398872 kB) free!
[1566058080] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566058133] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 31.3% (7672592 kB) free.
[1566058140] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 379 processes
[1566058168] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11603164 kB) free.
[1566058172] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.3% (4019340 kB) free!
[1566058182] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 586.42, 414.57, 191.30
[1566058200] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566058302] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 587.05, 471.68, 239.30
[1566058422] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 254.93, 419.41, 250.70
[1566058740] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 379 processes
[1566058768] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7642248 kB) free!
[1566058772] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 23.7% (5841328 kB) free.
[1566058800] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566058888] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11594668 kB) free.
[1566058920] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1566059085] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.2% (14655072 kB) free.
[1566059127] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 67.8% (33428652 kB) free.
[1566059161] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 88.5% (21679048 kB) free.
[1566059399] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1566059488] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7819952 kB) free!
[1566059520] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566059608] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11602648 kB) free.
[1566059640] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566059999] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1566060208] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7971892 kB) free!
[1566060240] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566060328] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11592672 kB) free.
[1566060360] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1566060822] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.37, 0.55, 19.26
[1566060928] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (8003032 kB) free!
[1566060960] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1566061049] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11603184 kB) free.
[1566061081] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566061648] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8241356 kB) free!
[1566061680] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566061768] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.4% (6581600 kB) free!
[1566061800] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1566061888] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11592756 kB) free.
[1566062400] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566062488] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7110508 kB) free!
[1566062520] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566062608] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11602204 kB) free.
[1566062622] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.26, 0.38, 3.11
[1566062640] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1566063208] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7704220 kB) free!
[1566063240] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1566063328] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11595972 kB) free.
[1566063360] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1566063502] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1566063622] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566063894] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.6% (1878572 kB) free!
[1566063928] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7998420 kB) free!
[1566063961] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1566064015] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.6% (1879588 kB) free!
[1566064048] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6227396 kB) free!
[1566064081] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566064135] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 11.5% (1862480 kB) free!
[1566064168] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11600420 kB) free.
[1566064681] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566064768] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6800276 kB) free!
[1566064801] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566064889] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11595432 kB) free.
[1566064921] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1566065335] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (151960 kB) free!
[1566065488] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7392360 kB) free!
[1566065521] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566065608] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11601204 kB) free.
[1566065641] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1566066208] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7716612 kB) free!
[1566066241] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1566066328] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11591668 kB) free.
[1566066361] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1566066534] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)