[1579113014] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 58.8% (28998368 kB) free.
[1579113585] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1579113759] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18082580 kB) free!
[1579113879] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.6% (46897912 kB) free.
[1579114185] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 352 processes
[1579114305] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1579114480] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24436688 kB) free!
[1579114600] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.5% (20421940 kB) free!
[1579114720] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16717092 kB) free!
[1579114758] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 20.0% (4898728 kB) free!
[1579114881] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 20.0% (4890128 kB) free!
[1579114900] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1579114905] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 356 processes
[1579115001] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.9% (4882076 kB) free!
[1579115025] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1579115110] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171956 kB) free!
[1579115171] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178660 kB) free!
[1579115231] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (180712 kB) free!
[1579115291] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 71.0% (17413216 kB) free.
[1579115320] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27438428 kB) free.
[1579115468] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1579116008] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1579116225] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 353 processes
[1579116345] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1579116465] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1579116520] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18027844 kB) free!
[1579116640] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.4% (46605136 kB) free.
[1579117240] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (23647452 kB) free!
[1579117360] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.0% (19822368 kB) free!
[1579117480] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.3% (46511596 kB) free.
[1579117666] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1579117786] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1579118080] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25299524 kB) free!
[1579118200] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.3% (21514112 kB) free!
[1579118320] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17972980 kB) free!
[1579119520] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35351824 kB) free.
[1579119586] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 355 processes
[1579119706] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1579119826] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 353 processes
[1579120343] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.1% (14622420 kB) free!
[1579120426] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1579120721] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23926960 kB) free!
[1579120841] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.4% (20333260 kB) free!
[1579120961] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.9% (46015604 kB) free.
[1579121561] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25801112 kB) free!
[1579121626] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 355 processes
[1579121681] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.8% (22136008 kB) free!
[1579121746] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1579121801] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.2% (46406800 kB) free.
[1579121866] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1579122368] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (217360 kB) free!
[1579122401] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25829876 kB) free!
[1579122466] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 353 processes
[1579122488] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (242648 kB) free!
[1579122521] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.9% (22221788 kB) free!
[1579122586] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 353 processes
[1579122608] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.2% (42992120 kB) free.
[1579122641] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.2% (46369688 kB) free.
[1579122706] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 355 processes
[1579123241] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26355568 kB) free!
[1579123306] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1579123361] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.2% (22642808 kB) free!
[1579123400] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.7% (10210812 kB) free.
[1579123481] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.1% (46321200 kB) free.
[1579123937] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1850) < 2020/01/16 03:02:11 [error] 21675#0: *25095287 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.59.173, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579124057] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(46) < 2020/01/16 03:02:32 [error] 21677#0: *25096627 no live upstreams while connecting to upstream, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "backend_3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579124081] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25272144 kB) free!
[1579124177] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579124201] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.3% (21475340 kB) free!
[1579124232] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1509) < 49.50.64.122 [16/Jan/2020:03:02:32 +0530] TCP 502 0 0 60.000 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1579124321] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.1% (46233328 kB) free.
[1579124350] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/01/16 03:01:12 [error] 18360#0: *39633278 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.67.237.219, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579124352] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579124470] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579124922] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24526292 kB) free!
[1579125042] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.9% (20938340 kB) free!
[1579125106] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1579125162] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.0% (44798576 kB) free.
[1579125226] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1579125346] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 356 processes
[1579125762] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (23868456 kB) free!
[1579125882] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.1% (46253072 kB) free.
[1579125946] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1579126262] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1579126382] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1579127746] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 355 processes
[1579127866] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 353 processes
[1579127986] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 352 processes
[1579128182] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1579128303] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1579128882] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20833060 kB) free!
[1579129002] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (16907384 kB) free!
[1579129122] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.1% (44994620 kB) free.
[1579129722] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19413244 kB) free!
[1579129842] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (40588220 kB) free.
[1579130442] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21726184 kB) free!
[1579130563] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.5% (17793644 kB) free!
[1579130683] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.7% (45752300 kB) free.
[1579131201] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182184 kB) free!
[1579131283] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (23862100 kB) free!
[1579131321] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185876 kB) free!
[1579131403] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.5% (20487720 kB) free!
[1579131441] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182416 kB) free!
[1579131523] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42805400 kB) free.
[1579132207] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 54.6% (13464184 kB) free.
[1579132723] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19186532 kB) free!
[1579132843] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1579132843] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39210344 kB) free.
[1579133443] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19738968 kB) free!
[1579133563] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39188996 kB) free.
[1579134163] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21621972 kB) free!
[1579134283] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.9% (18367284 kB) free!
[1579134302] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1579134403] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38367044 kB) free.
[1579134422] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1579135003] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19312272 kB) free!
[1579135022] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 407 processes
[1579135123] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38601220 kB) free.
[1579135142] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 407 processes
[1579135343] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (11691660 kB) free!
[1579135724] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19496768 kB) free!
[1579135844] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39597644 kB) free.
[1579135943] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 32.1% (42295064 kB) free.
[1579136246] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.1% (39985044 kB) free.
[1579136285] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.7% (22731152 kB) free.
[1579136323] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.8% (20786936 kB) free.
[1579136444] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20760800 kB) free!
[1579136564] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39944948 kB) free.
[1579136927] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181336 kB) free!
[1579137007] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (532848 kB) free!
[1579137047] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191180 kB) free!
[1579137128] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (526636 kB) free!
[1579137164] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21870972 kB) free!
[1579137172] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193420 kB) free!
[1579137248] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (520332 kB) free!
[1579137284] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.3% (20105360 kB) free!
[1579137404] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (38015964 kB) free.
[1579138004] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19582684 kB) free!
[1579138124] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37954868 kB) free.
[1579138186] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1579138724] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18453704 kB) free!
[1579138844] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38074412 kB) free.
[1579139386] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 354 processes
[1579139444] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19703920 kB) free!
[1579139506] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1579139564] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39135348 kB) free.
[1579140107] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 357 processes
[1579140164] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19235956 kB) free!
[1579140201] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2405368 kB) free!
[1579140227] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1579140284] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38774920 kB) free.
[1579140347] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1579140884] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19501992 kB) free!
[1579141005] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39172448 kB) free.
[1579141547] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1579141604] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18908220 kB) free!
[1579141724] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (38964816 kB) free.
[1579142153] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.3% (21636100 kB) free.
[1579142186] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.5% (22815460 kB) free.
[1579142213] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.9% (22275560 kB) free.
[1579142324] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18224116 kB) free!
[1579142444] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38592244 kB) free.
[1579142747] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 355 processes
[1579142819] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189068 kB) free!
[1579142867] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1579142940] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186600 kB) free!
[1579143044] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19771272 kB) free!
[1579143060] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1579143060] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173640 kB) free!
[1579143164] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38367672 kB) free.
[1579143467] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 356 processes
[1579143588] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1579143764] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19813604 kB) free!
[1579143884] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38188752 kB) free.
[1579144142] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1579144262] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1579144484] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20016064 kB) free!
[1579144604] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37991084 kB) free.
[1579144787] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 357 processes
[1579144907] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1579145204] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20958896 kB) free!
[1579145324] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.0% (17190356 kB) free!
[1579145444] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36719644 kB) free.
[1579146044] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (14017968 kB) free!
[1579146165] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36272016 kB) free.
[1579146764] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14150056 kB) free!
[1579146884] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35947420 kB) free.
[1579147484] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14132568 kB) free!
[1579147604] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35168264 kB) free.
[1579148204] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14350036 kB) free!
[1579148324] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10334352 kB) free!
[1579148444] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32338052 kB) free.
[1579148777] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(24) < 2020/01/16 09:53:43 [error] 21677#0: *25114179 connect() failed (111: Connection refused) while connecting to upstream, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579148897] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579149044] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9805604 kB) free!
[1579149164] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32514528 kB) free.
[1579149764] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9437500 kB) free!
[1579149884] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32198988 kB) free.
[1579150151] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 94.7% (23217752 kB) free.
[1579150484] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9356788 kB) free!
[1579150604] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (5202576 kB) free!
[1579150724] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.8% (24810744 kB) free!
[1579151325] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (375700 kB) free!
[1579151507] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1579151628] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1579151675] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1579151748] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1579151757] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186172 kB) free!
[1579151795] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579151878] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177188 kB) free!
[1579151897] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/16 10:41:57 [error] 21679#0: *25116343 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:2712/106, bytes from/to upstream:106/2712
[1579151916] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1579151925] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14715104 kB) free!
[1579151957] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181436 kB) free!
[1579151998] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184872 kB) free!
[1579152017] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579152077] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194888 kB) free!
[1579152197] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173964 kB) free!
[1579152386] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164820 kB) free!
[1579152446] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (231888 kB) free!
[1579152506] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168200 kB) free!
[1579152566] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (238024 kB) free!
[1579152598] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 82.5% (20218796 kB) free.
[1579152626] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180680 kB) free!
[1579152686] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (213780 kB) free!
[1579153725] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1579153725] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (9238660 kB) free!
[1579154325] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16692152 kB) free!
[1579154459] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.2% (3411208 kB) free!
[1579154477] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.9% (21450464 kB) free.
[1579154533] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.0% (39920928 kB) free.
[1579154569] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.6% (22931268 kB) free.
[1579154580] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.4% (3013812 kB) free!
[1579154700] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.8% (2641956 kB) free!
[1579155088] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191108 kB) free!
[1579155136] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (251360 kB) free!
[1579155207] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.6% (22106204 kB) free.
[1579155256] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (250192 kB) free!
[1579155300] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (565028 kB) free!
[1579155376] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (239512 kB) free!
[1579155525] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (780636 kB) free!
[1579155614] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8036116 kB) free!
[1579155734] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7400392 kB) free!
[1579155777] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (1513804 kB) free!
[1579155818] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183852 kB) free!
[1579155854] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 13.8% (6800696 kB) free!
[1579155898] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (1365300 kB) free!
[1579155938] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187216 kB) free!
[1579155948] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1579156018] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (1224604 kB) free!
[1579156058] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173268 kB) free!
[1579156125] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17087924 kB) free!
[1579156379] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 12.9% (3181876 kB) free!
[1579156455] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 62.3% (30684932 kB) free.
[1579156499] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 11.1% (2730300 kB) free!
[1579156500] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 74.8% (16786496 kB) free.
[1579156619] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2271204 kB) free!
[1579156817] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/16 12:06:59 [error] 21683#0: *25120273 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:5894/106, bytes from/to upstream:106/5894
[1579156938] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579157220] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.4% (23046360 kB) free.
[1579157325] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2281576 kB) free!
[1579157537] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(8) < 2020/01/16 12:18:51 [error] 21674#0: *25120815 connect() failed (111: Connection refused) while connecting to upstream, client: 92.42.104.166, server: 0.0.0.0:13001, upstream: "10.147.212.143:13001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579157657] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579157925] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (14977388 kB) free!
[1579158257] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/16 12:25:18 [error] 21683#0: *25121143 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:17057/140, bytes from/to upstream:140/17057
[1579158377] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579158525] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27497604 kB) free.
[1579159125] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2875784 kB) free!
[1579159246] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (390796 kB) free!
[1579159366] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30066620 kB) free.
[1579159808] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9567544 kB) free!
[1579159928] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.9% (9311596 kB) free!
[1579159966] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4913228 kB) free!
[1579160048] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.3% (9030860 kB) free!
[1579160086] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (794784 kB) free!
[1579160206] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33408116 kB) free.
[1579160806] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7906692 kB) free!
[1579160818] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 15.5% (3795088 kB) free!
[1579160890] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (4635116 kB) free!
[1579160926] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3523800 kB) free!
[1579161010] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (4483880 kB) free!
[1579161046] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383372 kB) free!
[1579161130] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.5% (4299772 kB) free!
[1579161348] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1579161418] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174004 kB) free!
[1579161468] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1579161588] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1579161848] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4546876 kB) free!
[1579162246] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26139852 kB) free!
[1579162846] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1600740 kB) free!
[1579162930] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1635208 kB) free!
[1579163011] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.0% (22198216 kB) free.
[1579163040] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579163050] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 53.1% (70033568 kB) free.
[1579163084] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.4% (40589032 kB) free.
[1579163120] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.2% (45419796 kB) free.
[1579163149] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.2% (22592148 kB) free.
[1579163198] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 92.7% (22719956 kB) free.
[1579163640] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579163760] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579163880] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1579164407] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188912 kB) free!
[1579164527] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165596 kB) free!
[1579164598] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.6% (4803196 kB) free!
[1579164647] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166444 kB) free!
[1579164718] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.1% (4681340 kB) free!
[1579164759] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.3% (22864972 kB) free.
[1579164817] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167440 kB) free!
[1579164838] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.6% (4564536 kB) free!
[1579164865] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.6% (21951912 kB) free.
[1579164874] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190084 kB) free!
[1579164880] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182412 kB) free!
[1579164900] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.4% (3903888 kB) free!
[1579164934] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.3% (22038576 kB) free.
[1579164965] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 75.1% (16863312 kB) free.
[1579165080] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579165188] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1579165458] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23965304 kB) free!
[1579165578] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.7% (19410640 kB) free!
[1579165698] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15141380 kB) free!
[1579166298] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27714960 kB) free.
[1579166655] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8644224 kB) free!
[1579166775] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.9% (8314784 kB) free!
[1579166895] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.1% (7953784 kB) free!
[1579167420] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 18.9% (4675836 kB) free!
[1579167496] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1579167496] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.0% (30049700 kB) free.
[1579167498] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16496776 kB) free!
[1579167540] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4433200 kB) free!
[1579167588] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579167618] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12283396 kB) free!
[1579167660] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 17.1% (4228452 kB) free!
[1579167708] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1579167738] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.7% (41794248 kB) free.
[1579167876] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1579167997] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 331 processes
[1579168117] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1579168260] SERVICE FLAPPING ALERT: db1-vm2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1579168260] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.5% (23073312 kB) free.
[1579168338] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8203292 kB) free!
[1579168458] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3570108 kB) free!
[1579168578] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389776 kB) free!
[1579168881] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (230596 kB) free!
[1579169001] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (291336 kB) free!
[1579169121] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (286104 kB) free!
[1579169136] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (1449216 kB) free!
[1579169256] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (1296536 kB) free!
[1579169376] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (1163928 kB) free!
[1579169462] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1579169583] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1579169779] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.9% (22290624 kB) free!
[1579170273] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 753 processes
[1579170377] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/16 15:50:12 [error] 21681#0: *25130840 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:10243/104, bytes from/to upstream:104/10243
[1579170379] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399536 kB) free!
[1579170393] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 755 processes
[1579170497] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579170498] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9295536 kB) free!
[1579170513] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 756 processes
[1579170618] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (9215568 kB) free!
[1579170738] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (9051828 kB) free!
[1579170934] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1579170934] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.9% (21933852 kB) free.
[1579170965] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36758896 kB) free.
[1579171023] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.4% (14368936 kB) free.
[1579171486] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.1% (8410816 kB) free!
[1579171606] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.6% (8680716 kB) free!
[1579171698] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(659) < 2020/01/16 16:16:53 [error] 21678#0: *25133381 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579171726] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.6% (8682196 kB) free!
[1579171752] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(648) < 103.255.146.164 [16/Jan/2020:16:16:53 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1579171818] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579171825] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 693 processes
[1579171873] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579171945] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 360 processes
[1579172168] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (25113396 kB) free!
[1579172288] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.0% (19718720 kB) free!
[1579172409] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.3% (14926780 kB) free!
[1579173009] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38223444 kB) free.
[1579173357] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (4682120 kB) free!
[1579173478] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.7% (4575032 kB) free!
[1579173598] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.0% (4421192 kB) free!
[1579173609] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11664256 kB) free!
[1579173708] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1579173729] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (7159616 kB) free!
[1579173849] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.2% (43790764 kB) free.
[1579174117] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1579174237] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1579174449] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15494140 kB) free!
[1579174569] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (10707956 kB) free!
[1579174689] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (6317056 kB) free!
[1579174798] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165372 kB) free!
[1579174938] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4447800 kB) free!
[1579175289] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.2% (21350272 kB) free!
[1579175616] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1579175624] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (2897376 kB) free!
[1579175744] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (3133616 kB) free!
[1579175864] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3470836 kB) free.
[1579175889] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (41022020 kB) free.
[1579175926] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3793952 kB) free!
[1579176018] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/01/16 17:27:02 [error] 21680#0: *25136683 connect() failed (111: Connection refused) while connecting to upstream, client: 185.255.8.58, server: 0.0.0.0:13001, upstream: "10.147.212.155:13001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579176138] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579176337] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.3% (3520340 kB) free!
[1579176457] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.7% (3390152 kB) free!
[1579176489] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13690228 kB) free!
[1579176578] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.0% (3196312 kB) free!
[1579176609] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8530332 kB) free!
[1579176729] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (4166448 kB) free!
[1579176738] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/16 17:32:42 [error] 21675#0: *25136973 recv() failed (104: Connection reset by peer) while proxying connection, client: 125.19.22.226, server: 0.0.0.0:13001, upstream: "10.147.212.142:13001", bytes from/to client:39/0, bytes from/to upstream:0/39
[1579176858] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579176968] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.7% (3523724 kB) free!
[1579177088] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.5% (3247408 kB) free!
[1579177178] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2434236 kB) free!
[1579177208] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.5% (3034288 kB) free!
[1579177329] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.9% (20977456 kB) free!
[1579177808] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 75.0% (16839196 kB) free.
[1579177930] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390260 kB) free!
[1579178413] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.4% (20583788 kB) free.
[1579178443] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 29.2% (14385112 kB) free.
[1579178509] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28506768 kB) free.
[1579178545] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 54.3% (26752048 kB) free.
[1579178574] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 64.3% (15756216 kB) free.
[1579178897] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1579179616] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (1880216 kB) free!
[1579179660] SERVICE FLAPPING ALERT: db1-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1579179718] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24325948 kB) free!
[1579179736] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190428 kB) free!
[1579179838] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.1% (19882108 kB) free!
[1579179856] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172864 kB) free!
[1579179958] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15735340 kB) free!
[1579180664] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.4% (2501056 kB) free!
[1579180784] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.1% (2614644 kB) free!
[1579180904] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.2% (4427680 kB) free.
[1579181127] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9769784 kB) free!
[1579181158] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40880608 kB) free.
[1579181247] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.5% (9626280 kB) free!
[1579181367] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.3% (9506100 kB) free!
[1579181667] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.6% (4809340 kB) free!
[1579181758] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (14011512 kB) free!
[1579181787] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.4% (4745012 kB) free!
[1579181878] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (9445264 kB) free!
[1579181907] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.1% (4679260 kB) free!
[1579181998] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (5191404 kB) free!
[1579182598] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19828088 kB) free!
[1579183198] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39834088 kB) free.
[1579183347] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.0% (8884516 kB) free!
[1579183467] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8884768 kB) free!
[1579183587] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.0% (8874984 kB) free!
[1579183798] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11910252 kB) free!
[1579183918] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7297356 kB) free!
[1579184038] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (3067476 kB) free!
[1579184638] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20526568 kB) free!
[1579184658] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/16 19:46:54 [error] 21678#0: *25143166 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:4229/87, bytes from/to upstream:87/4229
[1579184778] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579185216] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (4459808 kB) free!
[1579185239] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40797992 kB) free.
[1579185336] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4388624 kB) free!
[1579185378] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/16 20:04:20 [error] 21679#0: *25143992 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:7258/70, bytes from/to upstream:70/7258
[1579185456] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.6% (4320268 kB) free!
[1579185498] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579185508] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2404624 kB) free!
[1579185704] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.0% (2769348 kB) free!
[1579185824] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.5% (4960492 kB) free.
[1579185839] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10045732 kB) free!
[1579185959] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (5297880 kB) free!
[1579186079] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (1008316 kB) free!
[1579186167] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4106296 kB) free!
[1579186679] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16926392 kB) free!
[1579187879] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392764 kB) free!
[1579188479] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14188236 kB) free!
[1579189079] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (441536 kB) free!
[1579189679] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (20069116 kB) free!
[1579190056] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1579190187] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4721072 kB) free!
[1579190279] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38878472 kB) free.
[1579190396] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.6% (21962848 kB) free.
[1579190436] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 67.6% (16683140 kB) free.
[1579190458] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 64.5% (31796988 kB) free.
[1579190491] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.9% (8817868 kB) free!
[1579190493] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1579190493] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 27.4% (13504868 kB) free.
[1579190879] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13292968 kB) free!
[1579190999] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9076912 kB) free!
[1579191036] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188448 kB) free!
[1579191119] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (5191236 kB) free!
[1579191157] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179216 kB) free!
[1579191277] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193716 kB) free!
[1579191719] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19606412 kB) free!
[1579192319] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38331480 kB) free.
[1579192920] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10991192 kB) free!
[1579193040] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (6690496 kB) free!
[1579193160] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2819716 kB) free!
[1579193256] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2415492 kB) free!
[1579193760] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17324424 kB) free!
[1579194360] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36766204 kB) free.
[1579194960] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9233652 kB) free!
[1579195080] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (5109972 kB) free!
[1579195200] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (1300088 kB) free!
[1579195800] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (16011888 kB) free!
[1579196400] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31937164 kB) free.
[1579197000] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8861848 kB) free!
[1579197120] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (5186924 kB) free!
[1579197240] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37842536 kB) free.
[1579197840] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13527568 kB) free!
[1579197960] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (9495056 kB) free!
[1579198080] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (5831436 kB) free!
[1579198096] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (245744 kB) free!
[1579198216] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (241220 kB) free!
[1579198336] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (245328 kB) free!
[1579198343] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.4% (25614064 kB) free!
[1579198464] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.4% (17601964 kB) free!
[1579198584] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (6074064 kB) free!
[1579198681] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17960924 kB) free!
[1579199281] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31275592 kB) free.