[1586629848] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19203280 kB) free!
[1586630449] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38359340 kB) free.
[1586631049] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12604120 kB) free!
[1586631169] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8502628 kB) free!
[1586631289] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (4779388 kB) free!
[1586631812] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 10.3% (2526844 kB) free!
[1586631872] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (2421464 kB) free!
[1586631889] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.1% (18630928 kB) free!
[1586631932] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 9.4% (2309776 kB) free!
[1586631992] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 8.9% (2183936 kB) free!
[1586632263] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (1551496 kB) free!
[1586632383] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (746804 kB) free!
[1586632489] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34168140 kB) free.
[1586632503] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (1370664 kB) free!
[1586633089] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9861956 kB) free!
[1586633192] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 77.2% (18929324 kB) free.
[1586633209] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6036924 kB) free!
[1586633329] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36201068 kB) free.
[1586633929] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15060360 kB) free!
[1586634049] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (10921496 kB) free!
[1586634169] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (6969828 kB) free!
[1586634769] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19893476 kB) free!
[1586635369] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35232468 kB) free.
[1586635969] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (11992924 kB) free!
[1586636089] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (8104152 kB) free!
[1586636121] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1586636210] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.9% (24899936 kB) free!
[1586637410] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28497132 kB) free.
[1586638010] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8153736 kB) free!
[1586638130] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37752692 kB) free.
[1586638730] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15811692 kB) free!
[1586638759] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.20, 3.78, 4.83
[1586638850] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (12065364 kB) free!
[1586638970] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (8415484 kB) free!
[1586639514] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 69.3% (34164604 kB) free.
[1586639570] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.7% (19369276 kB) free!
[1586639703] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.9% (2743848 kB) free!
[1586640170] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31239388 kB) free.
[1586640171] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.1% (43895016 kB) free.
[1586640589] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 38.7% (9489448 kB) free.
[1586640672] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(303) < 2020/04/12 03:00:31 [error] 22122#0: *83281327 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1586640770] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9887256 kB) free!
[1586640792] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2) < 2020/04/12 03:01:32 [error] 22125#0: *83281373 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1586640890] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6401304 kB) free!
[1586640912] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586640916] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2578) < 172.105.59.173 [12/Apr/2020:03:02:32 +0530] TCP 502 0 0 120.001 "10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1586641010] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35042848 kB) free.
[1586641036] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586641610] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12472532 kB) free!
[1586641730] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8845932 kB) free!
[1586641851] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37518872 kB) free.
[1586641955] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1586642075] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1586642450] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7458040 kB) free!
[1586642570] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3404092 kB) free!
[1586642690] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405656 kB) free!
[1586643890] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23173048 kB) free!
[1586644490] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (433432 kB) free!
[1586645090] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14163008 kB) free!
[1586645690] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30557176 kB) free.
[1586645703] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1586646559] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.53, 5.60, 5.13
[1586646679] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.36, 5.55, 5.17
[1586646799] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.15, 5.46, 5.19
[1586646890] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19576656 kB) free!
[1586647010] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (15715476 kB) free!
[1586647130] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39476844 kB) free.
[1586647399] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.24, 3.35, 4.34
[1586647503] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.5% (3661804 kB) free.
[1586647730] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20324024 kB) free!
[1586647850] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.6% (16613808 kB) free!
[1586647971] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39408056 kB) free.
[1586648389] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (238248 kB) free!
[1586648509] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181348 kB) free!
[1586648570] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20971048 kB) free!
[1586648629] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199096 kB) free!
[1586648690] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.2% (17410696 kB) free!
[1586648810] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41517316 kB) free.
[1586649253] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 72.7% (17811984 kB) free.
[1586649342] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 42.4% (10470536 kB) free.
[1586649410] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (23552308 kB) free!
[1586649530] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.1% (19955624 kB) free!
[1586649590] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 71.7% (17563668 kB) free.
[1586649650] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41456864 kB) free.
[1586649942] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.7% (3389696 kB) free!
[1586650062] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (891668 kB) free!
[1586650182] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194464 kB) free!
[1586650250] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (22695192 kB) free!
[1586650370] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (18909696 kB) free!
[1586650490] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41401216 kB) free.
[1586651090] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21839956 kB) free!
[1586651210] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.6% (17918004 kB) free!
[1586651330] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.3% (41210952 kB) free.
[1586651599] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.45, 6.31, 5.29
[1586651719] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.93, 6.24, 5.40
[1586651839] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.44, 5.95, 5.40
[1586651930] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19780512 kB) free!
[1586652050] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.0% (15822080 kB) free!
[1586652170] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39957212 kB) free.
[1586652771] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20294500 kB) free!
[1586652891] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.7% (16741356 kB) free!
[1586653011] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38927728 kB) free.
[1586653039] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1586653039] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 1.83, 2.84, 4.18
[1586653611] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19720264 kB) free!
[1586653731] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40772156 kB) free.
[1586653892] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 17.1% (4184960 kB) free!
[1586653952] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 17.1% (4191280 kB) free!
[1586654012] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 17.1% (4190516 kB) free!
[1586654072] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 17.1% (4189144 kB) free!
[1586654331] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22537848 kB) free!
[1586654451] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (18895148 kB) free!
[1586654571] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (40644076 kB) free.
[1586655171] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (21218564 kB) free!
[1586655291] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.9% (16977324 kB) free!
[1586655411] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40469956 kB) free.
[1586656011] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19350544 kB) free!
[1586656131] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (15655096 kB) free!
[1586656251] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38581408 kB) free.
[1586656851] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17917952 kB) free!
[1586656971] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40058692 kB) free.
[1586657572] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20185624 kB) free!
[1586657692] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.3% (16262732 kB) free!
[1586657703] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3200752 kB) free!
[1586657812] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39156888 kB) free.
[1586657824] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3383820 kB) free.
[1586658412] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18533156 kB) free!
[1586658532] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36419208 kB) free.
[1586659132] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18760416 kB) free!
[1586659252] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.2% (14781924 kB) free!
[1586659372] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39017892 kB) free.
[1586659972] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18478224 kB) free!
[1586660092] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.0% (14519912 kB) free!
[1586660212] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37796444 kB) free.
[1586660223] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3222184 kB) free!
[1586660343] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3382620 kB) free.
[1586660812] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (17041080 kB) free!
[1586660932] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1586660932] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31615660 kB) free.
[1586661439] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.32, 5.60, 5.05
[1586661532] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19128016 kB) free!
[1586661543] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3242512 kB) free!
[1586661559] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.09, 5.51, 5.09
[1586661652] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.4% (15082688 kB) free!
[1586661663] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3362184 kB) free.
[1586661679] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.04, 5.34, 5.07
[1586661772] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38060016 kB) free.
[1586662279] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.98, 4.29, 4.70
[1586662373] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16141460 kB) free!
[1586662493] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11807504 kB) free!
[1586662613] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36338820 kB) free.
[1586662879] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.12, 5.35, 5.06
[1586662999] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.44, 5.28, 5.07
[1586663119] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.66, 5.07, 5.02
[1586663213] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14172664 kB) free!
[1586663333] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26917020 kB) free.
[1586663463] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3127576 kB) free!
[1586663583] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 20.0% (3252240 kB) free!
[1586663704] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (3005208 kB) free!
[1586663719] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.38, 4.39, 4.75
[1586663933] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15184992 kB) free!
[1586664053] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (10709376 kB) free!
[1586664173] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35794040 kB) free.
[1586664319] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.69, 7.03, 5.45
[1586664441] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.78, 8.49, 6.18
[1586664561] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.65, 7.57, 6.12
[1586664773] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12746544 kB) free!
[1586664893] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (8138812 kB) free!
[1586665013] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (3652000 kB) free!
[1586666213] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.9% (22309900 kB) free!
[1586666813] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31241852 kB) free.
[1586667413] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8329888 kB) free!
[1586667533] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (3936924 kB) free!
[1586667654] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1586667654] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17562640 kB) free!
[1586667980] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1586668100] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1586668190] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4903140 kB) free!
[1586668220] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1586668231] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1586668254] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (8202188 kB) free!
[1586668310] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4720572 kB) free!
[1586668351] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1586668430] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (4561020 kB) free!
[1586668453] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183572 kB) free!
[1586668471] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1586668573] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185916 kB) free!
[1586668693] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191580 kB) free!
[1586668820] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1586668854] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.1% (18582600 kB) free!
[1586669071] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1586669361] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 36.90, 12.35, 9.36
[1586669454] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29081152 kB) free.
[1586669630] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 21.7% (5328532 kB) free.
[1586669961] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.13, 8.60, 8.81
[1586670020] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1586670054] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4338576 kB) free!
[1586670140] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1586670174] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (377516 kB) free!
[1586670230] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (4303212 kB) free!
[1586670260] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1586670294] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402688 kB) free!
[1586670312] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 62.9% (15427892 kB) free.
[1586670350] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.0% (4165324 kB) free!
[1586670470] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.4% (4025488 kB) free!
[1586670724] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.0% (3360336 kB) free!
[1586670844] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.6% (3044112 kB) free!
[1586670964] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.6% (2608300 kB) free!
[1586671494] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19909680 kB) free!
[1586671512] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/04/12 11:30:36 [error] 22135#0: *83373293 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1586671564] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (946104 kB) free!
[1586671633] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586671636] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1468) < 172.105.59.173 [12/Apr/2020:11:36:17 +0530] TCP 502 0 0 180.002 "10.147.212.155:3001, 10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.142:3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 60.001, 60.000, 60.001"
[1586671757] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586672094] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27292456 kB) free.
[1586672270] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1586672270] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184884 kB) free!
[1586672694] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (396676 kB) free!
[1586672814] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (376052 kB) free!
[1586672934] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402224 kB) free!
[1586673364] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.9% (16142864 kB) free.
[1586673534] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16698884 kB) free!
[1586674135] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (40628360 kB) free.
[1586674156] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(657) < 49.50.86.152 [12/Apr/2020:12:19:16 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1586674276] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(169) < 203.212.70.170 [12/Apr/2020:12:20:40 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1586674397] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(3) < 172.105.59.173 [12/Apr/2020:12:22:21 +0530] TCP 502 0 0 180.000 "10.147.212.143:3001, 10.147.212.142:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 0.000"
[1586674460] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1586674460] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1586674735] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14250260 kB) free!
[1586674855] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9754604 kB) free!
[1586674975] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5402436 kB) free!
[1586674997] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586675232] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/04/12 12:34:43 [error] 22135#0: *83385050 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/dashboard HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/dashboard", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/dashboard"
[1586675353] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586675575] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17773128 kB) free!
[1586675704] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1533432 kB) free!
[1586675893] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 33.9% (8302408 kB) free.
[1586676175] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34119024 kB) free.
[1586676304] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.8% (2074344 kB) free!
[1586676714] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8935404 kB) free!
[1586676775] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6879424 kB) free!
[1586676834] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (8679204 kB) free!
[1586676895] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2410640 kB) free!
[1586676904] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (524436 kB) free!
[1586676912] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.2% (3980532 kB) free!
[1586676954] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (8419508 kB) free!
[1586677015] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401744 kB) free!
[1586677032] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.3% (3749980 kB) free!
[1586677152] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.5% (3550664 kB) free!
[1586677615] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14780052 kB) free!
[1586678215] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35873528 kB) free.
[1586678352] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1589124 kB) free!
[1586678704] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1586678704] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.0% (2111644 kB) free!
[1586678815] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8671824 kB) free!
[1586678893] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4888492 kB) free!
[1586678935] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4092040 kB) free!
[1586679013] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (4768764 kB) free!
[1586679055] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392152 kB) free!
[1586679133] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (4640320 kB) free!
[1586679171] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.2% (8461296 kB) free!
[1586679291] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.5% (8144672 kB) free!
[1586679304] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3370856 kB) free.
[1586679411] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.9% (7823920 kB) free!
[1586679655] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19679700 kB) free!
[1586679904] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.3% (2322928 kB) free!
[1586679954] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4114224 kB) free!
[1586679964] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.4% (3004592 kB) free!
[1586680024] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.0% (1957660 kB) free!
[1586680084] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.6% (2601456 kB) free!
[1586680145] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3312764 kB) free.
[1586680204] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2231936 kB) free!
[1586680256] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390160 kB) free!
[1586680611] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4880812 kB) free!
[1586680744] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.4% (2018528 kB) free!
[1586680856] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14744732 kB) free!
[1586680864] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3400812 kB) free.
[1586681361] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1586681456] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381676 kB) free!
[1586681464] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3222304 kB) free!
[1586681584] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (3168644 kB) free!
[1586681597] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(175) < 92.42.104.166 [12/Apr/2020:14:20:16 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1586681704] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (3135604 kB) free!
[1586681718] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586682133] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2280564 kB) free!
[1586682260] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1586682470] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1586682656] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 20.0% (26386200 kB) free!
[1586682904] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1410676 kB) free!
[1586683256] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383916 kB) free!
[1586683504] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.0% (2431076 kB) free!
[1586683761] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 16.5% (5796128 kB) free!
[1586683804] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 32.3% (7253980 kB) free.
[1586683882] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.9% (31834556 kB) free.
[1586684117] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(158) < 92.42.104.166 [12/Apr/2020:15:00:13 +0530] TCP 502 0 0 0.002 "10.147.212.143:13001, 10.147.212.154:13001, 10.147.212.142:13001, 10.147.212.155:13001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 0.000, 0.000, 0.001"
[1586684237] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586684660] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 352 processes
[1586684780] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 359 processes
[1586684901] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 4 crit, 0 warn out of 359 processes
[1586685056] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.8% (23516744 kB) free!
[1586685185] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1586685305] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1586685425] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1586685555] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1586685656] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387272 kB) free!
[1586685675] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1586686101] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1586686152] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/04/12 15:31:39 [error] 22135#0: *83417516 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1586686256] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28345236 kB) free.
[1586686272] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586686504] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (785372 kB) free!
[1586686701] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1586686821] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 360 processes
[1586686856] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (400052 kB) free!
[1586686941] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 4 crit, 0 warn out of 360 processes
[1586686976] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (391188 kB) free!
[1586687096] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385240 kB) free!
[1586687104] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.2% (1659996 kB) free!
[1586687541] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1586687541] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1586687697] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17655120 kB) free!
[1586688275] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.8% (20824292 kB) free!
[1586688297] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (404212 kB) free!
[1586688304] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (739372 kB) free!
[1586688395] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.1% (30417536 kB) free.
[1586688425] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1586688741] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1586688861] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1586688897] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13576764 kB) free!
[1586688904] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.2% (2305912 kB) free!
[1586688981] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1586689037] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(561) < 203.212.70.170 [12/Apr/2020:16:20:34 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1586689157] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586690097] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (11149084 kB) free!
[1586690181] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1586690697] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36851788 kB) free.
[1586691297] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6386448 kB) free!
[1586691304] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (274852 kB) free!
[1586691381] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 343 processes
[1586691417] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1828028 kB) free!
[1586691502] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 343 processes
[1586691537] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401504 kB) free!
[1586691605] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.2% (4094452 kB) free!
[1586691622] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 350 processes
[1586691725] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.9% (3572844 kB) free!
[1586691845] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.0% (3135268 kB) free!
[1586691904] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.8% (2576084 kB) free!
[1586692137] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.4% (24315816 kB) free!
[1586692222] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1586692445] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2026052 kB) free!
[1586692737] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (454936 kB) free!
[1586693104] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (947576 kB) free!
[1586693337] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16176728 kB) free!
[1586693705] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.6% (2697428 kB) free!
[1586693795] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.5% (20418580 kB) free!
[1586693915] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.6% (32464376 kB) free.
[1586693937] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403452 kB) free!
[1586694305] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3383920 kB) free.
[1586694537] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (14002276 kB) free!
[1586694845] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1586694845] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 73.0% (16395332 kB) free.
[1586694905] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3045156 kB) free!
[1586695025] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (263308 kB) free!
[1586695137] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400936 kB) free!
[1586695145] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (982324 kB) free!
[1586696315] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.8% (18246264 kB) free!
[1586696436] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.4% (32148784 kB) free.
[1586696937] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16652420 kB) free!
[1586697035] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 10.4% (13644928 kB) free!
[1586697155] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.5% (25695876 kB) free!
[1586697275] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.3% (14960936 kB) free!
[1586697537] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (375840 kB) free!
[1586697622] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1586697742] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 342 processes
[1586697862] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 342 processes
[1586697875] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.2% (29207744 kB) free.
[1586698137] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13495400 kB) free!
[1586698145] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (2585832 kB) free!
[1586698462] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1586698475] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12314972 kB) free!
[1586698595] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.5% (29606064 kB) free.
[1586698737] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (374908 kB) free!
[1586699195] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11276288 kB) free!
[1586699315] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.3% (28046928 kB) free.
[1586699337] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29269232 kB) free.
[1586699345] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3326984 kB) free.
[1586699915] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9509172 kB) free!
[1586699937] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (396892 kB) free!
[1586700035] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 16.9% (22238352 kB) free!
[1586700057] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (408616 kB) free!
[1586700155] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (12194392 kB) free!
[1586700177] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (375808 kB) free!
[1586700755] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1586700755] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.8% (28737280 kB) free.
[1586700777] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19743540 kB) free!
[1586701145] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (587724 kB) free!
[1586701265] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.0% (2119104 kB) free!
[1586701355] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11841756 kB) free!
[1586701377] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (413752 kB) free!
[1586701386] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3373004 kB) free.
[1586701445] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4296908 kB) free!
[1586701475] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.1% (26554412 kB) free.
[1586701565] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.9% (3789936 kB) free!
[1586701685] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.9% (3338448 kB) free!
[1586701977] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30825480 kB) free.
[1586702075] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9138788 kB) free!
[1586702195] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.4% (25527776 kB) free!
[1586702285] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (674904 kB) free!
[1586702315] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.3% (14931248 kB) free!
[1586702577] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (453180 kB) free!
[1586702697] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (470704 kB) free!
[1586702818] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400340 kB) free!
[1586703185] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (162412 kB) free!
[1586703305] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (1545804 kB) free!
[1586703418] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.8% (24758764 kB) free!
[1586703425] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.9% (3237964 kB) free!
[1586704018] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (421208 kB) free!
[1586704025] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.1% (3914484 kB) free.
[1586704625] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (1370800 kB) free!
[1586704716] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (7823208 kB) free!
[1586704745] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (3239548 kB) free!
[1586704865] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.4% (3960228 kB) free.
[1586704881] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 12.5% (4365408 kB) free!
[1586705001] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 10.9% (3824324 kB) free!
[1586705121] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (3378560 kB) free!
[1586705316] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.7% (16765200 kB) free!
[1586705465] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (999592 kB) free!
[1586705585] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (2969832 kB) free!
[1586705706] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.2% (4097496 kB) free.
[1586705721] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.9% (31834316 kB) free.
[1586705885] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 73.7% (16541404 kB) free.
[1586705916] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (515772 kB) free!
[1586706305] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (1356276 kB) free!
[1586706418] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13391612 kB) free!
[1586706425] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (793032 kB) free!
[1586706516] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.6% (20588308 kB) free!
[1586706545] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (440252 kB) free!
[1586707018] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391864 kB) free!
[1586707116] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.8% (31410648 kB) free.
[1586707145] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (2914908 kB) free!
[1586707177] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.8% (4864980 kB) free!
[1586707297] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.9% (4866208 kB) free!
[1586707417] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.9% (4865788 kB) free!
[1586707618] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30146944 kB) free.
[1586707716] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.3% (16224424 kB) free!
[1586707745] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.1% (3917372 kB) free.
[1586707836] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.1% (33054012 kB) free.
[1586708218] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (2710916 kB) free!
[1586708338] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (442384 kB) free!
[1586708345] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.9% (1937932 kB) free!
[1586708436] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.8% (18151792 kB) free!
[1586708458] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391496 kB) free!
[1586708465] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.8% (4187168 kB) free.
[1586708556] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.1% (31788684 kB) free.
[1586708721] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1656 processes
[1586709058] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.9% (23540900 kB) free!
[1586709065] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (920924 kB) free!
[1586709156] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.9% (19663216 kB) free!
[1586709185] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.8% (2574012 kB) free!
[1586709276] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.1% (31785272 kB) free.
[1586709305] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.3% (4106312 kB) free.
[1586709322] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1852 processes
[1586709658] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396392 kB) free!
[1586709862] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1586709876] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.9% (19615892 kB) free!
[1586709905] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (481000 kB) free!
[1586709996] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.2% (30574136 kB) free.
[1586710026] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.9% (2261092 kB) free!
[1586710146] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.8% (4024960 kB) free.
[1586710258] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27590040 kB) free.
[1586710746] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (424472 kB) free!
[1586710859] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (3818172 kB) free!
[1586710866] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.9% (1933968 kB) free!
[1586710979] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (388812 kB) free!
[1586710986] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.3% (3460112 kB) free.
[1586711099] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391896 kB) free!
[1586712186] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (955464 kB) free!
[1586712306] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.4% (2337552 kB) free!
[1586712426] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.3% (3780424 kB) free.
[1586714099] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19687528 kB) free!
[1586714285] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1586714699] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (413184 kB) free!
[1586715299] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.0% (22368992 kB) free!
[1586715899] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (377024 kB) free!