[1549909879] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.4% (12026264 kB) free.
[1549913455] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549913479] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9323020 kB) free!
[1549913575] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549913599] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.7% (12658896 kB) free.
[1549913695] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549919696] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549921464] SERVICE ALERT: sig4;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 303 processes
[1549921584] SERVICE ALERT: sig4;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1549921999] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9054832 kB) free!
[1549922119] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.9% (12771712 kB) free.
[1549923319] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9402668 kB) free!
[1549923439] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.0% (12820916 kB) free.
[1549924496] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549924616] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549924736] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549925336] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1549930736] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549930856] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549930976] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549931577] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1549931954] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 320 processes
[1549931955] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1549932074] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 320 processes
[1549932075] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1549932111] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.6% (3021136 kB) free!
[1549932194] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 319 processes
[1549932195] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1549932231] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.8% (2892852 kB) free!
[1549932351] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 17.1% (2781192 kB) free!
[1549932794] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 323 processes
[1549932951] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (601232 kB) free!
[1549933551] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 34.0% (5527316 kB) free.
[1549934752] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (981532 kB) free!
[1549934872] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (907204 kB) free!
[1549934992] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1549934992] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (151636 kB) free!
[1549935592] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 27.6% (4482592 kB) free.
[1549938195] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 14 users currently logged in
[1549938777] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549938795] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1549938897] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549939017] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549939039] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9864040 kB) free!
[1549939159] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.9% (12291008 kB) free.
[1549940218] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1549942618] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549942738] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549942858] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1549943458] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549943579] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549943699] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1549943957] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.59, 5.45, 5.12
[1549944077] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.60, 4.22, 4.71
[1549944299] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549944419] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1549944559] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9495840 kB) free!
[1549944679] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.5% (12063292 kB) free.
[1549945019] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549945139] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549945259] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549945394] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1549945514] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 320 processes
[1549945792] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (3051896 kB) free!
[1549945879] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8832208 kB) free!
[1549945912] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.8% (3049960 kB) free!
[1549945999] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.1% (11408408 kB) free.
[1549946032] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (3052084 kB) free!
[1549946355] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 93.0% (15122976 kB) free.
[1549947199] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9299216 kB) free!
[1549947319] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.5% (11557280 kB) free.
[1549948164] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47010 MB (20% inode=99%):
[1549948284] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46858 MB (20% inode=99%):
[1549948404] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46748 MB (20% inode=99%):
[1549948519] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.3% (7043712 kB) free!
[1549948537] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.4% (4031616 kB) free!
[1549948639] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9198248 kB) free!
[1549948657] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.9% (4131296 kB) free!
[1549948759] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.3% (9980852 kB) free.
[1549948777] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.8% (4117564 kB) free!
[1549949359] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.5% (6158236 kB) free!
[1549949377] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198104 kB) free!
[1549949479] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 17.6% (8666896 kB) free!
[1549949599] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.1% (5476516 kB) free!
[1549950659] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1549951259] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1549951380] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1549951979] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549951999] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4436560 kB) free!
[1549952099] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549952219] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1549952599] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.9% (5873152 kB) free!
[1549952613] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 52 users currently logged in
[1549952673] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 52 users currently logged in
[1549952733] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 52 users currently logged in
[1549952793] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 52 users currently logged in
[1549952819] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549952926] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3169816 kB) free!
[1549952939] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549953046] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 21.5% (3498836 kB) free.
[1549953059] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1549953646] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3244584 kB) free!
[1549953659] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549953684] SERVICE ALERT: webservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 16 users currently logged in
[1549953766] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.7% (3203268 kB) free!
[1549953780] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549953800] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1549953800] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (3064192 kB) free!
[1549953827] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9465924 kB) free!
[1549953886] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.5% (3163792 kB) free!
[1549953900] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549953947] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12163588 kB) free.
[1549955000] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.3% (6048184 kB) free!
[1549955484] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 15 users currently logged in
[1549955600] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (2869668 kB) free!
[1549956200] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.8% (5314428 kB) free!
[1549956563] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1549956800] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (1564236 kB) free!
[1549956896] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.7% (22960480 kB) free.
[1549957904] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 69.1% (34075376 kB) free.
[1549957938] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 91.2% (14823860 kB) free.
[1549958747] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9537596 kB) free!
[1549958867] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11976560 kB) free.
[1549959296] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (4480772 kB) free!
[1549959416] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4414780 kB) free!
[1549959467] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9307192 kB) free!
[1549959536] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.7% (4348288 kB) free!
[1549959587] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11663784 kB) free.
[1549959763] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (25.6% change >= 20.0% threshold)
[1549959763] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.6% (22943772 kB) free.
[1549959993] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 61 users currently logged in
[1549961864] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.1% (2940760 kB) free!
[1549961984] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.5% (2848924 kB) free!
[1549962104] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.0% (2758904 kB) free!
[1549962795] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1549962981] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.5% (15041260 kB) free.
[1549963395] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1549963501] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1549963995] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1549964700] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1549965127] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9306388 kB) free!
[1549965204] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 69806 MB (31% inode=99%):
[1549965247] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.3% (9996008 kB) free.
[1549965297] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.34, 5.45, 5.09
[1549965300] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549965417] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.08, 5.74, 5.24
[1549965420] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1549965537] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.44, 5.61, 5.25
[1549965847] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 13.5% (6638468 kB) free!
[1549965967] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (3220528 kB) free!
[1549966020] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549966087] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.6% (6215888 kB) free!
[1549966140] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549966260] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1549966687] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (2691320 kB) free!
[1549967287] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.4% (6087876 kB) free!
[1549967460] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1549967533] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 69.9% (34431500 kB) free.
[1549967563] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.6% (4060652 kB) free!
[1549967683] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.4% (4028184 kB) free!
[1549967705] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.2% (23089176 kB) free.
[1549968060] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1549968180] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549968300] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1549968300] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549968914] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1549969034] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1549970113] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1549970701] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1549971301] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549971421] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1549972021] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549972141] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1549972261] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1549972588] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (164092 kB) free!
[1549972708] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (220588 kB) free!
[1549972737] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.87, 4.84, 5.00
[1549972828] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (147768 kB) free!
[1549972861] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1549973337] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.82, 5.08, 5.04
[1549973458] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.18, 5.11, 5.06
[1549973461] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1549973578] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.83, 5.35, 5.15
[1549973582] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1549973702] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1549973963] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.0% (1781532 kB) free!
[1549974083] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (341904 kB) free!
[1549974168] SERVICE ALERT: primary-server;Nginx Status;CRITICAL;SOFT;1;NGINX CRITICAL - 500 Can't connect to 10.147.212.141:80 (Connection refused)
[1549974203] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (340144 kB) free!
[1549974289] SERVICE ALERT: primary-server;Nginx Status;CRITICAL;SOFT;2;NGINX CRITICAL - 500 Can't connect to 10.147.212.141:80 (Connection refused)
[1549974409] SERVICE ALERT: primary-server;Nginx Status;OK;HARD;1;NGINX OK -  0.089 sec. response time, Active: 0 (Writing: 0 Reading: 0 Waiting: 0) ReqPerSec: 0.000 ConnPerSec: 0.000 ReqPerConn: 0.000
[1549974629] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 28.0% (4557312 kB) free.
[1549974778] SERVICE FLAPPING ALERT: qservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1549974778] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 5.00, 4.69, 4.87
[1549974799] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.2% (14166036 kB) free.
[1549975363] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 332 processes
[1549977302] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1549977393] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 59 users currently logged in
[1549977903] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1549977993] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 50 users currently logged in
[1549978023] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1549978468] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.22, 13.32, 9.16
[1549978588] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.57, 9.06, 8.10
[1549978623] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549978708] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.49, 6.21, 7.17
[1549978743] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549978863] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1549978893] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 51 users currently logged in
[1549978933] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1549978953] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 51 users currently logged in
[1549978995] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1549979013] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 51 users currently logged in
[1549979073] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 51 users currently logged in
[1549979308] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.17, 1.13, 3.94
[1549979484] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1549979673] SERVICE FLAPPING ALERT: webserver2;Current Users;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1549979673] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 41 users currently logged in
[1549980063] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1549980195] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1549980663] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549980784] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549980904] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549981333] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.7% (7741064 kB) free!
[1549981453] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.4% (11529996 kB) free.
[1549981514] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195764 kB) free!
[1549981635] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186940 kB) free!
[1549981755] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193736 kB) free!
[1549982053] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9115196 kB) free!
[1549982173] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.4% (11539820 kB) free.
[1549982704] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1549983304] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549983373] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.1% (7928312 kB) free!
[1549983424] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549983493] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.4% (11049368 kB) free.
[1549983544] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549984034] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549984154] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549984274] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1549984693] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.1% (6967280 kB) free!
[1549984744] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1549984813] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.7% (10716280 kB) free.
[1549984874] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1549985073] SERVICE FLAPPING ALERT: webserver2;Current Users;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1549985414] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8662820 kB) free!
[1549985534] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.0% (10833680 kB) free.
[1549985578] SERVICE FLAPPING ALERT: qservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1549985944] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549986064] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1549986205] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.8% (2073176 kB) free!
[1549986325] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.7% (2071548 kB) free!
[1549986445] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 12.7% (2070984 kB) free!
[1549986734] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.1% (8435516 kB) free!
[1549986854] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.8% (10739680 kB) free.
[1549987264] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1549987385] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549987505] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1549988054] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.7% (8224604 kB) free!
[1549988174] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.6% (10666168 kB) free.
[1549988313] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 93.2% (15152908 kB) free.
[1549988356] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 65.6% (16068980 kB) free.
[1549988705] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1549989305] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1549989374] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.1% (7463416 kB) free!
[1549989425] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549989494] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.6% (10647752 kB) free.
[1549989545] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549990094] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9212360 kB) free!
[1549990145] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1549990214] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.3% (6076100 kB) free!
[1549990335] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.4% (10060196 kB) free.
[1549990746] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549990866] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1549990934] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.2% (7503440 kB) free!
[1549990986] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1549991054] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.6% (10630836 kB) free.
[1549991586] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549991596] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1549991654] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9484560 kB) free!
[1549991706] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549991774] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.2% (5996392 kB) free!
[1549991826] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549991894] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 20.0% (9848584 kB) free!
[1549993094] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.2% (10439076 kB) free.
[1549993694] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9368008 kB) free!
[1549993814] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.6% (5723216 kB) free!
[1549993934] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.6% (9652936 kB) free!
[1549995134] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1549995134] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.1% (10393860 kB) free.
[1549995735] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9049620 kB) free!
[1549995855] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.4% (5599576 kB) free!
[1549995975] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.5% (9589656 kB) free!
[1549996026] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes