[1559846299] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1559846367] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (10002728 kB) free.
[1559846685] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 414 processes
[1559846899] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559846995] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1559847019] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559847139] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559847195] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.6% (1725436 kB) free!
[1559847568] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (1412220 kB) free!
[1559847688] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.5% (9594052 kB) free!
[1559847808] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.6% (7697228 kB) free!
[1559847885] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 407 processes
[1559848194] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1559848395] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166864 kB) free!
[1559848680] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1559849539] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1559850139] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559850195] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.4% (1849172 kB) free!
[1559850208] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10338960 kB) free.
[1559850259] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559850379] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559850795] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (163612 kB) free!
[1559850808] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (1461604 kB) free!
[1559850928] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.8% (9773464 kB) free!
[1559851048] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.7% (7723352 kB) free!
[1559851395] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.1% (1793980 kB) free!
[1559851485] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 396 processes
[1559851995] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (164320 kB) free!
[1559852085] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 394 processes
[1559852180] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1559852685] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 375 processes
[1559852849] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (9988540 kB) free.
[1559853285] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 408 processes
[1559853380] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559853421] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2467132 kB) free!
[1559853500] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559853672] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559854049] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (1577748 kB) free!
[1559854169] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.9% (9824568 kB) free!
[1559854289] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.7% (7746584 kB) free!
[1559854701] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559854821] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559854941] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559855595] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.1% (1965748 kB) free!
[1559856089] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10315156 kB) free.
[1559856285] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 391 processes
[1559856741] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559856795] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1281092 kB) free!
[1559857289] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (1521804 kB) free!
[1559857409] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.7% (9727928 kB) free!
[1559857485] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 408 processes
[1559857529] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.5% (7614196 kB) free!
[1559857941] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559857995] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.8% (2401864 kB) free!
[1559858062] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559858182] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559859195] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (163488 kB) free!
[1559859329] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10066592 kB) free.
[1559859422] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1559859982] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1559860395] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.9% (3559180 kB) free.
[1559860529] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (2732484 kB) free!
[1559860582] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559860649] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10236308 kB) free.
[1559860702] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559860822] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559860995] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (158584 kB) free!
[1559861115] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.3% (2971508 kB) free!
[1559861235] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.4% (2827464 kB) free!
[1559861249] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (1992672 kB) free!
[1559861369] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10080996 kB) free.
[1559861836] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.4% (3789204 kB) free.
[1559862569] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (1353048 kB) free!
[1559862622] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1559862689] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.2% (9476500 kB) free!
[1559862810] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.3% (7526632 kB) free!
[1559863222] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559863342] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559863462] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559864836] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.9% (3069584 kB) free!
[1559864956] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.4% (4126792 kB) free.
[1559865157] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 326.59, 95.13, 33.90
[1559865210] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9941264 kB) free.
[1559865277] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 489.62, 234.08, 92.24
[1559865397] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 511.02, 326.82, 143.42
[1559865672] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 582.69, 349.42, 149.31
[1559865792] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 593.06, 430.57, 203.31
[1559865912] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 216.68, 364.58, 208.51
[1559865924] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 723 processes
[1559866021] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 22.4% (5519256 kB) free.
[1559866044] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 464 processes
[1559866294] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.4% (2791056 kB) free!
[1559866410] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (1410704 kB) free!
[1559866462] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559866530] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.1% (9399092 kB) free!
[1559866650] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.4% (7600692 kB) free!
[1559866680] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.05, 9.98, 6.78
[1559866800] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.19, 6.71, 5.97
[1559866920] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.09, 4.52, 5.26
[1559867062] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559867085] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 427 processes
[1559867182] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559867302] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559867494] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2060472 kB) free!
[1559867520] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.40, 0.86, 2.88
[1559867685] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 401 processes
[1559867956] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (706624 kB) free!
[1559868076] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (157808 kB) free!
[1559868196] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (808492 kB) free!
[1559868312] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.13, 0.57, 18.10
[1559868397] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.75, 1.93, 18.62
[1559868450] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10176088 kB) free.
[1559868594] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1559868714] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1559868796] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.0% (3566308 kB) free.
[1559868834] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1559868923] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 35.5% (8712400 kB) free.
[1559869021] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (844556 kB) free!
[1559869051] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (1026328 kB) free!
[1559869103] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1559869141] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (834820 kB) free!
[1559869171] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.4% (9555000 kB) free!
[1559869224] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 17.6% (4304884 kB) free!
[1559869261] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (836264 kB) free!
[1559869284] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 16.6% (4063236 kB) free!
[1559869291] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.3% (7554204 kB) free!
[1559869294] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.5% (2563776 kB) free!
[1559869344] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 14.6% (3576104 kB) free!
[1559869404] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 14.6% (3579800 kB) free!
[1559869605] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 20.5% (3337856 kB) free.
[1559869703] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559869823] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559869943] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559869996] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.1% (2611976 kB) free!
[1559870112] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.19, 0.23, 2.85
[1559870116] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.6% (3507104 kB) free.
[1559870126] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 618 processes
[1559870186] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 609 processes
[1559870197] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.05, 1.70, 4.13
[1559870486] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 612 processes
[1559870494] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1559870494] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2455916 kB) free!
[1559870546] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 611 processes
[1559870607] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 608 processes
[1559870805] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (412752 kB) free!
[1559870925] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (412140 kB) free!
[1559871045] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (412864 kB) free!
[1559871234] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1559871316] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.1% (3097804 kB) free!
[1559871436] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.5% (3480416 kB) free.
[1559871691] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9952508 kB) free.
[1559872291] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3617612 kB) free!
[1559872343] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1559872411] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 18.2% (8955072 kB) free!
[1559872531] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.3% (7062812 kB) free!
[1559872636] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.4% (2978704 kB) free!
[1559872756] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.2% (3433360 kB) free.
[1559873543] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559873663] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559873783] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559874885] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 411 processes
[1559875531] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9894852 kB) free.
[1559875756] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.9% (3220452 kB) free!
[1559875876] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.5% (3165088 kB) free!
[1559875996] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.6% (3183240 kB) free!
[1559876086] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 421 processes
[1559876731] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (808200 kB) free!
[1559876784] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1559876851] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.0% (9368664 kB) free!
[1559876971] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.9% (7345680 kB) free!
[1559877384] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559877504] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559877624] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559878396] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.7% (3363484 kB) free.
[1559878486] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 425 processes
[1559878630] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24676 MB (10% inode=99%):
[1559878771] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10114896 kB) free.
[1559878996] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.6% (3025344 kB) free!
[1559879086] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 433 processes
[1559879116] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.1% (2933392 kB) free!
[1559879236] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.4% (2010720 kB) free!
[1559879247] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (1819756 kB) free!
[1559879371] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (224044 kB) free!
[1559879424] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1559879491] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.1% (9428128 kB) free!
[1559879612] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.4% (7083268 kB) free!
[1559879847] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (224736 kB) free!
[1559880625] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559880745] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559880865] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559881047] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (2215932 kB) free!
[1559881091] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 363 processes
[1559881211] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 362 processes
[1559881294] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1559881331] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 362 processes
[1559881412] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10078424 kB) free.
[1559881486] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 441 processes
[1559881636] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1192652 kB) free!
[1559881647] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1559881647] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (167196 kB) free!
[1559882087] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 434 processes
[1559882531] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1559882612] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7728196 kB) free!
[1559882665] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1559882687] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 418 processes
[1559882732] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.7% (9725244 kB) free!
[1559882831] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 108400 MB (48% inode=99%):
[1559882852] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.2% (7479744 kB) free!
[1559882916] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 64.8% (10538640 kB) free.
[1559883046] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 33.2% (5400208 kB) free.
[1559883112] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 9 users currently logged in
[1559883131] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 362 processes
[1559883232] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 9 users currently logged in
[1559883251] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 361 processes
[1559883265] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559883287] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 472 processes
[1559883352] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1559883371] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 363 processes
[1559883385] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559883505] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559883516] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (172084 kB) free!
[1559883636] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (186940 kB) free!
[1559883646] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (177604 kB) free!
[1559883756] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (167748 kB) free!
[1559883766] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (165400 kB) free!
[1559883834] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1559883886] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174248 kB) free!
[1559883952] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1559883954] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1559884037] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.1% (1796108 kB) free!
[1559884074] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1559884320] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 582.06, 410.23, 197.07
[1559884440] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 222.09, 336.29, 196.42
[1559884526] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.56, 5.86, 5.49
[1559884560] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 160.96, 275.11, 190.96
[1559884646] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.17, 6.24, 5.66
[1559884654] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1559884674] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1559884766] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.04, 6.07, 5.67
[1559884769] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 21.8% (5348860 kB) free.
[1559885152] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1559885171] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1559885171] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1559885237] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (186528 kB) free!
[1559885252] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10205116 kB) free.
[1559885837] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.0% (2755120 kB) free!
[1559885969] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (4512536 kB) free!
[1559886089] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (4422124 kB) free!
[1559886209] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.7% (4328296 kB) free!
[1559886397] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 70.17, 58.06, 31.17
[1559886452] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (215000 kB) free!
[1559886506] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559886517] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 33.91, 52.49, 32.66
[1559886573] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9940416 kB) free.
[1559886637] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 5.96, 35.68, 28.91
[1559886960] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.27, 0.52, 17.00
[1559887237] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.43, 6.62, 16.10
[1559887637] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (158968 kB) free!
[1559887706] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559887826] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559887946] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559888010] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2360832 kB) free!
[1559888048] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 18.8% (1385548 kB) free!
[1559888168] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 14.3% (1050840 kB) free!
[1559888288] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (720040 kB) free!
[1559888760] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.52, 0.98, 4.27
[1559888972] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (230048 kB) free!
[1559889037] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.00, 3.03, 4.86
[1559889092] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10983692 kB) free.
[1559889489] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 91.7% (6741952 kB) free.
[1559890292] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236608 kB) free!
[1559890346] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1559890412] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.1% (10405004 kB) free.
[1559890946] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559891066] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559891152] SERVICE FLAPPING ALERT: secondary-server;Current Users;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1559891152] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1559891186] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559891237] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.8% (1756708 kB) free!
[1559891612] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (234148 kB) free!
[1559891732] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10588172 kB) free.
[1559891837] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (180952 kB) free!
[1559891854] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1559892932] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (556336 kB) free!
[1559892986] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1559893047] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559893052] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10862536 kB) free.
[1559893552] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1559893586] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559893672] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1559893706] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559893793] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1559893827] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559894252] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (217368 kB) free!
[1559894372] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10633680 kB) free.
[1559894687] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559895371] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559895572] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (223960 kB) free!
[1559895627] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1559895637] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 514.55, 398.89, 197.33
[1559895692] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1559895692] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10825648 kB) free.
[1559895757] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 514.95, 437.39, 235.88
[1559895877] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 514.20, 462.78, 269.62
[1559895885] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5383 MB (10% inode=97%):
[1559896037] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.4% (2504636 kB) free!
[1559896227] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1559896347] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559896467] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559896485] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6985 MB (13% inode=98%):
[1559896560] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 14.66, 18.60, 8.69
[1559896638] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174936 kB) free!
[1559896680] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.33, 12.55, 7.67
[1559896689] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (699472 kB) free!
[1559896800] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.59, 8.50, 6.78
[1559896809] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (396656 kB) free!
[1559896892] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (436140 kB) free!
[1559896930] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (123972 kB) free!
[1559897013] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11408256 kB) free.
[1559897401] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1559897401] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.36, 1.49, 3.73
[1559897667] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1559898212] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (1211812 kB) free!
[1559898267] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559898332] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11816672 kB) free.
[1559898388] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1559898438] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 32.9% (5334816 kB) free.
[1559898508] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559898878] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.16, 3.10, 16.85
[1559898932] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (227124 kB) free!
[1559899038] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (999396 kB) free!
[1559899052] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10569732 kB) free.
[1559899158] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (160600 kB) free!
[1559899279] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155156 kB) free!
[1559899707] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 604 processes
[1559899767] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 611 processes
[1559899930] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 91.4% (6725584 kB) free.
[1559900252] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (4847088 kB) free!
[1559900372] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (2121092 kB) free!
[1559900492] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (216880 kB) free!
[1559901278] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.71, 2.05, 4.50
[1559901454] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1559901679] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.3% (3286864 kB) free.
[1559901692] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.1% (6939192 kB) free!
[1559902193] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1559902279] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (1500944 kB) free!
[1559902399] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (159272 kB) free!
[1559902519] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.6% (2859792 kB) free!
[1559902893] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10517064 kB) free.
[1559903336] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.1% (2785972 kB) free!
[1559903457] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.2% (2964092 kB) free!
[1559903493] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (215556 kB) free!
[1559903577] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.7% (2877528 kB) free!
[1559903613] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (10004368 kB) free.
[1559903687] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 403 processes
[1559903720] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (173572 kB) free!
[1559904287] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 441 processes
[1559904813] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (225464 kB) free!
[1559904933] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11658588 kB) free.
[1559905377] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1417468 kB) free!
[1559905478] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 135.76, 114.69, 61.91
[1559905533] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (238240 kB) free!
[1559905598] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 26.76, 83.48, 57.16
[1559905653] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10982368 kB) free.
[1559905719] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1559905719] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.91, 56.39, 50.43
[1559906254] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1559906853] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (220308 kB) free!
[1559906908] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1559906919] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.46, 2.71, 15.11
[1559906973] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11374060 kB) free.
[1559907130] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 15.1% (1109996 kB) free!
[1559907250] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 12.9% (951796 kB) free!
[1559907370] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 10.8% (790720 kB) free!
[1559907508] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559907566] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.50, 3.61, 4.45
[1559907592] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 62 users currently logged in
[1559907629] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559907652] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 62 users currently logged in
[1559907712] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 58 users currently logged in
[1559907749] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559907811] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.2% (3485988 kB) free!
[1559907970] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1559907970] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (168976 kB) free!
[1559907971] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 358 processes
[1559908054] SERVICE ALERT: sigqscluster1;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1559908091] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 363 processes
[1559908193] SERVICE FLAPPING ALERT: secondary-server;Current Users;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1559908211] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 362 processes
[1559908411] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189324 kB) free!
[1559908487] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 413 processes
[1559908521] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.2% (1810568 kB) free!
[1559908719] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.35, 1.84, 3.87
[1559908773] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (1677120 kB) free!
[1559908801] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559908893] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12261476 kB) free.
[1559909087] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1559909087] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 415 processes
[1559909121] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161444 kB) free!
[1559909254] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1559909374] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1559909412] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1559909493] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (454712 kB) free!
[1559909494] SERVICE FLAPPING ALERT: sigqscluster1;Current Users;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1559909494] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1559909549] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1559909613] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11298040 kB) free.
[1559910149] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1559910269] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1559910389] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559910612] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 367 processes
[1559910732] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1559910777] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 83.3% (13535040 kB) free.
[1559910813] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (515300 kB) free!
[1559910933] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11336684 kB) free.
[1559910971] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.0% (6842876 kB) free.
[1559911332] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 363 processes
[1559911452] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1559911487] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 417 processes
[1559912087] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 391 processes
[1559912112] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 144.44, 111.62, 54.95
[1559912133] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (436036 kB) free!
[1559912189] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1559912232] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 144.09, 122.36, 65.75
[1559912254] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11156176 kB) free.
[1559912352] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 20.85, 82.45, 58.00
[1559912652] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 365 processes
[1559912688] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 401 processes
[1559912772] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 364 processes
[1559912789] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559912892] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 361 processes
[1559912909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559913029] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559913288] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 390 processes
[1559913453] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (212268 kB) free!
[1559913519] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.43, 4.65, 5.16
[1559913552] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.60, 1.97, 16.26
[1559913573] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12900424 kB) free.
[1559913639] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.06, 3.79, 4.78
[1559913921] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.4% (3795004 kB) free.
[1559914467] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 569 processes
[1559914521] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (166332 kB) free!
[1559914527] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 576 processes
[1559914641] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (174668 kB) free!
[1559914752] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.67, 0.64, 4.86
[1559914761] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (167720 kB) free!
[1559914773] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (1698240 kB) free!
[1559914827] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 574 processes
[1559914830] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1559914887] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 568 processes
[1559914893] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12380696 kB) free.
[1559915430] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1559915550] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559915670] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559916093] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (2459932 kB) free!
[1559916213] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12790836 kB) free.
[1559916813] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (1170924 kB) free!
[1559916870] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1559916933] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11870012 kB) free.
[1559917470] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559917590] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559917710] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559918088] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 422 processes
[1559918133] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (1337168 kB) free!
[1559918253] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11809384 kB) free.
[1559918688] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 425 processes
[1559918961] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.0% (4058004 kB) free.
[1559919039] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559919289] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 415 processes
[1559919453] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (1022088 kB) free!
[1559919510] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1559919562] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (157836 kB) free!
[1559919573] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11542568 kB) free.
[1559919682] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (165020 kB) free!
[1559919802] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156148 kB) free!
[1559919889] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 441 processes
[1559919971] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1559920110] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559920173] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (2516800 kB) free!
[1559920230] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1559920293] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11081540 kB) free.
[1559920294] SERVICE FLAPPING ALERT: sigqscluster1;Current Users;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1559920380] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1559920489] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 406 processes
[1559920500] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1559920620] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1559920830] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559920950] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559921070] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559921089] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 418 processes
[1559921493] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (1745032 kB) free!
[1559921613] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11828736 kB) free.
[1559922601] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16.68, 89.85, 60.49
[1559922721] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.42, 60.19, 53.18
[1559922802] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.6% (2045904 kB) free!
[1559922812] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (2750296 kB) free!
[1559922814] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3094980 kB) free!
[1559922841] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.59, 40.37, 46.77
[1559922934] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12324040 kB) free.
[1559923544] SERVICE ALERT: sigqscluster1;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 8245 MB (16% inode=98%):
[1559923664] SERVICE ALERT: sigqscluster1;HDD_Root status;CRITICAL;SOFT;2;DISK CRITICAL - free space: / 5370 MB (10% inode=97%):
[1559923785] SERVICE ALERT: sigqscluster1;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5370 MB (10% inode=97%):
[1559924002] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165604 kB) free!
[1559924041] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.09, 0.90, 12.98
[1559924089] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 394 processes
[1559924134] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (2677228 kB) free!
[1559924254] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12104704 kB) free.
[1559924385] SERVICE ALERT: sigqscluster1;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6557 MB (12% inode=98%):
[1559924689] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 436 processes
[1559924854] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (1943472 kB) free!
[1559924974] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11528260 kB) free.
[1559925202] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.6% (2691752 kB) free!
[1559925290] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 392 processes
[1559925420] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1559925694] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 42.7% (10478256 kB) free.
[1559925841] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.40, 0.69, 4.81
[1559925890] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 400 processes
[1559926020] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1559926174] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (2478308 kB) free!
[1559926294] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11922200 kB) free.
[1559926402] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (163824 kB) free!
[1559926692] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 374 processes
[1559926812] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 372 processes
[1559927090] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 422 processes
[1559927494] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3082104 kB) free!
[1559927615] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11854788 kB) free.
[1559927690] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 392 processes
[1559928802] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.2% (2948172 kB) free!
[1559928814] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (2530068 kB) free!
[1559928870] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559928934] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11691676 kB) free.
[1559929185] SERVICE ALERT: sigqscluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 17763 MB (34% inode=99%):
[1559929470] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559929490] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 383 processes
[1559929590] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1559929710] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559930090] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 365 processes
[1559930134] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (2603700 kB) free!
[1559930220] SERVICE FLAPPING ALERT: webservercluster1;Current Users;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1559930220] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 9 users currently logged in
[1559930254] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11594576 kB) free.
[1559930602] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (1501592 kB) free!
[1559931454] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (2894508 kB) free!
[1559931510] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559931574] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11787544 kB) free.
[1559932110] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes