[1583951445] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33335480 kB) free.
[1583952024] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189752 kB) free!
[1583952045] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10434920 kB) free!
[1583952144] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195364 kB) free!
[1583952165] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (6962000 kB) free!
[1583952203] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190984 kB) free!
[1583952264] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174824 kB) free!
[1583952286] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32441692 kB) free.
[1583952323] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167972 kB) free!
[1583952443] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171216 kB) free!
[1583952885] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10046408 kB) free!
[1583952965] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1583953005] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.5% (25715164 kB) free!
[1583953125] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31699712 kB) free.
[1583953607] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186812 kB) free!
[1583953667] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177060 kB) free!
[1583953725] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9811000 kB) free!
[1583953727] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (184376 kB) free!
[1583953787] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (186484 kB) free!
[1583953845] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34884232 kB) free.
[1583954087] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 71.0% (17408868 kB) free.
[1583954190] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583954445] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13000200 kB) free!
[1583954565] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (9570292 kB) free!
[1583954685] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34858816 kB) free.
[1583954765] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 38.4% (9464332 kB) free.
[1583955285] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12244860 kB) free!
[1583955365] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.9% (3421648 kB) free!
[1583955405] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (8928704 kB) free!
[1583955485] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (1730252 kB) free!
[1583955525] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33680140 kB) free.
[1583955605] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193384 kB) free!
[1583956125] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11019800 kB) free!
[1583956245] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34766576 kB) free.
[1583956636] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (220888 kB) free!
[1583956756] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (256568 kB) free!
[1583956845] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13410316 kB) free!
[1583956876] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (257028 kB) free!
[1583956966] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (10070812 kB) free!
[1583957086] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34693984 kB) free.
[1583957216] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13011908 kB) free!
[1583957336] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (4000364 kB) free!
[1583957450] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 210649 MB (19% inode=99%):
[1583957456] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402728 kB) free!
[1583957570] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 200963 MB (18% inode=99%):
[1583957686] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12333616 kB) free!
[1583957690] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 191349 MB (17% inode=99%):
[1583957806] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9068712 kB) free!
[1583957926] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32523760 kB) free.
[1583958526] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10743936 kB) free!
[1583958646] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34644644 kB) free.
[1583958676] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 20.6% (10139812 kB) free.
[1583959246] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (11941816 kB) free!
[1583959366] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34609516 kB) free.
[1583959644] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 45.3% (11109252 kB) free.
[1583959966] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13323964 kB) free!
[1583960086] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (10007168 kB) free!
[1583960206] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32692420 kB) free.
[1583960354] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236828 kB) free!
[1583960474] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (227524 kB) free!
[1583960594] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 87.3% (43028264 kB) free.
[1583960806] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11017352 kB) free!
[1583960926] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34476200 kB) free.
[1583961314] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.6% (43154084 kB) free.
[1583961526] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11825704 kB) free!
[1583961646] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34475440 kB) free.
[1583962247] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12963136 kB) free!
[1583962367] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26799284 kB) free.
[1583962414] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(217) < 103.16.101.72 [12/Mar/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1583962536] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583962646] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2040) < 203.212.70.170 [12/Mar/2020:03:01:32 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583962704] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(269) < 2020/03/12 03:00:32 [error] 12249#0: *69092593 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583962766] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583962825] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583962967] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13378192 kB) free!
[1583963087] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (10083236 kB) free!
[1583963207] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31338612 kB) free.
[1583964407] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13552648 kB) free!
[1583964527] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34144684 kB) free.
[1583965127] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6316280 kB) free!
[1583965247] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1583965247] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33701636 kB) free.
[1583965847] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5904536 kB) free!
[1583965967] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33572592 kB) free.
[1583966567] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8924996 kB) free!
[1583966668] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 696 processes
[1583966687] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.5% (25744076 kB) free!
[1583966807] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29828212 kB) free.
[1583968007] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11440324 kB) free!
[1583968127] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35650448 kB) free.
[1583968727] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15698944 kB) free!
[1583968847] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33720188 kB) free.
[1583969844] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165048 kB) free!
[1583969964] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (153832 kB) free!
[1583970047] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15581988 kB) free!
[1583970084] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177656 kB) free!
[1583970167] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33485560 kB) free.
[1583970864] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 76.3% (18705536 kB) free.
[1583971205] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 50.0% (12334072 kB) free.
[1583971367] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17331240 kB) free!
[1583971487] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35438032 kB) free.
[1583972966] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(148) < 49.50.67.175 [12/Mar/2020:05:55:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1583973086] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583973287] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19953056 kB) free!
[1583973407] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34164872 kB) free.
[1583973818] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 88.0% (116036832 kB) free.
[1583973864] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.1% (38964228 kB) free.
[1583973909] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.9% (21299908 kB) free.
[1583974286] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(132) < 92.42.104.166 [12/Mar/2020:06:20:16 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1583974406] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583974490] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 229843 MB (21% inode=99%):
[1583975405] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184708 kB) free!
[1583975525] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177988 kB) free!
[1583975645] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164220 kB) free!
[1583976362] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.20948.31> (10.147.212.158:47004 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1583976482] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1583976909] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.0% (20983036 kB) free.
[1583983193] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.5% (4773912 kB) free!
[1583983313] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.3% (4724704 kB) free!
[1583983433] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.1% (4674436 kB) free!
[1583985312] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4882968 kB) free!
[1583985432] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.6% (4099008 kB) free!
[1583985552] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.4% (4050628 kB) free!
[1583985833] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2362392 kB) free!
[1583987352] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2284828 kB) free!
[1583988275] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9543648 kB) free!
[1583988396] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.4% (9543456 kB) free!
[1583988407] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24815632 kB) free!
[1583988495] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 334.44, 370.70, 206.20
[1583988516] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.4% (9543216 kB) free!
[1583988527] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.3% (43881744 kB) free.
[1583988615] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 68.06, 257.03, 184.49
[1583988735] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 34.58, 181.53, 165.61
[1583988868] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 709 processes
[1583988988] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 708 processes
[1583989006] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583989108] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 709 processes
[1583989127] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20722696 kB) free!
[1583989136] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1583989140] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1159 processes
[1583989227] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 261.60, 210.13, 99.71
[1583989247] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42137764 kB) free.
[1583989257] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1583989262] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1157 processes
[1583989347] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 218.38, 207.69, 112.13
[1583989383] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1156 processes
[1583989468] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 149.69, 191.83, 118.44
[1583989847] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17537608 kB) free!
[1583989967] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40210680 kB) free.
[1583990147] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.6% (20391664 kB) free.
[1583990170] SERVICE ALERT: secondary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 367 processes
[1583990228] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.2% (21610672 kB) free.
[1583990291] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1583990461] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.5% (3925300 kB) free!
[1583990567] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14198592 kB) free!
[1583990581] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.0% (3366324 kB) free!
[1583990583] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 403 processes
[1583990688] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8640804 kB) free!
[1583990701] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.4% (2794216 kB) free!
[1583990808] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.6% (46931900 kB) free.
[1583991206] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(35) < 103.255.146.164 [12/Mar/2020:11:01:35 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583991302] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (183668 kB) free!
[1583991326] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583991408] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24885088 kB) free!
[1583991528] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.6% (48207796 kB) free.
[1583991927] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(41) < 103.255.146.164 [12/Mar/2020:11:14:21 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583992047] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583992128] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24428244 kB) free!
[1583992248] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.8% (20865648 kB) free!
[1583992303] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(4) < Channel error on connection <0.1364.3861> (10.147.212.181:56046 -> 10.147.212.160:5672, vhost: '/', user: 'guest'), channel 1:
[1583992327] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 12.1% (2993752 kB) free!
[1583992368] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37595128 kB) free.
[1583992423] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1583992448] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 10.5% (2593632 kB) free!
[1583992468] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.78, 4.37, 14.28
[1583992552] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.1% (2750528 kB) free!
[1583992568] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1934092 kB) free!
[1583992672] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (1294520 kB) free!
[1583992716] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3142128 kB) free!
[1583992792] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1583992792] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169652 kB) free!
[1583992968] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (11939592 kB) free!
[1583993088] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8611116 kB) free!
[1583993208] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40220452 kB) free.
[1583993437] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 17.9% (6229636 kB) free!
[1583993535] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 35.29, 33.04, 37.55
[1583993539] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 17.7% (6193876 kB) free!
[1583993557] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 16.3% (5686756 kB) free!
[1583993659] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 16.5% (5768808 kB) free!
[1583993677] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 14.8% (5140412 kB) free!
[1583993779] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 15.2% (5322564 kB) free!
[1583993808] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11902468 kB) free!
[1583993847] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(78) < 103.255.146.164 [12/Mar/2020:11:42:32 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583993928] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8747264 kB) free!
[1583993967] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583994048] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41172572 kB) free.
[1583994076] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8962528 kB) free!
[1583994196] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.0% (8400144 kB) free!
[1583994268] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 2.96, 2.44, 4.10
[1583994277] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (2538216 kB) free!
[1583994316] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.9% (7815512 kB) free!
[1583994379] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (3230408 kB) free!
[1583994463] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.1% (20513664 kB) free.
[1583994582] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 50.5% (17593372 kB) free.
[1583994624] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 89.7% (31399576 kB) free.
[1583994648] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16226276 kB) free!
[1583994712] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.2% (15544432 kB) free.
[1583994741] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 63.45, 72.50, 51.92
[1583994742] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583994743] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.0% (38954116 kB) free.
[1583994768] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (12505868 kB) free!
[1583994819] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.2% (22769932 kB) free.
[1583994847] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583994854] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583994864] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1583994889] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 63.1% (83156180 kB) free.
[1583994913] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (4684488 kB) free!
[1583994918] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.9% (29995428 kB) free.
[1583994969] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1583994975] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1583995033] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (4475348 kB) free!
[1583995153] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.3% (4243356 kB) free!
[1583995663] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195756 kB) free!
[1583995767] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(66) < 103.255.146.164 [12/Mar/2020:12:19:08 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583995783] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177612 kB) free!
[1583995887] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583995903] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183288 kB) free!
[1583995941] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1583995941] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 38.03, 35.52, 38.65
[1583996064] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (4638260 kB) free!
[1583996185] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4498772 kB) free!
[1583996235] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.5% (4530704 kB) free!
[1583996305] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (4381072 kB) free!
[1583996353] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2020112 kB) free!
[1583996355] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.0% (4155836 kB) free!
[1583996475] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 15.4% (3786164 kB) free!
[1583996487] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(64) < 103.255.146.164 [12/Mar/2020:12:30:58 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583996607] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583996699] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 342 processes
[1583996819] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1583997075] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2266188 kB) free!
[1583997142] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 17.71, 19.61, 26.73
[1583998395] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9005144 kB) free!
[1583998407] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(54) < 103.255.146.164 [12/Mar/2020:12:56:48 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583998515] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (8762040 kB) free!
[1583998527] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583998635] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (8510136 kB) free!
[1583998705] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1645596 kB) free!
[1583999219] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1583999265] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1583999340] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1583999363] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1583999385] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1583999483] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1583999505] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1583999603] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1583999691] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24789892 kB) free!
[1583999811] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.8% (20886916 kB) free!
[1583999931] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17323176 kB) free!
[1584000105] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1584000147] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.0% (7374456 kB) free!
[1584000203] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1584000267] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 14.0% (6899636 kB) free!
[1584000387] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.0% (6894724 kB) free!
[1584000435] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4885344 kB) free!
[1584001131] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33322436 kB) free.
[1584001527] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(56) < 103.255.146.164 [12/Mar/2020:13:51:25 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584001647] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584001731] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13118044 kB) free!
[1584001739] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 4 warn out of 342 processes
[1584001851] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30508928 kB) free.
[1584001859] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1584002451] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16813212 kB) free!
[1584002571] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12894508 kB) free!
[1584002692] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (9354280 kB) free!
[1584002787] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1584002787] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (2470416 kB) free!
[1584002955] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 696 processes
[1584003112] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.5% (3473884 kB) free!
[1584003232] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.1% (3163192 kB) free!
[1584003292] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1584003292] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17422784 kB) free!
[1584003352] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.7% (2852192 kB) free!
[1584003892] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27735684 kB) free.
[1584003952] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (1121704 kB) free!
[1584004047] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(96) < 103.255.146.164 [12/Mar/2020:14:31:02 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584004167] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584004424] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 17.3% (4274152 kB) free!
[1584004492] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7632336 kB) free!
[1584004545] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 15.8% (3906324 kB) free!
[1584004591] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.8% (2563576 kB) free!
[1584004612] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4267752 kB) free!
[1584004665] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 14.3% (3518564 kB) free!
[1584004711] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (3046808 kB) free!
[1584004732] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (31042496 kB) free.
[1584004831] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3399356 kB) free.
[1584005118] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9380800 kB) free!
[1584005238] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.1% (8936220 kB) free!
[1584005265] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1718788 kB) free!
[1584005332] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9308692 kB) free!
[1584005358] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.2% (8490004 kB) free!
[1584005367] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(55) < 103.255.146.164 [12/Mar/2020:14:55:11 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584005431] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (2899668 kB) free!
[1584005452] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (5577200 kB) free!
[1584005487] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(191) < 92.42.104.166 [12/Mar/2020:15:00:12 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1584005551] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.4% (1848504 kB) free!
[1584005572] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2065620 kB) free!
[1584005608] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584005671] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.5% (2195864 kB) free!
[1584005731] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 86.0% (21085956 kB) free.
[1584005753] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1584005753] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.9% (15692248 kB) free.
[1584005811] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.5% (21600840 kB) free.
[1584005883] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 90.9% (44811124 kB) free.
[1584005950] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 57.5% (75824032 kB) free.
[1584005958] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1584005958] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.3% (29721720 kB) free.
[1584006013] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.0% (38461544 kB) free.
[1584006024] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(43) < 2020/03/12 15:08:37 [error] 12249#0: *69480293 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584006065] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.9% (21298348 kB) free.
[1584006124] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.4% (20697020 kB) free.
[1584006144] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584006412] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177668 kB) free!
[1584006465] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 91.9% (22683984 kB) free.
[1584006532] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186792 kB) free!
[1584006652] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.4% (21574108 kB) free.
[1584006807] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(55) < 103.255.146.164 [12/Mar/2020:15:18:11 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584006927] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584007471] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.5% (3818608 kB) free.
[1584007824] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.3% (8522236 kB) free!
[1584007874] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.4% (4296704 kB) free!
[1584007942] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584007944] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.9% (8323152 kB) free!
[1584007994] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.2% (3988956 kB) free!
[1584008064] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.7% (8247996 kB) free!
[1584008114] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.0% (3709040 kB) free!
[1584008234] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (866212 kB) free!
[1584008354] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (856432 kB) free!
[1584008474] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (873340 kB) free!
[1584008646] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.4% (14045864 kB) free.
[1584008672] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3164348 kB) free!
[1584008714] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9804148 kB) free!
[1584008714] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2292188 kB) free!
[1584008792] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (3202876 kB) free!
[1584008805] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.5% (21590216 kB) free.
[1584008834] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.5% (9601588 kB) free!
[1584008912] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (1010032 kB) free!
[1584008954] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.1% (9414160 kB) free!
[1584009222] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1584009222] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 49.2% (7994748 kB) free.
[1584009327] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(40) < 103.255.146.164 [12/Mar/2020:16:01:15 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584009447] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584009864] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4209656 kB) free!
[1584010012] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167700 kB) free!
[1584010014] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.8% (38824200 kB) free.
[1584010129] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.3% (20055652 kB) free.
[1584010514] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44957 MB (20% inode=99%):
[1584010634] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44957 MB (20% inode=99%):
[1584010648] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(77) < 103.255.146.164 [12/Mar/2020:16:25:13 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584010729] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171812 kB) free!
[1584010754] SERVICE ALERT: primary-server;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44957 MB (20% inode=99%):
[1584010768] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584010821] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.4% (20329024 kB) free.
[1584010873] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164596 kB) free!
[1584010931] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186900 kB) free!
[1584010993] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186828 kB) free!
[1584011047] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.8% (20797448 kB) free.
[1584011051] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173448 kB) free!
[1584011111] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.3% (20167368 kB) free.
[1584011223] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9114560 kB) free!
[1584011343] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.9% (8802812 kB) free!
[1584011350] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11399860 kB) free!
[1584011421] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182656 kB) free!
[1584011463] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.3% (8529800 kB) free!
[1584011470] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7954920 kB) free!
[1584011541] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174396 kB) free!
[1584011564] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 53.6% (70646092 kB) free.
[1584011653] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180328 kB) free!
[1584011661] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179848 kB) free!
[1584011701] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.9% (20956964 kB) free.
[1584011712] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188004 kB) free!
[1584011773] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186652 kB) free!
[1584011832] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182468 kB) free!
[1584011893] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168368 kB) free!
[1584011952] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176676 kB) free!
[1584011954] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (3981844 kB) free!
[1584011968] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(67) < 103.255.146.164 [12/Mar/2020:16:42:48 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584012088] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584012223] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (422104 kB) free!
[1584012342] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.0% (4653136 kB) free!
[1584012343] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.6% (1731056 kB) free!
[1584012462] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.6% (4325440 kB) free!
[1584012463] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (3199348 kB) free!
[1584012582] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.6% (4062356 kB) free!
[1584012688] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(50) < 103.255.146.164 [12/Mar/2020:16:52:56 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1584012809] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584013063] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.2% (3926000 kB) free.
[1584013263] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4444208 kB) free!
[1584013663] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.7% (2711484 kB) free!
[1584013782] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182636 kB) free!
[1584013783] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.8% (4037744 kB) free.
[1584013906] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1584013935] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1584013935] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 86.7% (21249128 kB) free.
[1584013960] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 77.1% (37972008 kB) free.
[1584014003] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1584014008] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1076) < 103.255.146.164 [12/Mar/2020:17:15:16 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584014026] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 347 processes
[1584014032] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.0% (45828044 kB) free.
[1584014062] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.5% (20950368 kB) free.
[1584014122] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 80.9% (19834428 kB) free.
[1584014123] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1584014128] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584014146] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 350 processes
[1584014244] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1584014355] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 706 processes
[1584014383] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (215384 kB) free!
[1584014475] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 709 processes
[1584014503] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.1% (1640512 kB) free!
[1584014542] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191028 kB) free!
[1584014595] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 705 processes
[1584014623] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.3% (3621460 kB) free.
[1584014662] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166144 kB) free!
[1584014669] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.4% (21667004 kB) free.
[1584014707] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197432 kB) free!
[1584014728] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177148 kB) free!
[1584014746] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1584014753] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.4% (3689500 kB) free!
[1584014827] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174660 kB) free!
[1584014844] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1584014848] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176652 kB) free!
[1584014873] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.3% (3426868 kB) free!
[1584014947] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176552 kB) free!
[1584014968] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164140 kB) free!
[1584014993] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.1% (3168260 kB) free!
[1584015163] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.3% (8055228 kB) free!
[1584015283] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.8% (7784380 kB) free!
[1584015403] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.4% (6589252 kB) free!
[1584015554] SERVICE ALERT: primary-server;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 98221 MB (43% inode=99%):
[1584015593] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1952812 kB) free!
[1584015777] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (21277920 kB) free!
[1584015823] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (1614552 kB) free!
[1584015897] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.6% (17943820 kB) free!
[1584015943] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.2% (3446796 kB) free.
[1584016017] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.1% (20756972 kB) free.
[1584016017] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.7% (47014328 kB) free.
[1584016077] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.8% (15667992 kB) free.
[1584016129] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1584016129] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.5% (20217884 kB) free.
[1584016617] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584016665] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 16.1% (3984008 kB) free!
[1584016785] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 15.0% (3691120 kB) free!
[1584016859] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 338 processes
[1584016905] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 13.9% (3425576 kB) free!
[1584016979] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1584017051] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.5% (2684760 kB) free!
[1584017143] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (1182764 kB) free!
[1584017171] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.5% (2686900 kB) free!
[1584017203] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2968060 kB) free!
[1584017263] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3354600 kB) free.
[1584017291] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 16.6% (2693508 kB) free!
[1584017311] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.0% (38942084 kB) free.
[1584017359] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584017505] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.2% (22747468 kB) free.
[1584017579] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1584017699] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1584017817] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21365328 kB) free!
[1584017863] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (739232 kB) free!
[1584017891] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1584017891] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (171328 kB) free!
[1584017937] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.5% (17737184 kB) free!
[1584017984] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.5% (2200736 kB) free!
[1584018057] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14366660 kB) free!
[1584018104] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.4% (3971908 kB) free.
[1584018239] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.6% (14080548 kB) free.
[1584018429] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189264 kB) free!
[1584018516] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.3% (7547748 kB) free!
[1584018549] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171464 kB) free!
[1584018636] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 14.4% (7116348 kB) free!
[1584018657] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27778544 kB) free.
[1584018669] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166384 kB) free!
[1584018756] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.9% (6853320 kB) free!
[1584019257] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9646108 kB) free!
[1584019304] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.6% (1886068 kB) free!
[1584019378] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6311340 kB) free!
[1584019424] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.0% (3731648 kB) free.
[1584019498] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37680264 kB) free.
[1584019528] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1174) < 103.255.146.164 [12/Mar/2020:18:48:20 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584019648] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584019956] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4644464 kB) free!
[1584020064] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4747672 kB) free!
[1584020098] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16037180 kB) free!
[1584020184] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.6% (4549948 kB) free!
[1584020218] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12238596 kB) free!
[1584020304] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.9% (4381480 kB) free!
[1584020338] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (8813288 kB) free!
[1584020624] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (965464 kB) free!
[1584020744] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.0% (2600236 kB) free!
[1584020864] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.4% (4134152 kB) free.
[1584020938] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1584020938] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.2% (24043720 kB) free!
[1584021538] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (4853808 kB) free!
[1584021795] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584022064] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (357152 kB) free!
[1584022104] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1584022104] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (1531404 kB) free!
[1584022184] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (1240608 kB) free!
[1584022304] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1398568 kB) free!
[1584022904] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.0% (2437972 kB) free!
[1584023072] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584024104] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.3% (3631932 kB) free.
[1584024272] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.1% (4676320 kB) free!
[1584024392] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.8% (4604028 kB) free!
[1584024512] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.5% (4527152 kB) free!
[1584024538] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23250084 kB) free!
[1584025138] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (431824 kB) free!
[1584025738] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28133044 kB) free.
[1584026288] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584026333] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584026338] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (1798292 kB) free!
[1584026459] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (386144 kB) free!
[1584026504] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (3067180 kB) free!
[1584026579] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (414388 kB) free!
[1584026624] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.5% (3824132 kB) free.
[1584027069] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1584027179] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.9% (20994132 kB) free!
[1584027224] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.3% (2814496 kB) free!
[1584027344] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3479288 kB) free.
[1584027779] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409568 kB) free!
[1584027944] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.5% (2198472 kB) free!
[1584028064] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (2900876 kB) free!
[1584028184] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3634632 kB) free.
[1584028379] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26584584 kB) free.
[1584028712] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2299440 kB) free!
[1584028784] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.6% (1877788 kB) free!
[1584028904] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.2% (2473428 kB) free!
[1584028979] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (3024608 kB) free!
[1584029024] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (2840308 kB) free!
[1584029042] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584029099] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (512788 kB) free!
[1584029219] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398768 kB) free!
[1584029624] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3529736 kB) free.
[1584029819] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17811844 kB) free!
[1584030195] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584030315] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1584030419] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (516384 kB) free!
[1584030756] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584031019] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17852264 kB) free!
[1584031619] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (482944 kB) free!
[1584032219] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (18486344 kB) free!
[1584032304] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584032819] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390800 kB) free!
[1584033419] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.9% (23571796 kB) free!
[1584034020] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (3548096 kB) free!
[1584034620] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.4% (25636636 kB) free!
[1584035220] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (8234952 kB) free!
[1584037620] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33207348 kB) free.