[1590777120] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(7974) < 182.18.144.246 [30/May/2020:00:02:00 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1590777160] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.9% (3155520 kB) free!
[1590777164] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.6% (3086716 kB) free!
[1590777241] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(6340) < 182.18.144.246 [30/May/2020:00:04:00 +0530] TCP 502 0 0 0.001 "10.147.212.61:4000, 10.147.212.38:4000, 10.147.212.37:4000, 10.147.212.49:4000" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 0.000, 0.001, 0.000"
[1590777253] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(16165) < 2020/05/30 00:04:12 [error] 22132#0: *91854731 no live upstreams while connecting to upstream, client: 182.18.144.246, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1590777278] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397912 kB) free!
[1590777333] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(53163) < 103.255.146.164 [30/May/2020:00:05:32 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1590777348] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1590777361] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;HARD;3;(3009) < 172.105.57.57 [30/May/2020:00:05:11 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1590777373] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2592) < 2020/05/30 00:05:35 [error] 22121#0: *91857721 recv() failed (104: Connection reset by peer) while proxying connection, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:37/0, bytes from/to upstream:0/37
[1590777453] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1035) < 185.255.8.58 [30/May/2020:00:07:25 +0530] TCP 502 0 0 0.002 "10.147.212.143:13001, 10.147.212.154:13001, 10.147.212.155:13001, 10.147.212.142:13001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 0.000, 0.000, 0.001"
[1590777493] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590777573] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(518) < 103.255.146.164 [30/May/2020:00:08:32 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590777760] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199804 kB) free!
[1590777765] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189640 kB) free!
[1590777878] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26604884 kB) free.
[1590777909] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3180948 kB) free!
[1590777961] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590778029] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3761536 kB) free.
[1590778174] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590778478] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (392544 kB) free!
[1590778598] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (402712 kB) free!
[1590778629] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3166988 kB) free!
[1590778718] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399704 kB) free!
[1590778749] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.9% (3718428 kB) free.
[1590779318] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16055068 kB) free!
[1590779349] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590779918] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394488 kB) free!
[1590779949] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (863484 kB) free!
[1590780069] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (1480996 kB) free!
[1590780189] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.4% (1860348 kB) free!
[1590780719] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 2.8% (678332 kB) free!
[1590781718] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19479592 kB) free!
[1590781853] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42191 MB (18.77% inode=100%):
[1590781913] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39124 MB (17.40% inode=100%):
[1590781973] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;3;DISK WARNING - free space: /home 39667 MB (17.65% inode=100%):
[1590782033] SERVICE ALERT: webserver2;HDD_Home status;WARNING;HARD;4;DISK WARNING - free space: /home 39667 MB (17.65% inode=100%):
[1590782220] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 88.0% (21566824 kB) free.
[1590782318] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9556320 kB) free!
[1590782318] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386644 kB) free!
[1590782333] SERVICE ALERT: webserver2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 107147 MB (47.67% inode=100%):
[1590782438] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.4% (9553364 kB) free!
[1590782558] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.4% (9551232 kB) free!
[1590782589] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (919700 kB) free!
[1590782626] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.0% (25090520 kB) free!
[1590782746] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.1% (25127952 kB) free!
[1590782866] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.1% (25153784 kB) free!
[1590783189] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.1% (1973196 kB) free!
[1590783466] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.5% (28276860 kB) free.
[1590784066] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.7% (22021304 kB) free!
[1590784118] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14638624 kB) free!
[1590784186] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 16.8% (22088324 kB) free!
[1590784306] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.8% (22128576 kB) free!
[1590784358] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 36.2% (17819248 kB) free.
[1590784718] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (420520 kB) free!
[1590784906] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1590784906] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.5% (29710408 kB) free.
[1590784965] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.5% (3564336 kB) free!
[1590785320] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.80, 5.33, 5.09
[1590785440] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.45, 5.38, 5.14
[1590785560] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.92, 5.59, 5.24
[1590785565] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1590785565] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (1075608 kB) free!
[1590786160] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.23, 4.61, 4.86
[1590786511] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 92.2% (45461820 kB) free.
[1590787555] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.8% (43267480 kB) free.
[1590788162] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(449) < 103.16.101.72 [30/May/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1590788284] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590788318] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28581108 kB) free.
[1590788374] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4535) < 172.105.59.173 [30/May/2020:03:02:32 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1590788494] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590788918] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (407364 kB) free!
[1590789038] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (385844 kB) free!
[1590789158] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408824 kB) free!
[1590789189] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1590789189] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.5% (3817172 kB) free.
[1590789758] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (15048032 kB) free!
[1590790358] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401232 kB) free!
[1590791561] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.1% (4686936 kB) free!
[1590791589] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.9% (1935560 kB) free!
[1590791709] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.7% (1747208 kB) free!
[1590791829] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.4% (3807860 kB) free.
[1590792761] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (204532 kB) free!
[1590793958] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.9% (26225712 kB) free!
[1590794558] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (3508584 kB) free!
[1590794561] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 38.9% (9544368 kB) free.
[1590794845] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 37.2% (18353604 kB) free.
[1590795158] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.4% (19020424 kB) free!
[1590795161] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (213932 kB) free!
[1590795281] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (161016 kB) free!
[1590795401] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1590795401] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195396 kB) free!
[1590795429] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1590795445] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (264296 kB) free!
[1590795565] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (258984 kB) free!
[1590795685] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (255996 kB) free!
[1590795758] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403680 kB) free!
[1590795792] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (368148 kB) free!
[1590796306] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590796358] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (13984504 kB) free!
[1590796924] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 39.1% (9654640 kB) free.
[1590796958] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30441816 kB) free.
[1590796965] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 43.2% (10588936 kB) free.
[1590797205] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 71.7% (17575620 kB) free.
[1590797558] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10509236 kB) free!
[1590797678] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (6696828 kB) free!
[1590797798] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34183496 kB) free.
[1590798398] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12971064 kB) free!
[1590798518] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8767888 kB) free!
[1590798638] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (4635212 kB) free!
[1590799239] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16531568 kB) free!
[1590799839] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29216108 kB) free.
[1590799920] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 15.3% (3752916 kB) free!
[1590799980] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 14.3% (3515460 kB) free!
[1590800040] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 14.3% (3514784 kB) free!
[1590800100] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 14.3% (3511036 kB) free!
[1590800162] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1590800282] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1590800439] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9864188 kB) free!
[1590800559] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6044828 kB) free!
[1590800679] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33360104 kB) free.
[1590801279] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13393592 kB) free!
[1590801399] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (9484444 kB) free!
[1590801519] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (5850204 kB) free!
[1590802119] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15272204 kB) free!
[1590802968] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1590803319] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (5203224 kB) free!
[1590803524] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187692 kB) free!
[1590803644] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195552 kB) free!
[1590803764] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196536 kB) free!
[1590803919] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (14993408 kB) free!
[1590804519] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26539824 kB) free.
[1590804792] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 11.0% (1784816 kB) free!
[1590805119] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5765604 kB) free!
[1590805239] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (1967528 kB) free!
[1590805359] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30078036 kB) free.
[1590805960] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (7001888 kB) free!
[1590806080] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (2487708 kB) free!
[1590806200] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.0% (25066868 kB) free!
[1590806800] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (11072212 kB) free!
[1590806801] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590807400] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (22749448 kB) free!
[1590808000] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (1012384 kB) free!
[1590808600] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13533484 kB) free!
[1590808968] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4899060 kB) free!
[1590809088] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (4807632 kB) free!
[1590809208] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (4733432 kB) free!
[1590809800] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1636408 kB) free!
[1590810160] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 16.28, 12.25, 7.86
[1590810280] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.72, 10.77, 7.85
[1590810400] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18343148 kB) free!
[1590810400] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.98, 10.97, 8.28
[1590810408] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2136416 kB) free!
[1590811000] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399208 kB) free!
[1590812200] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28999516 kB) free.
[1590812800] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (3869836 kB) free!
[1590812829] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.6% (1728356 kB) free!
[1590812921] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (387328 kB) free!
[1590812949] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.8% (3870324 kB) free.
[1590813041] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408068 kB) free!
[1590813549] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3172788 kB) free!
[1590813669] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3544592 kB) free.
[1590813952] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1590814005] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186712 kB) free!
[1590814051] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1590814072] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1590814125] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186460 kB) free!
[1590814171] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1590814192] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1590814241] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.1% (22480236 kB) free!
[1590814246] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196504 kB) free!
[1590814269] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (3160204 kB) free!
[1590814292] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1590814389] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.5% (3657080 kB) free.
[1590814792] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1590814841] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382500 kB) free!
[1590814892] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1590814990] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3124424 kB) free!
[1590815110] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3284444 kB) free.
[1590815179] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 62.1% (15217752 kB) free.
[1590815392] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1590815441] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.8% (22130392 kB) free!
[1590815512] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1590815633] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1590816041] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405156 kB) free!
[1590816310] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (3062668 kB) free!
[1590816400] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 37.29, 18.38, 11.55
[1590816430] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.8% (2728520 kB) free!
[1590816516] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.3% (3891624 kB) free!
[1590816550] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.1% (2621272 kB) free!
[1590816636] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.6% (3279284 kB) free!
[1590816641] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19607032 kB) free!
[1590816756] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.8% (2650652 kB) free!
[1590817008] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.2% (2739712 kB) free!
[1590817241] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392812 kB) free!
[1590817356] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (195428 kB) free!
[1590817555] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9382768 kB) free!
[1590817608] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1590817608] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (1484492 kB) free!
[1590817675] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9076680 kB) free!
[1590817795] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.3% (8538480 kB) free!
[1590817841] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23041744 kB) free!
[1590817929] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 17.45, 23.96, 22.15
[1590818177] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 942 processes
[1590818297] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 939 processes
[1590818417] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 939 processes
[1590818441] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401508 kB) free!
[1590818529] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 17.26, 16.31, 18.39
[1590819041] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.0% (21113264 kB) free!
[1590819129] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 32.94, 26.62, 22.34
[1590819233] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1590819233] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1590819379] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.6% (4554912 kB) free!
[1590819494] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 12.0% (4204948 kB) free!
[1590819499] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.9% (4377508 kB) free!
[1590819614] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 10.4% (3641532 kB) free!
[1590819619] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.2% (4208788 kB) free!
[1590819641] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (435004 kB) free!
[1590819734] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (3105512 kB) free!
[1590820093] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/30 11:50:26 [error] 22131#0: *91976078 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590820195] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4266520 kB) free!
[1590820213] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590820217] SERVICE ALERT: qservercluster2;Total Processes;OK;HARD;1;PROCS OK: 619 processes
[1590820241] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (25877040 kB) free!
[1590820334] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 91.3% (31974060 kB) free.
[1590820356] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.6% (16068064 kB) free.
[1590820620] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.1218.945> (10.147.212.158:49630 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590820740] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590820749] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 13.4% (3315248 kB) free!
[1590820818] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 989 processes
[1590820819] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2003768 kB) free!
[1590820841] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (414252 kB) free!
[1590820869] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 11.9% (2927884 kB) free!
[1590820938] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 987 processes
[1590820989] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 10.5% (2587888 kB) free!
[1590821058] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 984 processes
[1590821351] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1231236 kB) free!
[1590821442] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.2% (25253992 kB) free!
[1590821589] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.8% (22898332 kB) free.
[1590821951] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.8% (1915808 kB) free!
[1590822000] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 1.0% (254568 kB) free!
[1590822042] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (429948 kB) free!
[1590822106] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.8% (26037100 kB) free!
[1590822129] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1590822129] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 20.29, 19.69, 19.45
[1590822226] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.4% (19044288 kB) free!
[1590822346] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.6% (28509252 kB) free.
[1590822551] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1590822551] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.7% (3857364 kB) free.
[1590822642] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34027344 kB) free.
[1590823242] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4372276 kB) free!
[1590823330] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 25.81, 21.30, 20.24
[1590823362] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (377240 kB) free!
[1590823482] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (425376 kB) free!
[1590823546] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.5% (24407876 kB) free!
[1590823666] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.9% (32846728 kB) free.
[1590823751] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3038212 kB) free!
[1590823871] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.0% (2768888 kB) free!
[1590823991] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (1682844 kB) free!
[1590824082] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19790444 kB) free!
[1590824682] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (426628 kB) free!
[1590824866] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13042852 kB) free!
[1590824986] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.1% (22480612 kB) free!
[1590825106] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.8% (15561788 kB) free!
[1590825282] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17993388 kB) free!
[1590825400] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 12.34, 13.52, 19.26
[1590825695] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 172.105.59.173 [30/May/2020:13:23:24 +0530] TCP 502 0 0 120.000 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1590825791] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1558048 kB) free!
[1590825815] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590825882] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384876 kB) free!
[1590826000] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 27.21, 21.88, 20.52
[1590826330] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 13.36, 17.75, 19.98
[1590826600] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1590826600] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 15.39, 17.62, 18.95
[1590826711] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8194232 kB) free!
[1590826831] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7864792 kB) free!
[1590826951] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 15.3% (7542792 kB) free!
[1590826956] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.4% (3239060 kB) free!
[1590827076] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.4% (2784848 kB) free!
[1590827196] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.3% (2317396 kB) free!
[1590827200] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 28.83, 22.04, 20.22
[1590827282] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 335 processes
[1590827340] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.1081.915> (10.147.212.147:57508 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590827402] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1590827461] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590827591] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.2% (1986792 kB) free!
[1590827592] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (528492 kB) free!
[1590827633] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1590827797] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (24.9% change >= 20.0% threshold)
[1590827797] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (206956 kB) free!
[1590828130] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 22.84, 21.70, 20.09
[1590828258] SERVICE ALERT: qservercluster2;Total Processes;OK;HARD;1;PROCS OK: 618 processes
[1590828283] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14823188 kB) free!
[1590828660] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.4926.980> (10.147.212.158:35742 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590828706] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.4% (26936188 kB) free.
[1590828751] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3727576 kB) free!
[1590828780] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590828858] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 960 processes
[1590828883] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403652 kB) free!
[1590828978] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 899 processes
[1590829008] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590829098] SERVICE FLAPPING ALERT: qservercluster2;Total Processes;STARTED; Service appears to have started flapping (25.3% change >= 20.0% threshold)
[1590829098] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 894 processes
[1590829330] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 27.11, 20.41, 19.78
[1590829391] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (294376 kB) free!
[1590829483] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.0% (23682872 kB) free!
[1590829698] SERVICE ALERT: qservercluster2;Total Processes;WARNING;HARD;3;PROCS WARNING: 843 processes
[1590829802] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 344 processes
[1590829906] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 11.0% (14546784 kB) free!
[1590829922] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1590829934] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 14.8% (5168344 kB) free!
[1590829980] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.16742.981> (10.147.212.158:45400 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590829991] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.3% (2323964 kB) free!
[1590830027] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.6% (27126408 kB) free.
[1590830042] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 336 processes
[1590830054] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 14.1% (4949272 kB) free!
[1590830083] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (415912 kB) free!
[1590830100] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590830174] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 13.6% (4759216 kB) free!
[1590830299] SERVICE ALERT: qservercluster2;Total Processes;OK;HARD;1;PROCS OK: 613 processes
[1590830626] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9213348 kB) free!
[1590830642] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1590830683] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31704296 kB) free.
[1590830701] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.8736.984> (10.147.212.158:49794 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590830746] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.0% (23777096 kB) free!
[1590830800] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.48, 6.34, 13.85
[1590830821] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590830866] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (12646592 kB) free!
[1590831013] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/30 14:52:27 [error] 22133#0: *92033750 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590831133] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590831215] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(108) < 185.255.8.58 [30/May/2020:15:00:03 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1590831283] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (690788 kB) free!
[1590831335] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590831374] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 91.2% (31940780 kB) free.
[1590831397] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.1% (16199188 kB) free.
[1590831403] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393236 kB) free!
[1590831421] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.31337.951> (10.147.212.158:54256 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590831466] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1590831466] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.4% (21624408 kB) free!
[1590831523] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (423648 kB) free!
[1590831541] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590832066] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.4% (29521516 kB) free.
[1590832123] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19289488 kB) free!
[1590832723] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381800 kB) free!
[1590832991] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (753548 kB) free!
[1590833266] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 11.6% (15344824 kB) free!
[1590833289] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.16945.6406> (10.147.212.53:48020 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1590833323] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.5% (24404240 kB) free!
[1590833386] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.6% (27160332 kB) free.
[1590833411] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590833591] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (2577996 kB) free!
[1590833923] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (415256 kB) free!
[1590834523] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33076752 kB) free.
[1590835123] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (1785216 kB) free!
[1590835243] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (391896 kB) free!
[1590835363] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396960 kB) free!
[1590835535] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [30/May/2020:16:14:21 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1590835655] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590835964] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28938368 kB) free.
[1590835991] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (926256 kB) free!
[1590836033] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1590836153] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 347 processes
[1590836273] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1590836342] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.1400.920> (10.147.212.158:57990 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590836386] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 10.3% (13642496 kB) free!
[1590836462] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590836506] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.6% (25828096 kB) free!
[1590836564] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (400472 kB) free!
[1590836591] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (3115092 kB) free!
[1590836626] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.4% (17722772 kB) free!
[1590836684] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (409732 kB) free!
[1590836804] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (458452 kB) free!
[1590837404] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20491504 kB) free!
[1590838004] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387388 kB) free!
[1590838056] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [30/May/2020:16:54:32 +0530] TCP 502 0 0 60.001 "10.147.212.155:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.000, 60.001, 0.000"
[1590838074] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1590838176] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590838604] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29825408 kB) free.
[1590839197] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.1% (2935872 kB) free!
[1590839200] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1590839204] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (436252 kB) free!
[1590839317] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.5% (2581756 kB) free!
[1590839324] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (437360 kB) free!
[1590839437] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2143084 kB) free!
[1590839444] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (411348 kB) free!
[1590839591] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3676632 kB) free.
[1590839874] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1590839994] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 350 processes
[1590840044] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.0% (22427828 kB) free!
[1590840114] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 350 processes
[1590840130] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1590840192] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (563100 kB) free!
[1590840312] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (1513316 kB) free!
[1590840432] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.8% (2560332 kB) free!
[1590840644] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (411096 kB) free!
[1590841244] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33569860 kB) free.
[1590841632] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3295720 kB) free.
[1590841699] SERVICE FLAPPING ALERT: qservercluster2;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590841837] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.2% (16209560 kB) free.
[1590841844] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (390424 kB) free!
[1590841964] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (378280 kB) free!
[1590842085] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393504 kB) free!
[1590842685] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26546696 kB) free.
[1590842832] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (1012756 kB) free!
[1590842952] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.6% (2377828 kB) free!
[1590843072] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3543076 kB) free.
[1590843285] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (439440 kB) free!
[1590843405] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (404916 kB) free!
[1590843525] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (417032 kB) free!
[1590843672] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (252636 kB) free!
[1590843792] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.7% (1907288 kB) free!
[1590843912] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.3% (3458632 kB) free.
[1590844125] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27513100 kB) free.
[1590844426] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1590844442] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1590844512] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (521768 kB) free!
[1590844562] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1590844632] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.8% (2247360 kB) free!
[1590844725] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (423540 kB) free!
[1590844753] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.2% (3606204 kB) free.
[1590844845] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (450724 kB) free!
[1590844965] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406768 kB) free!
[1590845514] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1590845514] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1590845565] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19772584 kB) free!
[1590845952] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (1554176 kB) free!
[1590846072] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (3126240 kB) free!
[1590846165] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (412676 kB) free!
[1590846192] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3753952 kB) free.
[1590846765] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34961380 kB) free.
[1590846792] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (695392 kB) free!
[1590846912] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.4% (2333532 kB) free!
[1590847032] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.9% (3718968 kB) free.
[1590847314] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1590847365] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (732660 kB) free!
[1590847486] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (394348 kB) free!
[1590847606] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (404648 kB) free!
[1590847632] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (251468 kB) free!
[1590847752] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.8% (1761992 kB) free!
[1590847872] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (3178608 kB) free!
[1590848162] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1590848206] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31277212 kB) free.
[1590848282] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1590848473] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.2% (3605388 kB) free.
[1590848800] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.09, 4.57, 4.78
[1590848806] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (416456 kB) free!
[1590848926] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (383172 kB) free!
[1590849046] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (419940 kB) free!
[1590849073] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (1289180 kB) free!
[1590849193] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (3129004 kB) free!
[1590849313] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.0% (3742904 kB) free.
[1590849646] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26091548 kB) free!
[1590849826] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.6% (33778520 kB) free.
[1590849913] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (781600 kB) free!
[1590850033] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.6% (2381220 kB) free!
[1590850153] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3638044 kB) free.
[1590850246] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384984 kB) free!
[1590850753] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.4% (1690740 kB) free!
[1590850837] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.3% (3882380 kB) free!
[1590850873] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3336936 kB) free.
[1590850957] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.3% (3659956 kB) free!
[1590851026] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.7% (26021024 kB) free!
[1590851077] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.5% (3474980 kB) free!
[1590851146] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.1% (18587060 kB) free!
[1590851266] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.2% (24042812 kB) free!
[1590852073] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.3% (1672348 kB) free!
[1590852193] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (3202844 kB) free!
[1590852277] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (1206860 kB) free!
[1590852314] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.5% (3816056 kB) free.
[1590852466] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 28.6% (37642076 kB) free.
[1590852646] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13617536 kB) free!
[1590852877] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.1% (16185312 kB) free.
[1590852913] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (942996 kB) free!
[1590853033] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.9% (2422108 kB) free!
[1590853153] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.3% (3790348 kB) free.
[1590853246] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402636 kB) free!
[1590853682] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1590853802] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 332 processes
[1590853846] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21274156 kB) free!
[1590853922] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1590854152] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46113 MB (20% inode=99%):
[1590854266] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.7% (26002784 kB) free!
[1590854272] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46113 MB (20% inode=99%):
[1590854353] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (1249280 kB) free!
[1590854386] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 27.4% (36141696 kB) free.
[1590854392] SERVICE ALERT: primary-server;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46113 MB (20% inode=99%):
[1590854446] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388184 kB) free!
[1590854473] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.6% (2530840 kB) free!
[1590854515] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1590854593] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3515096 kB) free.
[1590855046] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29090192 kB) free.
[1590855122] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1590855242] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1590855647] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (2243816 kB) free!
[1590855767] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (379308 kB) free!
[1590855793] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (1032312 kB) free!
[1590855887] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (456272 kB) free!
[1590855913] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.0% (2107852 kB) free!
[1590856033] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3268104 kB) free.
[1590856487] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19663916 kB) free!
[1590856633] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3236364 kB) free!
[1590856753] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (3233172 kB) free!
[1590856873] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1119452 kB) free!
[1590856892] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1590856915] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1590857012] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1590857035] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1590857087] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385808 kB) free!
[1590857132] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1590857155] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1590857392] SERVICE ALERT: primary-server;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 72070 MB (32% inode=99%):
[1590857474] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.9% (1933440 kB) free!
[1590857687] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.8% (23436092 kB) free!
[1590857733] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1590857755] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1590858287] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394716 kB) free!
[1590858619] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 65.3% (16013040 kB) free.
[1590858887] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27451408 kB) free.
[1590859487] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (2233600 kB) free!
[1590859607] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (399448 kB) free!
[1590859728] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (454548 kB) free!
[1590859874] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1404092 kB) free!
[1590860328] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17846604 kB) free!
[1590860928] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380856 kB) free!
[1590861074] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (2946460 kB) free!
[1590861528] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19057568 kB) free!
[1590862128] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (412388 kB) free!
[1590862274] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (1479488 kB) free!
[1590862728] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21283772 kB) free!
[1590862874] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.5% (4799652 kB) free.
[1590863328] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394412 kB) free!