[1578594621] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11648936 kB) free!
[1578594675] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.7% (4817960 kB) free!
[1578594689] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (10439868 kB) free!
[1578594741] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9179652 kB) free!
[1578594809] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (6670644 kB) free!
[1578594861] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (11457068 kB) free!
[1578595069] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.4% (30243732 kB) free.
[1578595409] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18846752 kB) free!
[1578596009] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31613704 kB) free.
[1578596609] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9195420 kB) free!
[1578596729] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5436304 kB) free!
[1578596809] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184624 kB) free!
[1578596849] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36203492 kB) free.
[1578596869] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196376 kB) free!
[1578596929] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (173484 kB) free!
[1578596990] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (208104 kB) free!
[1578597290] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 72.1% (17662284 kB) free.
[1578597449] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11465900 kB) free!
[1578597569] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7616840 kB) free!
[1578597689] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36967492 kB) free.
[1578597979] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1578598289] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13585884 kB) free!
[1578598409] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (9558192 kB) free!
[1578598529] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (5893144 kB) free!
[1578599061] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.9% (17016080 kB) free!
[1578599129] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (15906604 kB) free!
[1578600330] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36579128 kB) free.
[1578600379] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.88, 5.82, 5.05
[1578600500] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.86, 5.91, 5.18
[1578600620] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.35, 5.71, 5.20
[1578600930] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13691564 kB) free!
[1578601050] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (9907240 kB) free!
[1578601170] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36518768 kB) free.
[1578601770] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14325820 kB) free!
[1578601890] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (10593356 kB) free!
[1578602010] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36470348 kB) free.
[1578602420] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.70, 4.40, 4.67
[1578602610] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14581624 kB) free!
[1578602730] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11044848 kB) free!
[1578602850] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36398144 kB) free.
[1578602998] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (4392596 kB) free!
[1578603119] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (223824 kB) free!
[1578603197] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.5% (10409388 kB) free.
[1578603239] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (221956 kB) free!
[1578603352] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (239876 kB) free!
[1578603450] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15208528 kB) free!
[1578603472] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (244476 kB) free!
[1578603570] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (11272264 kB) free!
[1578603592] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (232228 kB) free!
[1578603690] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36309288 kB) free.
[1578603839] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 26.3% (12963148 kB) free.
[1578604220] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.77, 5.65, 5.31
[1578604290] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14871252 kB) free!
[1578604340] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.84, 6.17, 5.55
[1578604410] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11091752 kB) free!
[1578604460] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.68, 6.74, 5.83
[1578604531] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36223476 kB) free.
[1578604792] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.9% (43340316 kB) free.
[1578605130] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14299512 kB) free!
[1578605250] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (10594988 kB) free!
[1578605370] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35665104 kB) free.
[1578605457] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1578605577] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1578605970] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12864092 kB) free!
[1578606090] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (9357520 kB) free!
[1578606210] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32792100 kB) free.
[1578606810] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12094784 kB) free!
[1578606860] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1578606860] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.96, 4.24, 4.80
[1578606930] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (36030688 kB) free.
[1578607530] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14406752 kB) free!
[1578607650] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (10846668 kB) free!
[1578607770] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34389380 kB) free.
[1578608370] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13543900 kB) free!
[1578608490] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (26972108 kB) free.
[1578608660] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.47, 5.66, 5.36
[1578608780] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.44, 5.57, 5.37
[1578608900] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.73, 5.34, 5.31
[1578609090] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (16010312 kB) free!
[1578609210] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (12366884 kB) free!
[1578609331] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35847760 kB) free.
[1578609930] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13643836 kB) free!
[1578610050] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35758172 kB) free.
[1578610650] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14845928 kB) free!
[1578610770] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35720640 kB) free.
[1578611370] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15324900 kB) free!
[1578611490] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35671464 kB) free.
[1578612090] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16196636 kB) free!
[1578612210] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (12662624 kB) free!
[1578612330] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33708204 kB) free.
[1578612930] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13302868 kB) free!
[1578613050] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34648892 kB) free.
[1578613650] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13598796 kB) free!
[1578613770] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34876916 kB) free.
[1578614300] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.04, 3.00, 3.92
[1578614370] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13907560 kB) free!
[1578614491] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1578614491] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34914332 kB) free.
[1578615090] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7864084 kB) free!
[1578615210] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35009116 kB) free.
[1578615810] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10327048 kB) free!
[1578615930] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34946444 kB) free.
[1578616398] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186908 kB) free!
[1578616518] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167656 kB) free!
[1578616530] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10202380 kB) free!
[1578616638] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (231640 kB) free!
[1578616650] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34847500 kB) free.
[1578616977] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 408 processes
[1578617061] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 31.1% (41040452 kB) free.
[1578617098] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 408 processes
[1578617250] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11071656 kB) free!
[1578617370] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6531840 kB) free!
[1578617490] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31986896 kB) free.
[1578617522] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 66.0% (16177012 kB) free.
[1578618090] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7064388 kB) free!
[1578618210] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32480076 kB) free.
[1578618675] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2124732 kB) free!
[1578618810] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10599912 kB) free!
[1578618930] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (33954632 kB) free.
[1578619530] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17920744 kB) free!
[1578619650] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34188216 kB) free.
[1578620849] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.6% (21704904 kB) free.
[1578620851] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16753228 kB) free!
[1578620883] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 57.3% (14140840 kB) free.
[1578620900] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1578620923] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (8934024 kB) free!
[1578620926] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 48.5% (23889568 kB) free.
[1578620957] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 75.2% (18438456 kB) free.
[1578620971] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35978784 kB) free.
[1578621297] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 399 processes
[1578621417] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1578621486] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186248 kB) free!
[1578621526] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (264280 kB) free!
[1578621571] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15934248 kB) free!
[1578621606] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196364 kB) free!
[1578621646] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (211824 kB) free!
[1578621691] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35255764 kB) free.
[1578621726] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164868 kB) free!
[1578621766] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (242492 kB) free!
[1578622891] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15849156 kB) free!
[1578623011] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34884336 kB) free.
[1578624211] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14876376 kB) free!
[1578624331] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34729800 kB) free.
[1578624931] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15275204 kB) free!
[1578625051] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34394728 kB) free.
[1578625651] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13163424 kB) free!
[1578625772] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33141860 kB) free.
[1578626371] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12249732 kB) free!
[1578626491] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32497448 kB) free.
[1578626879] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(6) < 2020/01/10 08:57:52 [error] 21683#0: *23994525 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1578626999] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(21) < 2020/01/10 08:59:59 [error] 21683#0: *23994678 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1578627050] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 438 processes
[1578627091] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11149440 kB) free!
[1578627119] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(19) < 2020/01/10 09:01:53 [error] 21681#0: *23994731 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1578627170] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 911 processes
[1578627211] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31167724 kB) free.
[1578627388] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1094.33, 1076.99, 602.37
[1578627505] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1133.03, 444.33, 165.62
[1578627508] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 149.34, 721.36, 529.61
[1578627625] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1272.53, 725.59, 302.48
[1578627628] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 22.38, 483.51, 465.69
[1578627745] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1291.51, 913.87, 422.78
[1578627811] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8034404 kB) free!
[1578627931] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30882368 kB) free.
[1578628101] SERVICE ALERT: pservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 600 processes
[1578628221] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 343 processes
[1578628228] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 652 processes
[1578628345] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.90, 248.94, 350.86
[1578628348] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 651 processes
[1578628468] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 309 processes
[1578628531] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9316460 kB) free!
[1578628651] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30873168 kB) free.
[1578628945] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.38, 36.05, 185.19
[1578629251] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8263684 kB) free!
[1578629371] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (29995480 kB) free.
[1578629519] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578629971] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8198000 kB) free!
[1578630091] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4185824 kB) free!
[1578630211] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23008620 kB) free!
[1578630575] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 52.5% (12957860 kB) free.
[1578630635] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.4% (7602312 kB) free!
[1578630811] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410316 kB) free!
[1578631178] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191248 kB) free!
[1578631236] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (24.4% change >= 20.0% threshold)
[1578631236] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (212628 kB) free!
[1578631298] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166676 kB) free!
[1578631418] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1578631418] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (646312 kB) free!
[1578631743] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(49) < 103.255.146.164 [10/Jan/2020:10:13:38 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1578631864] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578631920] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(61) < 2020/01/10 10:13:38 [error] 21674#0: *24007255 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578632011] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.8% (23452500 kB) free!
[1578632040] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578632611] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1578632611] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28959672 kB) free.
[1578633211] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6604148 kB) free!
[1578633298] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.2% (22503156 kB) free.
[1578633331] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27525636 kB) free.
[1578633394] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.5% (40155092 kB) free.
[1578633907] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180456 kB) free!
[1578633997] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (2083400 kB) free!
[1578634027] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163488 kB) free!
[1578634117] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236800 kB) free!
[1578634147] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183892 kB) free!
[1578634237] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (256016 kB) free!
[1578634331] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.3% (4741004 kB) free!
[1578634451] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.1% (4680908 kB) free!
[1578634571] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 18.9% (4637456 kB) free!
[1578634868] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.3% (22293464 kB) free.
[1578634902] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.4% (40124512 kB) free.
[1578635477] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187900 kB) free!
[1578635505] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (238208 kB) free!
[1578635597] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182268 kB) free!
[1578635625] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (215948 kB) free!
[1578635717] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185460 kB) free!
[1578635746] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (225228 kB) free!
[1578636299] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.0% (3596196 kB) free!
[1578636357] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13525472 kB) free!
[1578636419] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.3% (3207808 kB) free!
[1578636477] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (9678852 kB) free!
[1578636539] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.6% (2827576 kB) free!
[1578636597] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35620176 kB) free.
[1578636914] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.14331.1905> (10.147.212.53:55686 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1578637035] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1578637069] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9393116 kB) free!
[1578637139] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (1154856 kB) free!
[1578637189] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.1% (8928716 kB) free!
[1578637197] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12269360 kB) free!
[1578637309] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.3% (8501608 kB) free!
[1578637317] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8475644 kB) free!
[1578637323] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189624 kB) free!
[1578637329] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.2% (22023648 kB) free.
[1578637361] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.9% (40830180 kB) free.
[1578637437] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35303192 kB) free.
[1578637443] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176536 kB) free!
[1578637484] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.6% (20985556 kB) free.
[1578637542] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 75.0% (16836352 kB) free.
[1578637578] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.6% (29394388 kB) free.
[1578637623] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 90.9% (22291128 kB) free.
[1578637929] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184368 kB) free!
[1578638037] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11259572 kB) free!
[1578638049] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167160 kB) free!
[1578638086] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188396 kB) free!
[1578638157] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7316620 kB) free!
[1578638169] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183532 kB) free!
[1578638207] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171172 kB) free!
[1578638277] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33524300 kB) free.
[1578638327] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189616 kB) free!
[1578638456] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 14.9% (3683224 kB) free!
[1578638567] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236876 kB) free!
[1578638576] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 13.5% (3331708 kB) free!
[1578638580] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.8% (22651720 kB) free.
[1578638612] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.8% (40319124 kB) free.
[1578638640] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/10 12:08:28 [error] 21674#0: *24027145 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:4695/121, bytes from/to upstream:121/4695
[1578638696] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.2% (23004272 kB) free.
[1578638760] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578638877] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10097176 kB) free!
[1578638927] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.3% (22390204 kB) free.
[1578638997] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6079252 kB) free!
[1578639117] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31768056 kB) free.
[1578639186] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191456 kB) free!
[1578639215] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (259168 kB) free!
[1578639306] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191040 kB) free!
[1578639335] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (217820 kB) free!
[1578639426] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166264 kB) free!
[1578639455] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237392 kB) free!
[1578639717] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9367244 kB) free!
[1578639837] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (5208476 kB) free!
[1578639958] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1528764 kB) free!
[1578640370] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1578640482] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1578640490] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1578640602] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1578640722] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1578640727] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180604 kB) free!
[1578640768] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4735144 kB) free!
[1578640792] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9719896 kB) free!
[1578640847] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182052 kB) free!
[1578640888] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.6% (4546940 kB) free!
[1578640912] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.2% (9451288 kB) free!
[1578640967] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196292 kB) free!
[1578641008] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.0% (4408620 kB) free!
[1578641032] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.7% (9213892 kB) free!
[1578641090] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1578641158] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19054832 kB) free!
[1578641210] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1578641331] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1578641758] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2071564 kB) free!
[1578642208] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2377848 kB) free!
[1578642958] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19608688 kB) free!
[1578643122] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1578643432] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4101144 kB) free!
[1578643558] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386812 kB) free!
[1578644075] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1578644195] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1578644758] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20532688 kB) free!
[1578644930] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1578645050] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1578645170] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1578645358] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382420 kB) free!
[1578645458] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.8% (4843232 kB) free!
[1578645522] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1578645578] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.5% (4785640 kB) free!
[1578645642] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1578645698] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.4% (4744316 kB) free!
[1578645762] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1578645770] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 338 processes
[1578646370] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1578646553] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1578646558] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.0% (22459312 kB) free!
[1578646600] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.2% (22020432 kB) free.
[1578646628] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.8% (40310696 kB) free.
[1578646694] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.0% (44825900 kB) free.
[1578646731] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.9% (22765240 kB) free.
[1578646751] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1578646751] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.8% (22500092 kB) free.
[1578646793] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.5% (21942088 kB) free.
[1578646864] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(774) < 103.255.146.164 [10/Jan/2020:14:28:40 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1578646874] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 48.1% (63385256 kB) free.
[1578646932] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3199376 kB) free!
[1578646971] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1578646971] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1578646984] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578647052] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (2897920 kB) free!
[1578647102] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(48) < 2020/01/10 14:28:34 [error] 18356#0: *36432375 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:226155/3394726, bytes from/to upstream:3394726/226155
[1578647145] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.4% (4142516 kB) free!
[1578647160] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(801) < 2020/01/10 14:28:40 [error] 21676#0: *24053182 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578647172] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.6% (3990804 kB) free.
[1578647206] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187424 kB) free!
[1578647222] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578647237] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (240476 kB) free!
[1578647265] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.4% (3896780 kB) free!
[1578647267] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 185.53, 54.97, 29.08
[1578647281] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578647326] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163588 kB) free!
[1578647357] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (211676 kB) free!
[1578647385] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.1% (3830796 kB) free!
[1578647387] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 581.57, 257.38, 106.94
[1578647446] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170032 kB) free!
[1578647477] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (230288 kB) free!
[1578647507] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 682.10, 398.04, 176.27
[1578647571] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1578647691] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1578647956] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183340 kB) free!
[1578648076] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197284 kB) free!
[1578648196] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179388 kB) free!
[1578648585] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (1044844 kB) free!
[1578648839] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9398000 kB) free!
[1578648891] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1578648959] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (9142284 kB) free!
[1578648972] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.3% (2815724 kB) free!
[1578649011] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1578649079] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (8910776 kB) free!
[1578649092] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.4% (2661716 kB) free!
[1578649131] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 338 processes
[1578649172] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.5% (22340732 kB) free.
[1578649186] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 74.9% (16821748 kB) free.
[1578649212] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.0% (4876732 kB) free.
[1578649244] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 92.4% (45555120 kB) free.
[1578649274] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.5% (40640472 kB) free.
[1578649306] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.4% (22159248 kB) free.
[1578649731] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1578649775] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187376 kB) free!
[1578649877] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10927276 kB) free!
[1578649879] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.1% (6928300 kB) free!
[1578649895] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188532 kB) free!
[1578649918] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (203160 kB) free!
[1578649997] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (7045612 kB) free!
[1578649999] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (214536 kB) free!
[1578650015] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1578650015] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197636 kB) free!
[1578650039] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165804 kB) free!
[1578650117] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (3178780 kB) free!
[1578650119] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (230968 kB) free!
[1578650159] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173140 kB) free!
[1578650280] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/10 15:19:23 [error] 21682#0: *24061991 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:4697/123, bytes from/to upstream:123/4697
[1578650400] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578650717] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16359120 kB) free!
[1578650932] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1578651107] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 18.59, 19.86, 32.95
[1578651317] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395348 kB) free!
[1578651532] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 338 processes
[1578651612] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3238132 kB) free!
[1578651652] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1578651708] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 24.41, 21.32, 26.86
[1578651732] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 20.0% (3245224 kB) free!
[1578651741] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1578651772] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1578651852] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.9% (3239828 kB) free!
[1578651861] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1578651917] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14306228 kB) free!
[1578651981] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1578652197] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175632 kB) free!
[1578652308] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 602.62, 229.33, 99.01
[1578652317] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190972 kB) free!
[1578652428] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 122.14, 180.51, 97.92
[1578652437] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191324 kB) free!
[1578652452] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.3% (4112380 kB) free.
[1578652517] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394624 kB) free!
[1578652548] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1578652548] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 30.59, 126.14, 88.01
[1578653117] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18386652 kB) free!
[1578653572] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1578653652] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.8% (2076324 kB) free!
[1578653717] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (38008872 kB) free.
[1578653773] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.6% (2206968 kB) free!
[1578653782] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1578653893] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.1% (2299164 kB) free!
[1578654172] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1578654198] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1578654198] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.5% (21947776 kB) free.
[1578654234] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.2% (22011364 kB) free.
[1578654292] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1578654307] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.4% (40121796 kB) free.
[1578654317] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11042704 kB) free!
[1578654348] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 17.70, 22.31, 32.48
[1578654355] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 94.2% (23089908 kB) free.
[1578654380] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (8933260 kB) free!
[1578654383] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (8825768 kB) free!
[1578654457] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 40.5% (53397868 kB) free.
[1578654543] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1578654543] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1605752 kB) free!
[1578654555] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.7% (4178508 kB) free.
[1578654843] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170224 kB) free!
[1578654948] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 18.67, 20.30, 26.47
[1578654963] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184496 kB) free!
[1578655083] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174808 kB) free!
[1578655155] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.9% (2252272 kB) free!
[1578655178] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.9% (22190764 kB) free.
[1578655275] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.0% (2442724 kB) free!
[1578655290] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 47.2% (7669000 kB) free.
[1578655509] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.2% (1983468 kB) free!
[1578655516] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4266456 kB) free!
[1578655629] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.2% (1988800 kB) free!
[1578655636] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (2843948 kB) free!
[1578655663] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26234820 kB) free!
[1578655749] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 12.2% (1983752 kB) free!
[1578655756] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (1577992 kB) free!
[1578655783] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.1% (22547008 kB) free!
[1578655903] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.7% (44415372 kB) free.
[1578656195] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1578656315] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1578656337] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.3% (3742752 kB) free!
[1578656387] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (2270772 kB) free!
[1578656457] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.6% (3576384 kB) free!
[1578656500] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (3090780 kB) free!
[1578656503] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (17090504 kB) free!
[1578656507] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192388 kB) free!
[1578656577] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 13.9% (3403068 kB) free!
[1578656620] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.0% (3419436 kB) free.
[1578656623] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (13050572 kB) free!
[1578656627] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180672 kB) free!
[1578656743] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (9171908 kB) free!
[1578656986] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.8% (3100420 kB) free!
[1578657000] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/10 17:12:52 [error] 21677#0: *24081766 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:14716/104, bytes from/to upstream:104/14716
[1578657106] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.4% (2779464 kB) free!
[1578657120] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/01/10 17:20:30 [error] 21680#0: *24083081 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:5629/87, bytes from/to upstream:87/5629
[1578657226] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.9% (2450932 kB) free!
[1578657240] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578657343] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.2% (25250524 kB) free!
[1578657348] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 22.61, 48.91, 39.78
[1578657468] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 17.18, 38.36, 37.03
[1578657588] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 20.64, 32.72, 35.13
[1578657777] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (1509068 kB) free!
[1578657820] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.2% (2630380 kB) free!
[1578657826] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1578657826] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (968240 kB) free!
[1578657892] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1578657940] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (3032972 kB) free!
[1578657943] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (3718148 kB) free!
[1578658012] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1578658060] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3474828 kB) free.
[1578658132] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1578658188] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 18.18, 21.43, 27.77
[1578658543] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16584812 kB) free!
[1578658555] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166240 kB) free!
[1578658660] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (1107900 kB) free!
[1578658675] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195028 kB) free!
[1578658780] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (1520800 kB) free!
[1578658795] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182688 kB) free!
[1578658900] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.9% (1940584 kB) free!
[1578659096] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 14.8% (3643952 kB) free!
[1578659143] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (374588 kB) free!
[1578659217] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 13.3% (3287148 kB) free!
[1578659337] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 11.9% (2934288 kB) free!
[1578659434] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.4% (22064992 kB) free.
[1578659472] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 45.1% (59412552 kB) free.
[1578659500] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 32.2% (5237232 kB) free.
[1578659503] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 75.3% (16915260 kB) free.
[1578659523] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 83.0% (40892168 kB) free.
[1578659556] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.7% (22242200 kB) free.
[1578659937] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (1156472 kB) free!
[1578660537] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.1% (22991620 kB) free.
[1578660638] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166300 kB) free!
[1578660700] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.6% (1891076 kB) free!
[1578660758] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175844 kB) free!
[1578660820] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.3% (2321204 kB) free!
[1578660878] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173228 kB) free!
[1578660940] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (2766632 kB) free!
[1578661149] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (185580 kB) free!
[1578661277] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9413752 kB) free!
[1578661327] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.5% (7641824 kB) free!
[1578661397] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (5778420 kB) free!
[1578661447] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 13.9% (6863936 kB) free!
[1578661517] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38892480 kB) free.
[1578661540] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.1% (4236476 kB) free.
[1578661567] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.8% (6791220 kB) free!
[1578661697] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8751456 kB) free!
[1578661749] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 26.7% (4342848 kB) free.
[1578661817] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.3% (8530344 kB) free!
[1578661937] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.8% (8283704 kB) free!
[1578662117] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12215516 kB) free!
[1578662237] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7910640 kB) free!
[1578662357] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (3967968 kB) free!
[1578662607] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1578662741] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (1441180 kB) free!
[1578662861] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.0% (1955380 kB) free!
[1578662957] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.9% (22314340 kB) free!
[1578662981] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.3% (2494304 kB) free!
[1578663558] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389652 kB) free!
[1578663581] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3298080 kB) free.
[1578663967] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3241072 kB) free!
[1578664158] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17944096 kB) free!
[1578664337] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4241680 kB) free!
[1578664758] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405152 kB) free!
[1578665358] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.7% (15464008 kB) free!
[1578665381] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.6% (1878948 kB) free!
[1578665501] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.3% (2331436 kB) free!
[1578665621] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (2905844 kB) free!
[1578665932] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1578665958] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39393256 kB) free.
[1578666221] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.3% (4270884 kB) free.
[1578666558] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13153172 kB) free!
[1578666596] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1578666678] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8300004 kB) free!
[1578666798] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (3883024 kB) free!
[1578666821] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.4% (2825316 kB) free!
[1578666941] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (3197876 kB) free!
[1578667061] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.5% (3978104 kB) free.
[1578667398] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26134980 kB) free!
[1578667661] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (1082672 kB) free!
[1578667781] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (1382780 kB) free!
[1578667901] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.1% (1796344 kB) free!
[1578667967] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.3% (4236908 kB) free!
[1578667998] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (2039816 kB) free!
[1578668087] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.8% (4116516 kB) free!
[1578668207] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.4% (4018388 kB) free!
[1578668502] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 33.9% (5511984 kB) free.
[1578668598] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.1% (22521272 kB) free!
[1578669102] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (2961732 kB) free!
[1578669105] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1578669198] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403468 kB) free!
[1578669222] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (3215468 kB) free!
[1578669342] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.0% (3407864 kB) free.
[1578669588] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1578669798] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20215188 kB) free!
[1578670007] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1995748 kB) free!
[1578670398] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393628 kB) free!
[1578670733] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1578670999] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20185208 kB) free!
[1578671007] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.3% (4729844 kB) free!
[1578671078] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1578671127] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.0% (4650280 kB) free!
[1578671142] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.1% (2130576 kB) free!
[1578671247] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.8% (4602920 kB) free!
[1578671262] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.4% (2336512 kB) free!
[1578671333] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1578671382] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.0% (2438324 kB) free!
[1578671453] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1578671573] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1578671599] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408496 kB) free!
[1578671703] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.9499.4401> (10.147.212.147:60460 -> 10.147.212.203:5672, vhost: '/', user: 'guest'), channel 1:
[1578671825] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1578672199] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19942940 kB) free!
[1578672773] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1578672799] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397928 kB) free!
[1578672893] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1578673013] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1578673182] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 31.7% (5160712 kB) free.
[1578673399] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17406796 kB) free!
[1578673567] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1578673999] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41132760 kB) free.
[1578674213] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1578674599] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14406732 kB) free!
[1578674719] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (9904240 kB) free!
[1578674814] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1578674839] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (5703452 kB) free!
[1578674934] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1578675054] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1578675054] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1578675439] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.7% (25995704 kB) free!
[1578675654] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1578676039] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (2559108 kB) free!
[1578676254] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1578676374] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1578676494] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1578676639] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24708660 kB) free!
[1578676647] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2363512 kB) free!
[1578676962] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1578677094] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1578677214] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1578677239] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (4780516 kB) free!
[1578677335] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1578677839] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21245684 kB) free!
[1578678439] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (40564648 kB) free.
[1578679022] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(35) < 2020/01/10 23:26:52 [error] 18370#0: *36871623 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:2088/17, bytes from/to upstream:17/2088
[1578679040] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16507160 kB) free!
[1578679135] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 333 processes
[1578679142] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2020/01/10 23:28:09 [error] 18367#0: *36872229 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:252/17, bytes from/to upstream:17/252
[1578679160] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (12369540 kB) free!
[1578679262] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578679280] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (8429728 kB) free!
[1578679384] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(871) < 103.255.146.164 [10/Jan/2020:23:28:44 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1578679440] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(908) < 2020/01/10 23:28:44 [error] 21679#0: *24147028 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578679505] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578679560] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578679735] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1578679862] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/01/10 23:40:41 [error] 18361#0: *36877256 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:252/17, bytes from/to upstream:17/252
[1578679880] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (25813800 kB) free!
[1578679983] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578680207] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 22.6% (5548220 kB) free.
[1578680335] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1578680455] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1578680480] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5403616 kB) free!
[1578680575] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1578680582] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/10 23:45:41 [error] 18366#0: *36879359 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:252/17, bytes from/to upstream:17/252
[1578680702] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578680792] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6882612 kB) free!
[1578680912] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6870724 kB) free!