[1552847621] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552847741] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552847797] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3092496 kB) free!
[1552847861] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1552847896] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.47, 4.66, 4.95
[1552847917] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11623492 kB) free.
[1552848326] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26734 MB (12% inode=99%):
[1552848462] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552848517] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.9% (5867164 kB) free!
[1552848582] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552848637] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (4566124 kB) free!
[1552848702] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552848757] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11610104 kB) free.
[1552848926] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 14624 MB (7% inode=99%):
[1552849357] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6583884 kB) free!
[1552849477] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.8% (5299936 kB) free!
[1552849597] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (3503964 kB) free!
[1552849902] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552850126] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 25902 MB (12% inode=99%):
[1552850197] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.7% (7236980 kB) free!
[1552850296] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.21, 5.97, 5.41
[1552850416] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.63, 5.82, 5.42
[1552850502] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552850536] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.11, 5.99, 5.54
[1552850622] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1552850727] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 14638 MB (7% inode=99%):
[1552850798] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11605044 kB) free.
[1552851222] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552851343] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552851398] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (3412644 kB) free!
[1552851463] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1552851518] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11621900 kB) free.
[1552851927] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26922 MB (12% inode=99%):
[1552852063] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552852118] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.9% (5864148 kB) free!
[1552852183] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552852238] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (4478032 kB) free!
[1552852303] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552852358] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11608172 kB) free.
[1552852527] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 14756 MB (7% inode=99%):
[1552852936] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 5.20, 4.88, 4.93
[1552852958] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6520516 kB) free!
[1552853078] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.6% (5228368 kB) free!
[1552853127] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 22955 MB (11% inode=99%):
[1552853198] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (3446564 kB) free!
[1552853503] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552853798] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.8% (7272928 kB) free!
[1552854103] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552854223] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1552854327] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 18855 MB (9% inode=99%):
[1552854398] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11602352 kB) free.
[1552854824] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552854927] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 22954 MB (11% inode=99%):
[1552854944] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552854999] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3766596 kB) free!
[1552855064] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552855119] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11621084 kB) free.
[1552855336] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.54, 5.96, 5.44
[1552855456] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.98, 6.01, 5.52
[1552855576] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.36, 6.13, 5.62
[1552855664] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552855719] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6116920 kB) free!
[1552855784] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552855839] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (4808468 kB) free!
[1552855904] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552855959] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11606248 kB) free.
[1552856128] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 14915 MB (7% inode=99%):
[1552856559] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6789568 kB) free!
[1552856679] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.1% (5489312 kB) free!
[1552856799] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3862084 kB) free!
[1552857104] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552857376] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 3.95, 4.54, 4.93
[1552857399] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.2% (7470440 kB) free!
[1552857705] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552857825] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552857928] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24990 MB (11% inode=99%):
[1552857999] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11599192 kB) free.
[1552858425] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552858545] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552858599] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (4055608 kB) free!
[1552858665] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552858719] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11618392 kB) free.
[1552859128] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 18830 MB (9% inode=99%):
[1552859265] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552859320] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6336324 kB) free!
[1552859385] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552859440] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.3% (5067916 kB) free!
[1552859505] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552859560] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3268964 kB) free!
[1552859728] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24977 MB (11% inode=99%):
[1552860160] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.6% (7186872 kB) free!
[1552860329] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 14736 MB (7% inode=99%):
[1552860376] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.02, 6.00, 5.52
[1552860496] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.52, 6.21, 5.65
[1552860616] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.34, 6.29, 5.76
[1552860706] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1552860760] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11615284 kB) free.
[1552861306] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552861360] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4081212 kB) free!
[1552861426] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552861480] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11600916 kB) free.
[1552861529] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26774 MB (12% inode=99%):
[1552862026] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552862080] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6376352 kB) free!
[1552862129] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 14735 MB (7% inode=99%):
[1552862146] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1552862200] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.3% (5095660 kB) free!
[1552862266] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552862320] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3231156 kB) free!
[1552862416] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.49, 4.53, 4.89
[1552862920] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.5% (7161948 kB) free!
[1552863329] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 27040 MB (12% inode=99%):
[1552863466] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552863521] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11600176 kB) free.
[1552863903] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 511.21, 298.05, 125.12
[1552863929] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 18848 MB (9% inode=99%):
[1552864023] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 517.30, 370.91, 172.79
[1552864067] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552864121] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3456000 kB) free!
[1552864143] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 190.73, 325.50, 182.30
[1552864187] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1552864231] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 12.3% (3009372 kB) free!
[1552864241] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11609368 kB) free.
[1552864529] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26021 MB (12% inode=99%):
[1552864787] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1552864816] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.35, 5.89, 5.23
[1552864841] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5971200 kB) free!
[1552864907] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552864936] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.84, 6.23, 5.44
[1552864961] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (4697980 kB) free!
[1552865027] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552865056] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.23, 6.21, 5.53
[1552865081] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11604756 kB) free.
[1552865681] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6646048 kB) free!
[1552865730] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 18860 MB (9% inode=99%):
[1552865801] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.9% (5357368 kB) free!
[1552865921] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3642828 kB) free!
[1552866227] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552866330] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 27058 MB (12% inode=99%):
[1552866521] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.1% (7464408 kB) free!
[1552866543] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.37, 0.63, 14.28
[1552866827] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552866930] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 14766 MB (7% inode=99%):
[1552866947] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1552867121] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11605240 kB) free.
[1552867456] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.91, 4.61, 4.86
[1552867530] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 22967 MB (11% inode=99%):
[1552867548] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552867668] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552867721] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (3985172 kB) free!
[1552867743] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.66, 0.58, 4.37
[1552867788] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1552867842] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11612340 kB) free.
[1552868388] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552868441] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6249644 kB) free!
[1552868508] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1552868561] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.1% (4970192 kB) free!
[1552868628] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552868681] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.1% (7421812 kB) free!
[1552868730] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 18848 MB (9% inode=99%):
[1552869330] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26789 MB (12% inode=99%):
[1552869828] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1552869856] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.81, 5.36, 5.01
[1552869881] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11613876 kB) free.
[1552869930] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 14756 MB (7% inode=99%):
[1552869976] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.71, 5.48, 5.10
[1552870096] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.88, 5.64, 5.20
[1552870429] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1552870481] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (3226472 kB) free!
[1552870549] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552870601] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11597820 kB) free.
[1552871149] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552871201] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.9% (5851572 kB) free!
[1552871269] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552871321] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (4479076 kB) free!
[1552871389] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552871441] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11613008 kB) free.
[1552871731] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26543 MB (12% inode=99%):
[1552872041] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6601948 kB) free!
[1552872162] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.8% (5323672 kB) free!
[1552872282] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3599244 kB) free!
[1552872331] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 18863 MB (9% inode=99%):
[1552872496] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.94, 4.73, 4.91
[1552872589] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552872882] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.9% (7350772 kB) free!
[1552873189] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552873309] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1552873482] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11612780 kB) free.
[1552873531] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 22964 MB (11% inode=99%):
[1552873910] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1552874030] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552874082] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (3668192 kB) free!
[1552874150] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552874202] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11598228 kB) free.
[1552874296] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.56, 5.12, 5.03
[1552874416] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.22, 5.15, 5.05
[1552874536] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.81, 4.99, 5.01
[1552874731] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 9324 MB (4% inode=99%):
[1552874750] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552874802] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6094456 kB) free!
[1552874870] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1552874922] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (4801968 kB) free!
[1552874990] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1552875042] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11611572 kB) free.
[1552875642] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6863424 kB) free!
[1552875762] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.3% (5582356 kB) free!
[1552875883] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4042516 kB) free!
[1552876190] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1552876483] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.5% (7632256 kB) free!
[1552876790] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1552876911] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1552877083] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11607512 kB) free.
[1552877132] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24795 MB (11% inode=99%):
[1552877510] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552877630] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552877683] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4322212 kB) free!
[1552877750] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552877803] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11582756 kB) free.
[1552878332] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 15275 MB (7% inode=99%):
[1552878403] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6475808 kB) free!
[1552878523] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.5% (5191084 kB) free!
[1552878643] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (3364412 kB) free!
[1552878950] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1552879243] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.6% (7184308 kB) free!
[1552879550] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552879670] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552879843] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11576344 kB) free.
[1552880271] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552880391] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552880443] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (3366492 kB) free!
[1552880511] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1552880564] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11095440 kB) free.
[1552880732] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26483 MB (12% inode=99%):
[1552880913] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 572.77, 339.73, 145.64
[1552881033] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 525.64, 398.80, 190.82
[1552881111] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552881153] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 137.53, 316.64, 187.77
[1552881211] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.9% (3431572 kB) free!
[1552881231] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552881332] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 14631 MB (7% inode=99%):
[1552881351] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552881763] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3732016 kB) free!
[1552881883] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11363132 kB) free.
[1552883083] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3952832 kB) free!
[1552883203] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11570460 kB) free.
[1552883498] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 8 users currently logged in
[1552883553] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.69, 0.57, 14.50
[1552883618] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1552883738] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1552883774] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.0% (2705340 kB) free!
[1552883803] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3047688 kB) free!
[1552883923] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10989312 kB) free.
[1552884632] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2354128 kB) free!
[1552884736] SERVICE FLAPPING ALERT: sigqscluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1552884753] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.45, 0.41, 4.28
[1552884974] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2299572 kB) free!
[1552885123] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3591152 kB) free!
[1552885143] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.56, 6.07, 5.22
[1552885232] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (2539608 kB) free!
[1552885232] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 7016 MB (13% inode=98%):
[1552885236] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 337 processes
[1552885243] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11414668 kB) free.
[1552885263] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.23, 4.13, 4.61
[1552885532] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23900 MB (11% inode=99%):
[1552885546] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.86, 7.52, 6.25
[1552885666] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.46, 5.49, 5.66
[1552885786] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.11, 4.34, 5.21
[1552885832] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2431152 kB) free!
[1552885844] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3604432 kB) free!
[1552885964] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10621160 kB) free.
[1552886132] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 13660 MB (6% inode=99%):
[1552886386] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.73, 2.00, 3.51
[1552886726] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 25.86, 207.61, 147.78
[1552886846] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.66, 139.02, 129.90
[1552886875] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 13.1% (3200156 kB) free!
[1552886966] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.99, 93.17, 114.20
[1552887164] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3931008 kB) free!
[1552887284] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11510116 kB) free.
[1552887884] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (3424896 kB) free!
[1552887951] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1552888004] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11212016 kB) free.
[1552888237] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 334 processes
[1552888552] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552888604] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (2989400 kB) free!
[1552888672] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1552888724] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10753836 kB) free.
[1552888766] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.40, 0.57, 16.73
[1552889272] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552889392] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1552889924] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4488244 kB) free!
[1552889966] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.38, 0.32, 4.79
[1552889992] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552889992] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1552890044] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11492132 kB) free.
[1552890112] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1552890112] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1552890232] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1552890342] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 430 processes
[1552890566] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.91, 37.98, 31.61
[1552890645] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4157132 kB) free!
[1552890686] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.74, 25.57, 27.83
[1552890712] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552890765] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11483512 kB) free.
[1552890806] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1552890806] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.47, 17.22, 24.50
[1552890832] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1552890832] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552890932] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 25351 MB (12% inode=99%):
[1552890953] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552891365] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3426000 kB) free!
[1552891406] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.13, 2.50, 12.95
[1552891485] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11206168 kB) free.
[1552891533] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 17159 MB (8% inode=99%):
[1552892085] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (2914404 kB) free!
[1552892153] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552892205] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10879880 kB) free.
[1552892606] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.37, 0.31, 3.72
[1552892733] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24309 MB (11% inode=99%):
[1552892753] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552892805] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1552892873] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552892993] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552893333] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 12016 MB (5% inode=99%):
[1552893405] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (3291124 kB) free!
[1552893525] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11442404 kB) free.
[1552894125] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (3813264 kB) free!
[1552894193] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1552894245] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11451844 kB) free.
[1552894793] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552894846] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4083436 kB) free!
[1552894914] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552894966] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11446000 kB) free.
[1552895034] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552895133] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23729 MB (11% inode=99%):
[1552895566] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3488808 kB) free!
[1552895634] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1552895686] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11109056 kB) free.
[1552895733] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 9967 MB (4% inode=99%):
[1552896234] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552896286] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3128468 kB) free!
[1552896354] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1552896406] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10832364 kB) free.
[1552896942] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 322 processes
[1552896954] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552897074] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1552897606] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4912652 kB) free!
[1552897674] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552897726] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (3121180 kB) free!
[1552897794] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1552897846] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11447836 kB) free.
[1552898133] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24401 MB (11% inode=99%):
[1552898387] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 203.91, 50.68, 17.98
[1552898395] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552898446] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4595032 kB) free!
[1552898507] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 472.36, 204.10, 78.16
[1552898515] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552898566] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11466608 kB) free.
[1552898627] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 508.98, 306.85, 131.08
[1552898635] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552898733] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 12106 MB (5% inode=99%):
[1552899032] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5392 MB (10% inode=97%):
[1552899167] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4116064 kB) free!
[1552899287] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11456764 kB) free.
[1552899887] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3609044 kB) free!
[1552900007] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11206744 kB) free.
[1552900539] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1552900607] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (3258516 kB) free!
[1552900727] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10920108 kB) free.
[1552901133] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24664 MB (11% inode=99%):
[1552901327] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3122880 kB) free!
[1552901447] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10877372 kB) free.
[1552901627] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.44, 1.45, 12.88
[1552901733] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 15448 MB (7% inode=99%):
[1552901739] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1552901859] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1552901979] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1552902647] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (4012664 kB) free!
[1552902767] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11466472 kB) free.
[1552902827] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.38, 1.58, 4.71
[1552903368] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (3384448 kB) free!
[1552903406] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1552903435] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552903488] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11452368 kB) free.
[1552903534] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23632 MB (11% inode=99%):
[1552903966] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 36.6% (5948168 kB) free.
[1552904035] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552904088] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.3% (5571464 kB) free!
[1552904134] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 12363 MB (5% inode=99%):
[1552904155] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552904208] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (4252968 kB) free!
[1552904275] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552904328] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11466028 kB) free.
[1552904928] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3479548 kB) free!
[1552905048] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11146324 kB) free.
[1552905166] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (170632 kB) free!
[1552905286] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (173260 kB) free!
[1552905406] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (179536 kB) free!
[1552905648] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3057164 kB) free!
[1552905768] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11458752 kB) free.
[1552905934] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24154 MB (11% inode=99%):
[1552906368] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4725996 kB) free!
[1552906432] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1552906488] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10871384 kB) free.
[1552906534] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 12628 MB (6% inode=99%):
[1552906552] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1552906672] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1552907088] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.3% (5589028 kB) free!
[1552907209] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (4208060 kB) free!
[1552907329] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11446728 kB) free.
[1552907379] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1552907929] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3919632 kB) free!
[1552908049] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11431956 kB) free.
[1552908649] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (3154000 kB) free!
[1552908769] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10992788 kB) free.
[1552909139] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 541 processes
[1552909179] SERVICE FLAPPING ALERT: secondary-server;Current Users;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1552909179] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1552909199] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 528 processes
[1552909893] SERVICE ALERT: sig2-vm1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 285 processes
[1552909969] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (3399292 kB) free!
[1552910013] SERVICE ALERT: sig2-vm1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 285 processes
[1552910089] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11198996 kB) free.
[1552910133] SERVICE ALERT: sig2-vm1;CPU_Procs;OK;HARD;1;CPU OK: 282 processes
[1552910575] SERVICE ALERT: sig2-vm1;Memory;OK;HARD;1;OK - 29.3% (2153176 kB) free.
[1552910689] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (3329052 kB) free!
[1552910809] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (11003076 kB) free.
[1552911409] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4075956 kB) free!
[1552911530] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11439592 kB) free.
[1552912129] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (3224756 kB) free!
[1552912249] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10855052 kB) free.
[1552912534] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24167 MB (11% inode=99%):
[1552913134] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 15975 MB (7% inode=99%):
[1552913449] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3847972 kB) free!
[1552913569] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11409284 kB) free.
[1552914169] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.5% (5175144 kB) free!
[1552914289] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10449160 kB) free.
[1552915489] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (2789916 kB) free!
[1552915609] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10430876 kB) free.
[1552915675] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1552916810] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3040916 kB) free!
[1552916875] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1552916930] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11077792 kB) free.
[1552917473] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1552917476] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552917530] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (2820100 kB) free!
[1552917579] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1552917596] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1552917650] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10634920 kB) free.
[1552918196] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552918316] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552918436] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552918850] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (3276908 kB) free!
[1552918970] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10969696 kB) free.
[1552920170] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (3649532 kB) free!
[1552920236] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1552920290] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11253172 kB) free.
[1552920837] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552920934] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23783 MB (11% inode=99%):
[1552920957] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552921077] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1552921077] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552921411] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2479976 kB) free!
[1552921491] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3847352 kB) free!
[1552921534] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 22887 MB (10% inode=99%):
[1552921611] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11379388 kB) free.
[1552922211] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (2862584 kB) free!
[1552922277] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1552922331] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10594984 kB) free.
[1552922877] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552922997] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552923117] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552923531] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (3184972 kB) free!
[1552923579] SERVICE FLAPPING ALERT: secondary-server;Current Users;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1552923651] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10870808 kB) free.
[1552924851] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3451980 kB) free!
[1552924917] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1552924971] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11378956 kB) free.
[1552925517] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552925571] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.9% (5853252 kB) free!
[1552925637] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552925692] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (4555636 kB) free!
[1552925757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552925812] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11364180 kB) free.
[1552926412] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6525236 kB) free!
[1552926532] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.6% (5236484 kB) free!
[1552926652] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3645376 kB) free!
[1552926957] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1552927252] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.7% (7262928 kB) free!
[1552927558] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552927678] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552927852] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11352272 kB) free.
[1552928278] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552928398] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552928452] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3763384 kB) free!
[1552928518] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1552928572] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11369444 kB) free.
[1552929118] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552929172] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.2% (5991008 kB) free!
[1552929238] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552929292] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (4700364 kB) free!
[1552929358] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552929413] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11356944 kB) free.
[1552930012] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6673812 kB) free!
[1552930132] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.9% (5361412 kB) free!
[1552930252] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1552930252] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3825404 kB) free!
[1552930558] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1552930852] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.9% (7359920 kB) free!
[1552931134] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24054 MB (11% inode=99%):
[1552931158] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552931278] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1552931452] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11348100 kB) free.
[1552931734] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 15860 MB (7% inode=99%):
[1552931878] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552931998] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552932052] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (4026340 kB) free!
[1552932118] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552932172] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11366712 kB) free.
[1552932772] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.6% (6188544 kB) free!
[1552932892] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (4917724 kB) free!
[1552932934] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24027 MB (11% inode=99%):
[1552933012] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (3053784 kB) free!
[1552933319] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552933534] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 14299 MB (6% inode=99%):
[1552933612] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.0% (6904156 kB) free!