[1566758135] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1566758255] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1566759575] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164600 kB) free!
[1566759695] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189656 kB) free!
[1566759815] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190952 kB) free!
[1566760292] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1566760412] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566760415] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 65.2% (16097948 kB) free.
[1566762214] SERVICE ALERT: secondary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 358 processes
[1566762334] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1566766412] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1566766415] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196096 kB) free!
[1566766532] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566766535] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198872 kB) free!
[1566766655] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198072 kB) free!
[1566768213] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 14.53, 45.90, 28.23
[1566768333] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.04, 31.05, 24.93
[1566768453] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.84, 20.97, 21.98
[1566768455] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1566768455] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.8% (2653644 kB) free!
[1566768697] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187548 kB) free!
[1566769053] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.23, 6.78, 13.50
[1566769056] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172964 kB) free!
[1566769297] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 48.4% (11856560 kB) free.
[1566770253] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.62, 0.80, 4.26
[1566779733] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 407 processes
[1566779853] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 408 processes
[1566780456] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1566781656] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 73.0% (18006552 kB) free.
[1566791013] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1566796653] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1566796773] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1566796953] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 442 processes
[1566797013] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 472 processes
[1566799413] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1566799533] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1566799653] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1566801031] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (4546920 kB) free!
[1566801098] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (4697716 kB) free!
[1566801151] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (4308988 kB) free!
[1566801218] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4634240 kB) free!
[1566801271] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (4192208 kB) free!
[1566801338] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.6% (4557072 kB) free!
[1566801453] SERVICE ALERT: sigqscluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1566802414] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 279 processes
[1566802534] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 277 processes
[1566802654] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 278 processes
[1566802655] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1566802775] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1566804271] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2233004 kB) free!
[1566807338] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2353224 kB) free!
[1566808054] SERVICE ALERT: sig5-vm2;CPU_Procs;OK;HARD;1;CPU OK: 219 processes
[1566809014] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 321 processes
[1566809134] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 321 processes
[1566809254] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 321 processes
[1566810806] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.0% (21803916 kB) free.
[1566810908] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.3% (22121532 kB) free.
[1566811055] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 321 processes
[1566811175] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 319 processes
[1566815135] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 158.93, 124.54, 62.43
[1566815255] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 125.11, 128.47, 71.70
[1566815375] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 18.47, 86.59, 63.24
[1566816453] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 57.56, 15.61, 5.61
[1566816573] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 101.22, 46.20, 18.04
[1566816575] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.43, 3.35, 18.54
[1566816693] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 107.40, 66.81, 29.00
[1566817413] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1566817533] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1566817653] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1566817893] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.31, 2.71, 13.19
[1566818375] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.41, 1.88, 4.00
[1566818853] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9801700 kB) free!
[1566818973] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.6% (9680392 kB) free!
[1566819093] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.4% (9556180 kB) free!
[1566819093] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.30, 0.42, 3.91
[1566819453] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1566820055] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 259 processes
[1566820175] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 259 processes
[1566820295] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 258 processes
[1566820454] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 52.9% (26062576 kB) free.
[1566829173] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1566829293] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1566836014] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4821632 kB) free!
[1566836134] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (4791208 kB) free!
[1566836254] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (4779088 kB) free!
[1566838293] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1566838413] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566840914] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4893400 kB) free!
[1566841035] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (4868396 kB) free!
[1566841155] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.8% (4853296 kB) free!