[1589222010] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11521564 kB) free!
[1589222130] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7677440 kB) free!
[1589222250] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (3934188 kB) free!
[1589222850] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15825696 kB) free!
[1589223450] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28925352 kB) free.
[1589223726] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 15.9% (3905472 kB) free!
[1589223786] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (896280 kB) free!
[1589223854] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (190976 kB) free!
[1589223902] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 347 processes
[1589223909] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1589223909] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.3% (3790432 kB) free.
[1589223914] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (192972 kB) free!
[1589224050] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7642600 kB) free!
[1589224170] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37591700 kB) free.
[1589224505] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1589224509] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.1% (1805248 kB) free!
[1589224629] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (1521948 kB) free!
[1589224749] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1185424 kB) free!
[1589224770] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14194684 kB) free!
[1589224814] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 80.8% (19792560 kB) free.
[1589224890] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (10305680 kB) free!
[1589225010] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (6703832 kB) free!
[1589225349] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.5% (1871104 kB) free!
[1589225611] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16640460 kB) free!
[1589225949] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1077552 kB) free!
[1589226211] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28608432 kB) free.
[1589226550] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.7% (1746376 kB) free!
[1589226811] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7704956 kB) free!
[1589226931] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37172560 kB) free.
[1589227104] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (219520 kB) free!
[1589227150] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (238068 kB) free!
[1589227224] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (227500 kB) free!
[1589227294] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (1776428 kB) free!
[1589227345] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235756 kB) free!
[1589227414] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1773640 kB) free!
[1589227531] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12828880 kB) free!
[1589227534] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (1351028 kB) free!
[1589227651] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9119976 kB) free!
[1589227750] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.3% (3629524 kB) free.
[1589227771] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29513972 kB) free.
[1589228350] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.4% (1855080 kB) free!
[1589228371] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14350348 kB) free!
[1589228470] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 31.1% (5057228 kB) free.
[1589228491] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (10504312 kB) free!
[1589228611] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (6986148 kB) free!
[1589229070] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3170584 kB) free!
[1589229145] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.9% (7834064 kB) free!
[1589229191] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 37.4% (6079836 kB) free.
[1589229211] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15121704 kB) free!
[1589229931] SERVICE ALERT: sig4-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 30268 MB (23% inode=99%):
[1589230411] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36803080 kB) free.
[1589231011] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14082884 kB) free!
[1589231131] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (10481372 kB) free!
[1589231251] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36753416 kB) free.
[1589231614] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.4% (44543640 kB) free.
[1589231777] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 89.7% (44181184 kB) free.
[1589231852] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14268064 kB) free!
[1589231972] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (10775092 kB) free!
[1589232092] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36701400 kB) free.
[1589232692] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14362936 kB) free!
[1589232746] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.1% (44427112 kB) free.
[1589232812] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (10839388 kB) free!
[1589232821] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(446) < 2020/05/12 03:01:31 [error] 22133#0: *87779709 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1589232877] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2707) < 203.212.70.170 [12/May/2020:03:02:03 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1589232905] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1589232932] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(350) < 182.18.144.246 [12/May/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1589232932] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36684792 kB) free.
[1589232941] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589232997] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589233052] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589233136] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 367 processes
[1589233256] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1589233532] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13355944 kB) free!
[1589233652] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (9929080 kB) free!
[1589233772] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34947916 kB) free.
[1589234372] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12756744 kB) free!
[1589234492] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28705712 kB) free.
[1589234657] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.0% (10061480 kB) free.
[1589235092] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15486512 kB) free!
[1589235212] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11792000 kB) free!
[1589235332] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36386976 kB) free.
[1589235932] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14788440 kB) free!
[1589236052] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11095376 kB) free!
[1589236173] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35263508 kB) free.
[1589236457] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (1575012 kB) free!
[1589236577] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (202300 kB) free!
[1589236697] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177984 kB) free!
[1589236772] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6178392 kB) free!
[1589236892] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (478556 kB) free!
[1589237012] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32672868 kB) free.
[1589237612] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5785948 kB) free!
[1589237732] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11904948 kB) free!
[1589237852] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31678088 kB) free.
[1589238452] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6697248 kB) free!
[1589238497] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 38.4% (9413748 kB) free.
[1589238572] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35688516 kB) free.
[1589239172] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11314768 kB) free!
[1589239292] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7386728 kB) free!
[1589239390] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589239412] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35682152 kB) free.
[1589240012] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10374436 kB) free!
[1589240132] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6206844 kB) free!
[1589240252] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32791432 kB) free.
[1589240297] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188056 kB) free!
[1589240417] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186400 kB) free!
[1589240537] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1589240537] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (220780 kB) free!
[1589240853] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12080680 kB) free!
[1589240973] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37661884 kB) free.
[1589241397] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 70.0% (17164524 kB) free.
[1589241573] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15822236 kB) free!
[1589241693] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37593780 kB) free.
[1589241716] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 54.6% (13470364 kB) free.
[1589241735] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 57.7% (14138668 kB) free.
[1589242293] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16023876 kB) free!
[1589242413] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37496284 kB) free.
[1589243013] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16325380 kB) free!
[1589243133] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30341692 kB) free.
[1589243733] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17700296 kB) free!
[1589243853] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.5% (13881448 kB) free!
[1589243973] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37529864 kB) free.
[1589244305] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 343 processes
[1589244425] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 342 processes
[1589244545] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1589244573] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16704716 kB) free!
[1589244693] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1589244693] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38067136 kB) free.
[1589245293] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17218540 kB) free!
[1589245413] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38537892 kB) free.
[1589245534] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 32.7% (43078136 kB) free.
[1589245814] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 14.4% (3517748 kB) free!
[1589245875] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 14.4% (3521672 kB) free!
[1589245935] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 14.4% (3522248 kB) free!
[1589245995] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 14.4% (3525876 kB) free!
[1589246013] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17781988 kB) free!
[1589246133] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38287960 kB) free.
[1589246517] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (209204 kB) free!
[1589246637] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185148 kB) free!
[1589246734] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18368084 kB) free!
[1589246757] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192320 kB) free!
[1589246854] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38117296 kB) free.
[1589247454] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18201496 kB) free!
[1589247574] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37884436 kB) free.
[1589248174] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18430732 kB) free!
[1589248294] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37642216 kB) free.
[1589248894] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17686000 kB) free!
[1589249014] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.9% (22223164 kB) free!
[1589249134] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34897328 kB) free.
[1589250334] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (17146716 kB) free!
[1589250454] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36560468 kB) free.
[1589250737] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1589251054] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16162960 kB) free!
[1589251174] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27440636 kB) free.
[1589251774] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16247200 kB) free!
[1589251895] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (12068344 kB) free!
[1589252015] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31498904 kB) free.
[1589252535] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (785892 kB) free!
[1589252615] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18623256 kB) free!
[1589252655] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186764 kB) free!
[1589252735] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29081784 kB) free.
[1589252767] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.98, 6.17, 5.61
[1589252775] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188352 kB) free!
[1589252859] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 327 processes
[1589252887] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.55, 5.30, 5.36
[1589252979] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1589253007] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.35, 4.21, 4.95
[1589253935] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10778888 kB) free!
[1589254055] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32836396 kB) free.
[1589254655] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9342292 kB) free!
[1589254745] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1589254775] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (5056036 kB) free!
[1589254807] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.46, 5.76, 5.18
[1589254895] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29378336 kB) free.
[1589254927] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.32, 5.90, 5.32
[1589255047] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.84, 6.01, 5.44
[1589255495] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5831576 kB) free!
[1589255615] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30225660 kB) free.
[1589256215] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6863256 kB) free!
[1589256335] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2284012 kB) free!
[1589256455] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28614128 kB) free.
[1589257056] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3334272 kB) free!
[1589257176] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (450676 kB) free!
[1589257296] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29934656 kB) free.
[1589257896] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6300640 kB) free!
[1589258016] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1752376 kB) free!
[1589258136] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27065524 kB) free.
[1589258197] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.9% (2662928 kB) free!
[1589258198] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(334) < 103.255.146.164 [12/May/2020:10:02:10 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589258317] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188040 kB) free!
[1589258318] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589258437] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186208 kB) free!
[1589258573] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1589258693] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1589258736] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (3446960 kB) free!
[1589258813] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1589258856] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (479132 kB) free!
[1589258976] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28613244 kB) free.
[1589259413] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1589259518] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(635) < 103.255.146.164 [12/May/2020:10:22:47 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589259576] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4813044 kB) free!
[1589259579] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1589259638] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589259696] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (685252 kB) free!
[1589259699] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1589259816] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28551812 kB) free.
[1589260416] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (415644 kB) free!
[1589260536] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (388540 kB) free!
[1589260657] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (2963176 kB) free!
[1589261015] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.7% (3522988 kB) free!
[1589261131] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 26731 MB (20% inode=99%):
[1589261135] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.2% (2954848 kB) free!
[1589261251] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 26667 MB (20% inode=99%):
[1589261255] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.6% (2374288 kB) free!
[1589261371] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 26667 MB (20% inode=99%):
[1589261804] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.12352.300> (10.147.212.158:33882 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1589261855] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (183480 kB) free!
[1589261857] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.4% (24248944 kB) free!
[1589261924] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589262457] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (1086992 kB) free!
[1589262468] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 40.9% (10036228 kB) free.
[1589262572] SERVICE ALERT: sig4-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 110811 MB (86% inode=99%):
[1589262847] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 36.26, 29.29, 21.65
[1589262934] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.3% (17560808 kB) free!
[1589263054] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.4% (26830520 kB) free.
[1589263448] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 16.13, 18.55, 19.55
[1589263546] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.8% (4841196 kB) free!
[1589263657] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1589263657] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26879656 kB) free.
[1589263666] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.6% (4815532 kB) free!
[1589263731] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (604856 kB) free!
[1589263786] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.6% (4794316 kB) free!
[1589264019] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 15.5% (5433676 kB) free!
[1589264048] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1589264048] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 17.10, 19.09, 20.32
[1589264139] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 13.7% (4803360 kB) free!
[1589264216] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 16.3% (4030360 kB) free!
[1589264254] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.2% (22643176 kB) free!
[1589264257] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (1612096 kB) free!
[1589264259] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 12.1% (4240812 kB) free!
[1589264268] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.3% (4004656 kB) free!
[1589264336] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 14.5% (3580352 kB) free!
[1589264374] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 12.0% (15798428 kB) free!
[1589264377] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (408488 kB) free!
[1589264389] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.0% (3684500 kB) free!
[1589264456] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 12.7% (3124808 kB) free!
[1589264494] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.2% (18720460 kB) free!
[1589264497] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23223284 kB) free!
[1589264509] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 13.5% (3302568 kB) free!
[1589264859] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (1327388 kB) free!
[1589265056] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (787908 kB) free!
[1589265095] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.0% (29005940 kB) free.
[1589265097] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393008 kB) free!
[1589265109] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (1449932 kB) free!
[1589265456] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.8% (16115180 kB) free.
[1589265462] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 91.1% (31907036 kB) free.
[1589265695] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.7% (21963200 kB) free!
[1589265697] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (25663796 kB) free!
[1589265699] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1589265815] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 10.8% (14220204 kB) free!
[1589265819] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1589265935] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.6% (21931108 kB) free!
[1589265977] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9560616 kB) free!
[1589266098] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.1% (8897328 kB) free!
[1589266218] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.2% (8455760 kB) free!
[1589266238] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(12) < 182.18.144.246 [12/May/2020:12:18:52 +0530] TCP 502 0 0 91.084 "10.147.212.155:3001, 10.147.212.154:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 31.083, 60.000, 0.000"
[1589266256] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.6% (22859168 kB) free.
[1589266297] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (2476656 kB) free!
[1589266358] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589266570] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 760 processes
[1589266690] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 760 processes
[1589266811] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 762 processes
[1589266897] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17195304 kB) free!
[1589267014] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8827368 kB) free!
[1589267134] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.4% (8563960 kB) free!
[1589267254] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.8% (8265444 kB) free!
[1589267497] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390596 kB) free!
[1589267648] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 13.25, 12.25, 16.33
[1589268018] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4350156 kB) free!
[1589268190] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.6% (2529668 kB) free!
[1589268310] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.4% (2670464 kB) free!
[1589268430] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.4% (2832460 kB) free!
[1589268698] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26872756 kB) free.
[1589269030] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3322220 kB) free.
[1589269054] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3827492 kB) free!
[1589269298] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4224480 kB) free!
[1589269418] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (377588 kB) free!
[1589269538] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34153200 kB) free.
[1589269630] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (1136400 kB) free!
[1589269750] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (1412816 kB) free!
[1589269871] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.3% (1670256 kB) free!
[1589269958] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(767) < 103.255.146.164 [12/May/2020:13:16:07 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589270078] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589270135] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1589270135] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.4% (28184096 kB) free.
[1589270138] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6357252 kB) free!
[1589270258] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2066356 kB) free!
[1589270346] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1589270378] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391436 kB) free!
[1589270978] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (14970192 kB) free!
[1589271279] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(544) < 103.255.146.164 [12/May/2020:13:43:44 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589271335] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.1% (22560824 kB) free!
[1589271399] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589271455] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 12.8% (16824272 kB) free!
[1589271575] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.6% (24523624 kB) free!
[1589271578] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33672300 kB) free.
[1589271671] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1589271671] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3360964 kB) free.
[1589272178] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9659232 kB) free!
[1589272271] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.0% (2758624 kB) free!
[1589272298] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5353640 kB) free!
[1589272391] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (3032112 kB) free!
[1589272418] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1154332 kB) free!
[1589272511] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (2949700 kB) free!
[1589272656] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4291420 kB) free!
[1589272746] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1589272776] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.2% (3860648 kB) free!
[1589272866] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 349 processes
[1589272896] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.1% (3395248 kB) free!
[1589272986] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1589273018] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14705772 kB) free!
[1589273375] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (11858924 kB) free!
[1589273496] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1589273496] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (956776 kB) free!
[1589273586] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1589273618] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29922260 kB) free.
[1589273711] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1173436 kB) free!
[1589273975] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.4% (13773172 kB) free!
[1589274219] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9090232 kB) free!
[1589274339] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4591348 kB) free!
[1589274459] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (378908 kB) free!
[1589274999] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(165) < 103.255.146.164 [12/May/2020:14:43:00 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589275013] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1589275119] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589275133] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1589275659] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18895912 kB) free!
[1589275786] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 32.0% (7833988 kB) free.
[1589275897] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.3% (16227244 kB) free.
[1589275909] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 64.6% (15837024 kB) free.
[1589276048] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1589276111] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.6% (2052248 kB) free!
[1589276259] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (499768 kB) free!
[1589276319] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(144) < 185.255.8.58 [12/May/2020:15:00:05 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1589276376] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (6139084 kB) free!
[1589276439] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589276711] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161664 kB) free!
[1589276859] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19820200 kB) free!
[1589277459] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (416052 kB) free!
[1589277941] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/12 15:34:01 [error] 22135#0: *87904171 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1589278059] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21230104 kB) free!
[1589278061] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/05/12 15:37:38 [error] 22135#0: *87904923 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1589278095] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 9.7% (2383744 kB) free!
[1589278181] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589278239] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(186) < 103.255.146.164 [12/May/2020:15:31:27 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589278359] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589278659] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387360 kB) free!
[1589278695] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 10.1% (2470528 kB) free!
[1589279111] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.0% (3744964 kB) free.
[1589279259] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17357564 kB) free!
[1589279295] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 10.0% (2458852 kB) free!
[1589279376] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.5% (15151048 kB) free!
[1589279559] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(185) < 103.255.146.164 [12/May/2020:16:01:18 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589279595] SERVICE FLAPPING ALERT: webserver2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1589279595] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 10.1% (2463796 kB) free!
[1589279680] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589279859] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (379400 kB) free!
[1589280195] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 10.0% (2459632 kB) free!
[1589280311] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (966872 kB) free!
[1589280431] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.5% (2192020 kB) free!
[1589280459] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15114212 kB) free!
[1589280551] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.9% (3233204 kB) free!
[1589280576] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410596 kB) free!
[1589280795] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 10.1% (2467956 kB) free!
[1589281059] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (441952 kB) free!
[1589281151] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3342056 kB) free.
[1589281176] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.7% (28579248 kB) free.
[1589281309] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 14.4% (3525200 kB) free!
[1589281429] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 13.3% (3256836 kB) free!
[1589281549] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 12.6% (3093808 kB) free!
[1589281751] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.4% (1853804 kB) free!
[1589281776] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11754796 kB) free!
[1589281871] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (3097144 kB) free!
[1589281896] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.3% (25486636 kB) free!
[1589281986] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1589281991] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (3158572 kB) free!
[1589282016] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.4% (19010948 kB) free!
[1589282150] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2141832 kB) free!
[1589282616] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (3511932 kB) free!
[1589282981] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/12 16:55:56 [error] 22135#0: *87919708 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1589283097] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.8% (3994032 kB) free!
[1589283101] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589283191] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3486400 kB) free.
[1589283217] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.7% (3526428 kB) free!
[1589283337] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.7% (3081636 kB) free!
[1589283459] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31978936 kB) free.
[1589283495] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 10.0% (2442804 kB) free!
[1589283937] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (1184952 kB) free!
[1589284059] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7545536 kB) free!
[1589284179] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3368540 kB) free!
[1589284299] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (451604 kB) free!
[1589284391] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.2% (1989088 kB) free!
[1589284416] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.4% (25564336 kB) free!
[1589284479] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 203.212.70.170 [12/May/2020:17:20:52 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1589284512] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.5% (3499836 kB) free.
[1589284599] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589284899] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19495348 kB) free!
[1589285016] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.8% (30070272 kB) free.
[1589285111] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (1293304 kB) free!
[1589285231] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (2865980 kB) free!
[1589285351] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.2% (3604092 kB) free.
[1589285499] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383652 kB) free!
[1589285586] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1589285616] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 10.3% (13592644 kB) free!
[1589285706] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 344 processes
[1589285736] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.7% (28612356 kB) free.
[1589285826] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1589285951] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.2% (2300356 kB) free!
[1589286071] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.5% (3822196 kB) free.
[1589286099] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18003368 kB) free!
[1589286336] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10524324 kB) free!
[1589286456] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.3% (24124688 kB) free!
[1589286576] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.4% (15044832 kB) free!
[1589286671] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.4% (1860524 kB) free!
[1589286699] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397584 kB) free!
[1589286791] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.1% (3915668 kB) free.
[1589286795] SERVICE FLAPPING ALERT: webserver2;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1589286937] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.6% (16305084 kB) free.
[1589287299] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17084256 kB) free!
[1589287391] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.5% (1707136 kB) free!
[1589287511] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3755348 kB) free.
[1589287899] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (458884 kB) free!
[1589288111] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (1566648 kB) free!
[1589288226] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1589288231] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3579704 kB) free.
[1589288376] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.9% (27574036 kB) free.
[1589288499] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16421632 kB) free!
[1589288595] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 10.1% (2464360 kB) free!
[1589288831] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (1492216 kB) free!
[1589288895] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 10.0% (2459000 kB) free!
[1589288952] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.8% (3704316 kB) free.
[1589289099] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (431968 kB) free!
[1589289551] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (1462144 kB) free!
[1589289671] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3672060 kB) free.
[1589289699] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15854544 kB) free!
[1589290176] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12766412 kB) free!
[1589290271] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (173240 kB) free!
[1589290296] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.9% (24918240 kB) free!
[1589290299] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398568 kB) free!
[1589290391] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (163020 kB) free!
[1589290416] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.6% (16618076 kB) free!
[1589290511] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (1378428 kB) free!
[1589290900] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17201036 kB) free!
[1589291500] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (380708 kB) free!
[1589291711] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.3% (4604620 kB) free.
[1589291854] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 20.8% (10242556 kB) free.
[1589292019] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (8919688 kB) free!
[1589292100] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (17025696 kB) free!
[1589292311] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (3158300 kB) free!
[1589292431] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.5% (4464804 kB) free.
[1589292454] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9558408 kB) free!
[1589292574] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9424244 kB) free!
[1589292619] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 38.2% (18804316 kB) free.
[1589292694] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.9% (9295352 kB) free!
[1589292700] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391780 kB) free!
[1589293031] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.0% (2596712 kB) free!
[1589293151] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.8% (4357264 kB) free.
[1589293211] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1589293300] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14786064 kB) free!
[1589293751] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (1626304 kB) free!
[1589293811] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1589293872] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.5% (3811748 kB) free.
[1589293900] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384100 kB) free!
[1589293932] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1589294048] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.00, 4.45, 4.97
[1589294471] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (1056796 kB) free!
[1589294591] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (3131472 kB) free!
[1589294711] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.8% (4198164 kB) free.
[1589295219] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1589295311] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.9% (1766088 kB) free!
[1589295339] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1589295431] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.2% (3607856 kB) free.
[1589295731] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 703 processes
[1589295816] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.6% (31150224 kB) free.
[1589295851] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1589295937] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1589295971] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1589296031] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (546276 kB) free!
[1589296151] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.4% (2175856 kB) free!
[1589296271] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.8% (3873708 kB) free.
[1589296416] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.0% (25061996 kB) free!
[1589296536] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.6% (17984152 kB) free!
[1589296656] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.1% (31821104 kB) free.
[1589296871] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (285024 kB) free!
[1589296991] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.5% (1872996 kB) free!
[1589297111] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3667676 kB) free.
[1589297171] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1589297291] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 703 processes
[1589297412] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 707 processes
[1589298012] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1589298094] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4495444 kB) free!
[1589298311] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.6% (2216248 kB) free!
[1589298431] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3750764 kB) free.
[1589298612] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1589298732] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1589298852] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1589298852] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 703 processes
[1589299300] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.3% (14905792 kB) free!
[1589299452] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1589299900] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410948 kB) free!
[1589300052] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1589300172] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1589300501] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17224380 kB) free!
[1589301101] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389912 kB) free!
[1589301701] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19553404 kB) free!
[1589301796] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 10.1% (2468976 kB) free!
[1589302031] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (471196 kB) free!
[1589302151] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.3% (1675676 kB) free!
[1589302237] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 23.4% (5738508 kB) free.
[1589302271] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (2934512 kB) free!
[1589302301] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (435180 kB) free!
[1589302656] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1589302901] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19668316 kB) free!
[1589303501] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397388 kB) free!
[1589304101] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20586248 kB) free!
[1589304671] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (910228 kB) free!
[1589304701] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409452 kB) free!
[1589305096] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 10.0% (2456348 kB) free!
[1589305237] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4766188 kB) free!
[1589305271] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (3076908 kB) free!
[1589305301] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20585996 kB) free!
[1589305357] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4747052 kB) free!
[1589305477] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (4724752 kB) free!
[1589305901] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405148 kB) free!
[1589306471] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (952992 kB) free!
[1589306501] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19925024 kB) free!
[1589307072] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (3011444 kB) free!
[1589307101] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396764 kB) free!
[1589307672] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (688388 kB) free!
[1589307701] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19902348 kB) free!
[1589307972] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1589308092] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes