[1552933919] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552934039] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552934213] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11365992 kB) free.
[1552934639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552934759] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552934813] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (3353516 kB) free!
[1552934879] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1552934933] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11351680 kB) free.
[1552935479] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552935533] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.7% (5777024 kB) free!
[1552935599] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552935653] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (4481520 kB) free!
[1552935720] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552935773] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11363060 kB) free.
[1552936373] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6452300 kB) free!
[1552936493] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.5% (5168844 kB) free!
[1552936613] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (3523220 kB) free!
[1552936920] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552937213] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.5% (7169852 kB) free!
[1552937520] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552937640] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552937813] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11361880 kB) free.
[1552938136] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.72, 4.79, 4.93
[1552938240] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552938360] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552938414] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3750288 kB) free!
[1552938480] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552938534] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11348308 kB) free.
[1552939080] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552939134] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.2% (6020588 kB) free!
[1552939201] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552939254] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (4730180 kB) free!
[1552939321] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552939374] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.0% (7382488 kB) free!
[1552939936] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.54, 5.91, 5.33
[1552940056] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.96, 5.89, 5.39
[1552940176] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.38, 6.07, 5.52
[1552940521] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1552940574] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11336500 kB) free.
[1552940734] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24082 MB (11% inode=99%):
[1552941121] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552941174] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (2967808 kB) free!
[1552941241] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1552941294] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11349552 kB) free.
[1552941334] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 16495 MB (7% inode=99%):
[1552941841] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552941894] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5526616 kB) free!
[1552941961] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552942014] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (4105452 kB) free!
[1552942082] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552942134] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11336672 kB) free.
[1552942534] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23068 MB (11% inode=99%):
[1552942735] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.6% (6201720 kB) free!
[1552942855] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (4911796 kB) free!
[1552942975] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (3073180 kB) free!
[1552943134] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 12312 MB (5% inode=99%):
[1552943176] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.46, 4.81, 4.91
[1552943282] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1552943575] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.9% (6874140 kB) free!
[1552943882] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552944002] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1552944175] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11335804 kB) free.
[1552944602] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552944722] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552944775] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (3174968 kB) free!
[1552944842] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1552944895] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11349716 kB) free.
[1552944976] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.62, 6.01, 5.42
[1552945096] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.20, 6.04, 5.50
[1552945216] SERVICE FLAPPING ALERT: sigqscluster2;CPU Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1552945216] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.94, 5.96, 5.53
[1552945443] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552945495] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.6% (5696712 kB) free!
[1552945534] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24108 MB (11% inode=99%):
[1552945563] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552945615] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (4403640 kB) free!
[1552945683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552945735] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11335084 kB) free.
[1552946134] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 14372 MB (6% inode=99%):
[1552946335] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6420476 kB) free!
[1552946456] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.4% (5111968 kB) free!
[1552946576] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (3373192 kB) free!
[1552946883] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1552947176] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.5% (7154480 kB) free!
[1552947483] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552947603] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1552947776] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11331192 kB) free.
[1552947934] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23085 MB (11% inode=99%):
[1552948203] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552948216] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.96, 4.87, 4.96
[1552948324] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552948376] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (3796932 kB) free!
[1552948444] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552948496] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11347284 kB) free.
[1552948535] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 12317 MB (5% inode=99%):
[1552949044] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552949096] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.3% (6041944 kB) free!
[1552949164] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552949216] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (4766356 kB) free!
[1552949284] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552949336] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (2848904 kB) free!
[1552949416] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.48, 5.37, 5.14
[1552949536] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.15, 5.23, 5.12
[1552949656] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.20, 5.19, 5.12
[1552949936] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.7% (6774884 kB) free!
[1552950335] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24082 MB (11% inode=99%):
[1552950484] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552950536] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11344360 kB) free.
[1552950935] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 14349 MB (6% inode=99%):
[1552951084] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552951136] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3071884 kB) free!
[1552951204] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1552951257] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11328904 kB) free.
[1552951805] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552951856] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5627540 kB) free!
[1552951925] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552951976] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (4330436 kB) free!
[1552952045] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552952096] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11346412 kB) free.
[1552952135] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24083 MB (11% inode=99%):
[1552952656] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.86, 4.74, 4.99
[1552952696] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6409748 kB) free!
[1552952735] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 15889 MB (7% inode=99%):
[1552952816] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.4% (5127084 kB) free!
[1552952936] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (3440500 kB) free!
[1552953245] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552953335] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23062 MB (11% inode=99%):
[1552953536] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.5% (7139976 kB) free!
[1552953845] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552953935] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 12310 MB (5% inode=99%):
[1552953965] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1552954136] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11345120 kB) free.
[1552954565] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552954685] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552954736] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3744108 kB) free!
[1552954806] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1552954856] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11324784 kB) free.
[1552955056] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.97, 6.05, 5.42
[1552955136] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24083 MB (11% inode=99%):
[1552955176] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.40, 6.16, 5.53
[1552955296] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.04, 6.34, 5.67
[1552955405] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552955457] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5971464 kB) free!
[1552955525] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552955577] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (4657780 kB) free!
[1552955645] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552955697] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11343592 kB) free.
[1552955736] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 15310 MB (7% inode=99%):
[1552956297] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6638524 kB) free!
[1552956417] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.9% (5349764 kB) free!
[1552956537] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3872316 kB) free!
[1552956845] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552956936] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24087 MB (11% inode=99%):
[1552957137] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.9% (7356448 kB) free!
[1552957445] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552957536] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 15895 MB (7% inode=99%):
[1552957565] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1552957696] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.78, 4.61, 4.86
[1552957737] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11339988 kB) free.
[1552958136] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23064 MB (11% inode=99%):
[1552958166] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552958286] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552958337] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3947076 kB) free!
[1552958406] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552958457] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11324480 kB) free.
[1552958736] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 12820 MB (6% inode=99%):
[1552959057] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6125572 kB) free!
[1552959178] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (4845508 kB) free!
[1552959298] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2968880 kB) free!
[1552959496] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.28, 5.26, 5.06
[1552959606] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1552959616] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.32, 4.97, 4.99
[1552959898] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.8% (6791396 kB) free!
[1552960206] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552960217] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.26, 6.34, 5.60
[1552960326] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552960337] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.24, 6.68, 5.81
[1552960457] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.70, 6.71, 5.93
[1552960498] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11327200 kB) free.
[1552960926] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552961047] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552961098] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (2997356 kB) free!
[1552961136] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23024 MB (11% inode=99%):
[1552961167] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1552961218] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11337800 kB) free.
[1552961736] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 12014 MB (5% inode=99%):
[1552961767] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552961818] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5513320 kB) free!
[1552961887] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552961938] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (4058808 kB) free!
[1552962007] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552962058] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11320944 kB) free.
[1552962658] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6583476 kB) free!
[1552962778] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.7% (5281804 kB) free!
[1552962857] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.51, 4.60, 4.91
[1552962898] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3697816 kB) free!
[1552963207] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552963499] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.9% (7360100 kB) free!
[1552963807] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552963928] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552964099] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11322720 kB) free.
[1552964136] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 23795 MB (11% inode=99%):
[1552964527] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552964647] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552964657] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.62, 5.28, 5.08
[1552964699] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (4057424 kB) free!
[1552964737] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 13809 MB (6% inode=99%):
[1552964767] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552964777] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.22, 5.47, 5.16
[1552964819] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11332812 kB) free.
[1552964897] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.05, 6.14, 5.45
[1552965419] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.6% (6211440 kB) free!
[1552965539] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (4937140 kB) free!
[1552965659] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3145360 kB) free!
[1552965967] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1552966259] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.0% (6922660 kB) free!
[1552966567] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552966687] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552966807] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552967137] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24034 MB (11% inode=99%):
[1552967553] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 61.19, 15.59, 5.79
[1552967673] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.91, 10.69, 5.18
[1552967737] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 13790 MB (6% inode=99%):
[1552967793] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.95, 7.37, 4.63
[1552968059] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10189884 kB) free.
[1552969259] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (3873012 kB) free!
[1552969379] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11309436 kB) free.
[1552969975] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;1;WARNING - 19.8% (1455228 kB) free!
[1552969979] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (3276744 kB) free!
[1552970095] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;2;WARNING - 19.6% (1439480 kB) free!
[1552970099] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11143068 kB) free.
[1552970215] SERVICE ALERT: sig2-vm1;Memory;WARNING;HARD;3;WARNING - 19.4% (1425468 kB) free!
[1552970979] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 8 users currently logged in
[1552971099] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1552971219] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1552971299] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (2846160 kB) free!
[1552971419] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11038096 kB) free.
[1552971475] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2329336 kB) free!
[1552972619] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (2965128 kB) free!
[1552972739] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11133740 kB) free.
[1552973215] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (649088 kB) free!
[1552973339] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (2030160 kB) free!
[1552973407] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1552973460] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10231988 kB) free.
[1552973963] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.4% (1277380 kB) free!
[1552974007] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552974083] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.8% (1236700 kB) free!
[1552974127] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552974203] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.3% (1198160 kB) free!
[1552974247] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552974659] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (2005148 kB) free!
[1552974779] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10308292 kB) free.
[1552975097] SERVICE ALERT: sig4-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (722412 kB) free!
[1552975098] SERVICE FLAPPING ALERT: sigqscluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1552975979] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (1977960 kB) free!
[1552976048] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1552976099] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10353188 kB) free.
[1552976648] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552976768] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552976888] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552977037] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 337 processes
[1552977299] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (3182996 kB) free!
[1552977419] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11234536 kB) free.
[1552978620] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4747188 kB) free!
[1552978740] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11763920 kB) free.
[1552979003] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (732644 kB) free!
[1552979137] SERVICE FLAPPING ALERT: db1;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1552979340] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4490908 kB) free!
[1552979460] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1552979460] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11754692 kB) free.
[1552980060] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (3994468 kB) free!
[1552980180] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11734808 kB) free.
[1552980780] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4071412 kB) free!
[1552980900] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11713824 kB) free.
[1552981500] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4217604 kB) free!
[1552981620] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11691160 kB) free.
[1552982220] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3588700 kB) free!
[1552982288] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1552982341] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11372480 kB) free.
[1552982888] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552983009] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552983129] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552983540] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3853896 kB) free!
[1552983660] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11577552 kB) free.
[1552984260] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (2978720 kB) free!
[1552984380] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10893612 kB) free.
[1552985580] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3497192 kB) free!
[1552985700] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11397004 kB) free.
[1552986300] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (2885600 kB) free!
[1552986420] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10801280 kB) free.
[1552987620] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (3217620 kB) free!
[1552987740] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11273656 kB) free.
[1552988340] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (3163040 kB) free!
[1552988460] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11126240 kB) free.
[1552989060] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (3230012 kB) free!
[1552989132] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1552989180] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11113772 kB) free.
[1552989732] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552989780] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (2899908 kB) free!
[1552989852] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552989900] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10765464 kB) free.
[1552989972] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552990237] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 335 processes
[1552991100] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (3890644 kB) free!
[1552991172] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552991220] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11471136 kB) free.
[1552991772] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552991893] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552992013] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552992421] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (4019136 kB) free!
[1552992541] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11567996 kB) free.
[1552993141] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (3050588 kB) free!
[1552993213] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1552993261] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10844624 kB) free.
[1552993813] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552993933] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552994053] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552994461] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (3295224 kB) free!
[1552994581] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11035844 kB) free.
[1552995181] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3477760 kB) free!
[1552995253] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1552995301] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11281308 kB) free.
[1552995407] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.8% (14755352 kB) free.
[1552995632] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 20.4% (4991896 kB) free.
[1552995637] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 336 processes
[1552995811] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.2% (3254136 kB) free!
[1552995854] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552995901] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (2831580 kB) free!
[1552995942] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 321 processes
[1552995974] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552995974] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.7% (3362428 kB) free!
[1552996021] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10716408 kB) free.
[1552996076] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.9% (4397536 kB) free!
[1552996094] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552996237] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 335 processes
[1552997221] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (220676 kB) free!
[1552997294] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1552997341] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.6% (13122480 kB) free.
[1552997894] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552997941] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.8% (5304084 kB) free!
[1552998014] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552998032] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4911872 kB) free!
[1552998061] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12094636 kB) free.
[1552998152] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4911416 kB) free!
[1552998272] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 20.0% (4910376 kB) free!
[1552998614] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552998661] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (4023548 kB) free!
[1552998735] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552998781] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12136844 kB) free.
[1552999334] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552999381] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (3889200 kB) free!
[1552999454] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1552999502] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11880652 kB) free.
[1553000054] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553000101] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (3347044 kB) free!
[1553000174] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553000221] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11527928 kB) free.
[1553000294] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553001342] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 322 processes
[1553001421] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4318316 kB) free!
[1553001494] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553001541] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12229320 kB) free.
[1553002095] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553002141] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3742592 kB) free!
[1553002215] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1553002261] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11757960 kB) free.
[1553002419] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1553002815] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553002935] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553003055] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553003461] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (4436896 kB) free!
[1553003581] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12326420 kB) free.
[1553004181] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (3880728 kB) free!
[1553004255] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553004302] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11832244 kB) free.
[1553004811] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2390920 kB) free!
[1553004855] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553004976] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553005096] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553005501] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4707312 kB) free!
[1553005621] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12320808 kB) free.
[1553006221] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3965104 kB) free!
[1553006296] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553006341] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11982056 kB) free.
[1553006896] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553006941] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8255924 kB) free!
[1553007016] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553007061] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11356424 kB) free.
[1553007136] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553008261] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4281696 kB) free!
[1553008337] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553008381] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12162916 kB) free.
[1553008937] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553008982] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3525880 kB) free!
[1553009057] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553009102] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11550140 kB) free.
[1553009657] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553009777] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553009898] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553010302] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4564236 kB) free!
[1553010422] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12263128 kB) free.
[1553011022] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4122520 kB) free!
[1553011098] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553011142] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11986988 kB) free.
[1553011698] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553011742] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3468960 kB) free!
[1553011818] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553011862] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11390196 kB) free.
[1553011938] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553013062] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4575940 kB) free!
[1553013138] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1553013182] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12250488 kB) free.
[1553013739] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553013783] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7206152 kB) free!
[1553013859] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553013903] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.1% (5966184 kB) free!
[1553013979] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553014023] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4306084 kB) free!
[1553014538] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26988 MB (12% inode=99%):
[1553014623] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.3% (8054328 kB) free!
[1553015179] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553015223] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12258472 kB) free.
[1553015738] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 20854 MB (9% inode=99%):
[1553015779] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553015823] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.6% (5246400 kB) free!
[1553015899] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553015943] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12273616 kB) free.
[1553016019] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553016543] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7353556 kB) free!
[1553016620] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553016663] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.5% (6150796 kB) free!
[1553016740] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553016783] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1553016783] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4928292 kB) free!
[1553016860] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553016938] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 27006 MB (12% inode=99%):
[1553017383] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.9% (8348576 kB) free!
[1553017984] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12269452 kB) free.
[1553018060] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553018138] SERVICE FLAPPING ALERT: db1;HDD_Home status;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1553018138] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 20216 MB (9% inode=99%):
[1553018584] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5636308 kB) free!
[1553018660] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553018704] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (3620508 kB) free!
[1553018738] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 22962 MB (11% inode=99%):
[1553018780] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553018824] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12251144 kB) free.
[1553019380] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553019424] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7027020 kB) free!
[1553019501] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553019544] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.7% (5779492 kB) free!
[1553019621] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553019664] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3910856 kB) free!