[1577731070] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.3% (20106816 kB) free!
[1577731106] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195036 kB) free!
[1577731226] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184700 kB) free!
[1577731288] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189080 kB) free!
[1577731347] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187688 kB) free!
[1577731350] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577731375] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (1209324 kB) free!
[1577731408] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193560 kB) free!
[1577731495] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (1197532 kB) free!
[1577731528] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184944 kB) free!
[1577731546] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187836 kB) free!
[1577731615] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (1191500 kB) free!
[1577731667] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170992 kB) free!
[1577731787] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192648 kB) free!
[1577731950] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(15) < 2019/12/31 00:21:19 [error] 20757#0: *21288693 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577732070] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577732671] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2019/12/31 00:31:41 [error] 20752#0: *21290414 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577732791] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577733148] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 53.3% (13144596 kB) free.
[1577733391] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(9) < 2019/12/31 00:45:38 [error] 20757#0: *21292666 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577733511] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2019/12/31 00:46:41 [error] 20756#0: *21292855 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577733631] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577734231] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(9) < 2019/12/31 00:57:57 [error] 20757#0: *21294685 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577734351] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2019/12/31 01:00:38 [error] 20757#0: *21295119 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577734472] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577734948] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (1933128 kB) free!
[1577735068] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (687352 kB) free!
[1577735071] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(6) < 2019/12/31 01:11:41 [error] 20757#0: *21296916 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577735188] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1577735188] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182824 kB) free!
[1577735191] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(6) < 2019/12/31 01:15:42 [error] 20754#0: *21297581 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577735311] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1577735311] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(3) < 2019/12/31 01:17:57 [error] 20757#0: *21297949 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577735530] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9384776 kB) free!
[1577735650] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9390556 kB) free!
[1577735770] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (9281148 kB) free!
[1577737187] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1577737187] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.4% (10392560 kB) free.
[1577737826] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.3% (22765692 kB) free.
[1577737890] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.6% (22458040 kB) free.
[1577737956] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.8% (22979220 kB) free.
[1577738003] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577738607] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/12/31 02:11:41 [error] 20757#0: *21305688 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577738727] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577739035] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198064 kB) free!
[1577739155] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195356 kB) free!
[1577739275] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188580 kB) free!
[1577739307] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (233652 kB) free!
[1577739328] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(15) < 2019/12/31 02:22:57 [error] 20757#0: *21307523 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577739427] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (217924 kB) free!
[1577739448] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577739547] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (212216 kB) free!
[1577739597] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 67.7% (33382388 kB) free.
[1577740048] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2019/12/31 02:36:41 [error] 20751#0: *21309789 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577740147] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 82.8% (40801436 kB) free.
[1577740168] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577740768] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2019/12/31 02:46:41 [error] 20753#0: *21311410 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577740888] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577741488] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2019/12/31 03:01:19 [error] 20752#0: *21313794 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577741609] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577742208] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2047) < 2019/12/31 03:12:58 [error] 20752#0: *21317020 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:117/0, bytes from/to upstream:0/117
[1577742216] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(40) < 49.50.86.152 [31/Dec/2019:03:10:17 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1577742283] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(47) < 2019/12/31 03:10:41 [error] 28863#0: *29977272 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1577742286] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1676) < 203.212.70.170 [31/Dec/2019:03:11:32 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1577742328] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577742338] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577742404] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577742406] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577742928] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(13) < 2019/12/31 03:22:57 [error] 20751#0: *21320558 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577743048] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2) < 2019/12/31 03:26:41 [error] 20755#0: *21321180 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577743134] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1577743168] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577743254] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1577743787] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1577744368] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(20) < 2019/12/31 03:46:19 [error] 20747#0: *21324414 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577744488] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577745088] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(13) < 2019/12/31 04:00:44 [error] 20761#0: *21326777 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577745208] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2019/12/31 04:01:41 [error] 20756#0: *21326933 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577745328] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577745928] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/12/31 04:12:57 [error] 20753#0: *21328762 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577746024] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.5% (22817468 kB) free.
[1577746049] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(8) < 2019/12/31 04:15:45 [error] 20751#0: *21329250 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577746058] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577746068] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1577746668] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(7) < 2019/12/31 04:26:19 [error] 20747#0: *21330976 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577746788] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/12/31 04:27:57 [error] 20747#0: *21331242 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577746908] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(10) < 2019/12/31 04:30:44 [error] 20760#0: *21331716 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577751427] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1577752787] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189104 kB) free!
[1577752907] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196632 kB) free!
[1577752966] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171288 kB) free!
[1577753027] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (210964 kB) free!
[1577753086] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 88.8% (21758212 kB) free.
[1577753201] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.2% (22840220 kB) free.
[1577753270] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 42.2% (55639880 kB) free.
[1577761254] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1577761374] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1577763174] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1577763294] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1577770494] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1577770614] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1577771170] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4073648 kB) free!
[1577771827] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.8% (3403348 kB) free!
[1577771947] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.4% (3314392 kB) free!
[1577772067] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2352092 kB) free!
[1577773404] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.8% (3537724 kB) free!
[1577773524] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.4% (3223116 kB) free!
[1577773616] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8882884 kB) free!
[1577773644] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.3% (2996512 kB) free!
[1577773736] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.4% (8573900 kB) free!
[1577773857] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 53.9% (26549700 kB) free.
[1577774086] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.8% (3860636 kB) free!
[1577774206] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.0% (3677752 kB) free!
[1577774244] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 76.3% (17141492 kB) free.
[1577774326] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.3% (3506840 kB) free!
[1577775406] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1287) < 103.255.146.164 [31/Dec/2019:12:17:06 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1577775526] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2327396 kB) free!
[1577775526] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577778502] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.9% (22689796 kB) free.
[1577778527] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.3% (22627916 kB) free.
[1577778579] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.5% (38662644 kB) free.
[1577779706] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.9% (3918372 kB) free!
[1577779826] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.4% (3791024 kB) free!
[1577779946] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.9% (3687204 kB) free!
[1577780811] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4665872 kB) free!
[1577780931] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.6% (4567372 kB) free!
[1577781051] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.0% (4415756 kB) free!
[1577781146] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1577781146] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1294968 kB) free!
[1577781272] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1577781392] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1577782251] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2387248 kB) free!
[1577782858] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 20.0% (4908424 kB) free!
[1577782978] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.7% (4839508 kB) free!
[1577783098] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.6% (4794428 kB) free!
[1577783947] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8765948 kB) free!
[1577784067] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.3% (8514564 kB) free!
[1577784187] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.8% (8281648 kB) free!
[1577784444] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.4% (3672644 kB) free!
[1577784564] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.5% (3487416 kB) free!
[1577784684] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 76.3% (17137784 kB) free.
[1577784695] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4788840 kB) free!
[1577784816] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (4672396 kB) free!
[1577784936] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (4592144 kB) free!
[1577785592] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 4 warn out of 335 processes
[1577785712] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1577786313] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1577786433] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1577786587] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4561060 kB) free!
[1577786698] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2120652 kB) free!
[1577787597] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9771840 kB) free!
[1577787717] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (9688020 kB) free!
[1577787837] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (9603540 kB) free!
[1577788536] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2196392 kB) free!
[1577790586] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9258576 kB) free!
[1577790706] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.8% (9259280 kB) free!
[1577790826] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (9258300 kB) free!
[1577791233] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1577791353] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1577791473] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1577791946] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1577792637] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4521224 kB) free!
[1577792919] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 90.1% (22086980 kB) free.
[1577792937] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (4442864 kB) free!
[1577792957] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.7% (22137420 kB) free.
[1577792990] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.8% (22256124 kB) free.
[1577793030] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 93.1% (45868224 kB) free.
[1577793057] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.9% (38880220 kB) free.
[1577793095] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.0% (45330140 kB) free.
[1577793120] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.2% (21862208 kB) free.
[1577793151] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.9% (22045680 kB) free.
[1577793262] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1577793382] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1577793502] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1577794169] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (1129364 kB) free!
[1577794289] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (1029040 kB) free!
[1577794409] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (931968 kB) free!
[1577794884] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.1% (3611764 kB) free!
[1577795004] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.4% (3463776 kB) free!
[1577795124] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.7% (3291564 kB) free!
[1577795640] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26377264 kB) free!
[1577795710] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 20.0% (4896276 kB) free!
[1577795724] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 76.2% (17114728 kB) free.
[1577795760] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.6% (48199740 kB) free.
[1577795830] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.9% (4871456 kB) free!
[1577795950] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.8% (4846408 kB) free!
[1577796960] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (26008216 kB) free!
[1577797080] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.9% (51241528 kB) free.
[1577797680] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (23547180 kB) free!
[1577797800] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.2% (47686884 kB) free.
[1577799404] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(10) < 2019/12/31 19:06:22 [error] 28871#0: *30096315 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:457759/6798101, bytes from/to upstream:6798101/457759
[1577799524] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/12/31 19:06:47 [error] 28869#0: *30426260 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1577799644] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577801673] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1577801793] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1577803266] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9175372 kB) free!
[1577803386] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.7% (8740472 kB) free!
[1577803506] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.3% (9021964 kB) free!
[1577804302] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1577806126] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1005) < 103.255.146.164 [31/Dec/2019:20:49:58 +0530] TCP 502 0 0 0.000 "backend_3000" "0" "0" "0.000"
[1577806247] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577813526] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4864288 kB) free!
[1577813646] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (4821896 kB) free!
[1577813767] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.5% (4790800 kB) free!
[1577814657] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (212800 kB) free!
[1577814777] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (213048 kB) free!
[1577814897] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (241364 kB) free!
[1577815014] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1577815070] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (1116036 kB) free!
[1577815134] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1577815190] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (378844 kB) free!
[1577815311] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (818644 kB) free!