[1563388220] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563388487] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1563388508] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1563388828] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199024 kB) free!
[1563388885] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (200496 kB) free!
[1563388948] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197000 kB) free!
[1563389005] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (203272 kB) free!
[1563389015] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189304 kB) free!
[1563389068] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191704 kB) free!
[1563389125] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199800 kB) free!
[1563389135] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (200424 kB) free!
[1563389255] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (208900 kB) free!
[1563392124] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (2113640 kB) free!
[1563392184] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (2112092 kB) free!
[1563392244] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 8.6% (2113592 kB) free!
[1563392304] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 8.6% (2116296 kB) free!
[1563394748] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.9% (22187716 kB) free.
[1563394812] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 83.4% (20434188 kB) free.
[1563394897] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.9% (38900060 kB) free.
[1563395006] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.0% (21079164 kB) free.
[1563395169] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 62.5% (15326220 kB) free.
[1563395225] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.2% (14329428 kB) free.
[1563395246] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 351 processes
[1563395367] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1563395508] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (250240 kB) free!
[1563395628] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (228612 kB) free!
[1563395748] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (262188 kB) free!
[1563395825] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (181544 kB) free!
[1563395945] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (172752 kB) free!
[1563396065] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (173920 kB) free!
[1563396088] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (213464 kB) free!
[1563396208] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (219668 kB) free!
[1563396328] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (225188 kB) free!
[1563397015] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.37, 15.89, 18.15
[1563397136] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.45, 10.76, 16.00
[1563397206] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.8% (3218084 kB) free!
[1563397256] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.19, 7.27, 14.08
[1563397326] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.8% (3213948 kB) free!
[1563397446] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.7% (3205864 kB) free!
[1563398040] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 13908.98, 6730.71, 2635.32
[1563398160] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 14482.89, 9325.29, 4082.66
[1563398280] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14560.62, 11061.53, 5354.58
[1563398456] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.30, 0.45, 4.13
[1563398943] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 54.87, 68.58, 30.36
[1563399063] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 7.62, 45.98, 26.71
[1563399183] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.20, 30.83, 23.50
[1563399301] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 4551 MB (8% inode=97%):
[1563399783] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.29, 4.40, 12.44
[1563399901] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6220 MB (12% inode=97%):
[1563400983] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.24, 0.35, 3.61
[1563404545] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 514 processes
[1563404605] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 506 processes
[1563404880] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.45, 2.48, 14.29
[1563406089] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.6% (39727012 kB) free.
[1563406137] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.7% (14744972 kB) free.
[1563406247] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 69.9% (11355584 kB) free.
[1563406680] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.63, 2.53, 4.23
[1563406705] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563406759] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9794496 kB) free!
[1563406826] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563406879] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.2% (8496800 kB) free!
[1563406946] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563406999] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12077156 kB) free.
[1563408199] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9626304 kB) free!
[1563408319] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8252508 kB) free!
[1563408439] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12060232 kB) free.
[1563409346] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563409639] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9175272 kB) free!
[1563409759] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7814344 kB) free!
[1563409879] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11917912 kB) free.
[1563409946] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563410066] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563410186] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563410787] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1563410787] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563411079] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9096372 kB) free!
[1563411199] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7662908 kB) free!
[1563411246] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1565600 kB) free!
[1563411319] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12018984 kB) free.
[1563411920] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7635096 kB) free!
[1563412040] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12025552 kB) free.
[1563412587] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563412640] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7387436 kB) free!
[1563412707] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563412760] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12031432 kB) free.
[1563413307] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563413360] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7075376 kB) free!
[1563413427] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563413480] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12020892 kB) free.
[1563414027] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563414148] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563414680] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8419560 kB) free!
[1563414747] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563414800] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12027540 kB) free.
[1563414867] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563415400] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8140548 kB) free!
[1563415467] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563415520] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12018124 kB) free.
[1563415587] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563416721] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8150348 kB) free!
[1563416787] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563416841] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12028916 kB) free.
[1563416907] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563417441] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8152644 kB) free!
[1563417508] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563417561] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12027180 kB) free.
[1563417628] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563418161] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8075684 kB) free!
[1563418228] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563418281] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12018644 kB) free.
[1563418348] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563418468] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563418695] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 10.4% (5131388 kB) free!
[1563418744] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (181436 kB) free!
[1563418815] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (216648 kB) free!
[1563418864] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (166240 kB) free!
[1563418935] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (225576 kB) free!
[1563418984] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (159460 kB) free!
[1563419481] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9325988 kB) free!
[1563419601] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7953380 kB) free!
[1563419668] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563419721] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12002872 kB) free.
[1563420268] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563420389] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563420509] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563420921] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8864288 kB) free!
[1563421042] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7128188 kB) free!
[1563421109] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563421162] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11388932 kB) free.
[1563422043] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.99, 7.58, 5.68
[1563422163] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.16, 7.35, 5.82
[1563422283] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.43, 7.12, 5.93
[1563422758] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 17.3% (1275452 kB) free!
[1563422878] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 15.0% (1102240 kB) free!
[1563422909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563422998] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 13.0% (958168 kB) free!
[1563423029] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563423562] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8199512 kB) free!
[1563423598] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (250468 kB) free!
[1563423629] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563423656] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 749.75, 233.65, 82.19
[1563423682] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11985992 kB) free.
[1563423750] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563423776] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1009.03, 503.57, 199.46
[1563423783] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1025.31, 688.66, 308.33
[1563423896] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1043.84, 684.12, 302.50
[1563423903] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1029.00, 801.45, 395.77
[1563424023] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 249.63, 622.10, 382.61
[1563424282] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8231304 kB) free!
[1563424349] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563424402] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11992756 kB) free.
[1563424469] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563424798] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 89.3% (6567988 kB) free.
[1563425002] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8271028 kB) free!
[1563425069] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563425122] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11985720 kB) free.
[1563425189] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563425722] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8222908 kB) free!
[1563425789] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563425843] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11993220 kB) free.
[1563425909] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563426019] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (8946020 kB) free!
[1563426140] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (18058640 kB) free.
[1563426442] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8247888 kB) free!
[1563426509] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563426562] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11994884 kB) free.
[1563426630] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563426739] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (6966408 kB) free!
[1563426859] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (17938808 kB) free.
[1563426994] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.0% (38952136 kB) free.
[1563427023] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.46, 0.58, 15.81
[1563427082] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 81.4% (13228612 kB) free.
[1563427132] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.6% (14079452 kB) free.
[1563427162] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8251184 kB) free!
[1563427229] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563427282] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11983624 kB) free.
[1563427349] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563427459] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (10723692 kB) free!
[1563427496] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.04, 0.97, 11.54
[1563427579] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (6364728 kB) free!
[1563427599] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (243348 kB) free!
[1563427699] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (17806376 kB) free.
[1563427719] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (230360 kB) free!
[1563427738] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (168316 kB) free!
[1563427839] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (226712 kB) free!
[1563427857] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (314440 kB) free!
[1563427858] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (173612 kB) free!
[1563427882] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8156964 kB) free!
[1563427949] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563427977] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (171724 kB) free!
[1563427978] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (172204 kB) free!
[1563428002] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11993244 kB) free.
[1563428069] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563428097] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (171028 kB) free!
[1563428299] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (10987212 kB) free!
[1563428419] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (6203344 kB) free!
[1563428539] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (16656720 kB) free.
[1563428602] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7999600 kB) free!
[1563428696] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.69, 0.91, 3.96
[1563428722] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11980664 kB) free.
[1563428823] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.38, 0.58, 2.95
[1563429139] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (11827852 kB) free!
[1563429259] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.3% (7461824 kB) free!
[1563429379] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (3034592 kB) free!
[1563429796] SERVICE ALERT: rq-2;HDD_VNS1 status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_VNS1 status' on host 'rq-2' was out of bounds)
[1563429861] SERVICE ALERT: rq-2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'rq-2' was out of bounds)
[1563429916] SERVICE ALERT: rq-2;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /dev 64349 MB (100% inode=99%):
[1563429979] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.3% (12037580 kB) free!
[1563429981] SERVICE ALERT: rq-2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1563429997] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (173340 kB) free!
[1563430071] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (3139764 kB) free!
[1563430117] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (118160 kB) free!
[1563430191] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.3% (2647244 kB) free!
[1563430237] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1563430237] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (133684 kB) free!
[1563430308] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (4551360 kB) free!
[1563430311] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 13.2% (2143564 kB) free!
[1563430338] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.39, 6.46, 5.75
[1563430428] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4402092 kB) free!
[1563430458] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.04, 6.66, 5.91
[1563430548] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.4% (4265852 kB) free!
[1563430578] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.63, 7.16, 6.19
[1563430811] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.3% (2810184 kB) free!
[1563430911] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (170424 kB) free!
[1563430931] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 15.8% (2565764 kB) free!
[1563431051] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.3% (2323864 kB) free!
[1563431116] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178300 kB) free!
[1563431179] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (1868208 kB) free!
[1563431236] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185904 kB) free!
[1563431356] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188152 kB) free!
[1563431438] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.0% (6767604 kB) free.
[1563431652] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 85.3% (13856260 kB) free.
[1563431753] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198880 kB) free!
[1563431779] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1563431779] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.9% (13066196 kB) free!
[1563432215] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (202612 kB) free!
[1563432335] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179180 kB) free!
[1563432455] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164640 kB) free!
[1563432897] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.4% (2011648 kB) free!
[1563432979] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (2585616 kB) free!
[1563433579] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (14759760 kB) free.
[1563434097] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (1488460 kB) free!
[1563434179] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (9222008 kB) free!
[1563434299] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (4754644 kB) free!
[1563434420] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (15978420 kB) free.
[1563434697] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1563434697] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.4% (1849272 kB) free!
[1563435019] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (12354252 kB) free!
[1563435139] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.2% (7379320 kB) free!
[1563435259] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (2885836 kB) free!
[1563435738] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.81, 7.50, 6.41
[1563435858] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.15, 7.73, 6.63
[1563435859] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (15926268 kB) free.
[1563435897] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (1517980 kB) free!
[1563435979] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.93, 7.38, 6.63
[1563436208] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 366 processes
[1563436328] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 362 processes
[1563436448] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 361 processes
[1563436459] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (12482408 kB) free!
[1563436579] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.7% (7689112 kB) free!
[1563436699] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (3414224 kB) free!
[1563436838] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (131504 kB) free!
[1563436959] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (123328 kB) free!
[1563437048] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 362 processes
[1563437052] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 15.3% (2487560 kB) free!
[1563437079] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (129240 kB) free!
[1563437168] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 364 processes
[1563437173] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 13.6% (2216944 kB) free!
[1563437288] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 360 processes
[1563437293] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 12.0% (1953044 kB) free!
[1563437670] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563437790] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563437893] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1563437893] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (662832 kB) free!
[1563437899] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (15608940 kB) free.
[1563437910] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563438378] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 5.29, 4.77, 4.91
[1563438499] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (11375496 kB) free!
[1563438619] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (6426680 kB) free!
[1563438739] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (1856500 kB) free!
[1563438978] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.72, 5.40, 5.12
[1563439088] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1563439093] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.7% (2389476 kB) free!
[1563439098] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.83, 5.50, 5.19
[1563439218] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.53, 5.42, 5.19
[1563439339] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (16815528 kB) free.
[1563439647] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 9.22, 33.19, 19.93
[1563439688] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1563439693] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (1439112 kB) free!
[1563439768] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.57, 23.50, 17.99
[1563439808] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1563439888] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.02, 16.71, 16.17
[1563439940] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (13122864 kB) free!
[1563440060] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (8405400 kB) free!
[1563440179] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.53, 2.43, 3.76
[1563440180] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (4031540 kB) free!
[1563440223] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.20, 6.71, 6.63
[1563440310] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563440343] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.23, 4.94, 6.00
[1563440463] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.49, 3.45, 5.32
[1563440722] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7154816 kB) free!
[1563440843] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11401824 kB) free.
[1563440910] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563441031] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563441063] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.56, 0.92, 3.04
[1563441112] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 35.2% (5718496 kB) free.
[1563441151] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563441284] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.6% (21378952 kB) free.
[1563441380] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (15381704 kB) free.
[1563441418] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 88.9% (21794252 kB) free.
[1563441460] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 34.0% (16747520 kB) free.
[1563441499] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 82.0% (13328084 kB) free.
[1563441542] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.5% (21434856 kB) free.
[1563441593] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.1% (14646204 kB) free.
[1563441608] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 364 processes
[1563441618] SERVICE FLAPPING ALERT: smscapp1;CPU Load;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1563441618] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 3.04, 3.67, 4.78
[1563441638] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 62.2% (10104276 kB) free.
[1563441676] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.7% (6820272 kB) free.
[1563441728] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1563441980] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (11242752 kB) free!
[1563442100] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.2% (6679780 kB) free!
[1563442197] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (174212 kB) free!
[1563442220] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (2638688 kB) free!
[1563442288] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.22, 3.41, 4.26
[1563442317] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (148884 kB) free!
[1563442438] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (178956 kB) free!
[1563442487] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179020 kB) free!
[1563442607] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189404 kB) free!
[1563442630] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166608 kB) free!
[1563442642] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8091876 kB) free!
[1563442669] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (219672 kB) free!
[1563442727] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181568 kB) free!
[1563442750] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181824 kB) free!
[1563442750] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163556 kB) free!
[1563442762] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11980436 kB) free.
[1563442789] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (239452 kB) free!
[1563442820] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.4% (12735856 kB) free!
[1563442870] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188376 kB) free!
[1563442870] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187968 kB) free!
[1563442909] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (219248 kB) free!
[1563442912] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.4% (2670840 kB) free!
[1563442951] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1563442951] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563442990] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172600 kB) free!
[1563443032] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.4% (2503556 kB) free!
[1563443152] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 14.2% (2311648 kB) free!
[1563443362] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8150852 kB) free!
[1563443447] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (855568 kB) free!
[1563443482] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11969408 kB) free.
[1563443567] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (217844 kB) free!
[1563443687] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (173988 kB) free!
[1563443752] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1260112 kB) free!
[1563444020] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (6094072 kB) free!
[1563444082] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8306148 kB) free!
[1563444151] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563444202] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11977476 kB) free.
[1563444271] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563444802] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8257960 kB) free!
[1563444872] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563444922] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11964896 kB) free.
[1563444992] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563445127] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.5% (21603964 kB) free.
[1563445156] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.7% (21979676 kB) free.
[1563445190] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (20775432 kB) free.
[1563445230] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 74.6% (36742552 kB) free.
[1563445268] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1563445268] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 83.8% (13620744 kB) free.
[1563445328] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1563445382] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.2% (22102556 kB) free.
[1563445437] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.9% (14453832 kB) free.
[1563445448] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 356 processes
[1563445479] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 61.5% (9988560 kB) free.
[1563445496] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 14623.59, 11335.23, 5573.10
[1563445523] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8335992 kB) free!
[1563445568] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1563445592] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563445616] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 14631.01, 12426.09, 6671.50
[1563445643] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11973456 kB) free.
[1563445712] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563445727] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (840772 kB) free!
[1563445737] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14631.84, 13155.96, 7636.70
[1563445761] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187684 kB) free!
[1563445839] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (235616 kB) free!
[1563445847] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168788 kB) free!
[1563445882] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195116 kB) free!
[1563445959] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (214676 kB) free!
[1563445967] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1563445967] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188528 kB) free!
[1563446002] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1563446002] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189432 kB) free!
[1563446037] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (168556 kB) free!
[1563446079] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1563446079] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (226388 kB) free!
[1563446090] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.5% (1712184 kB) free!
[1563446157] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (179328 kB) free!
[1563446210] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (1196016 kB) free!
[1563446243] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8309300 kB) free!
[1563446278] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1563446278] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166316 kB) free!
[1563446312] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563446330] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (533332 kB) free!
[1563446363] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11974460 kB) free.
[1563446432] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563446963] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8422332 kB) free!
[1563447032] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563447083] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11968356 kB) free.
[1563447152] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563447530] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.8% (5012324 kB) free.
[1563447598] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (12516132 kB) free!
[1563447683] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8588164 kB) free!
[1563447718] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.9% (8480932 kB) free!
[1563447753] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563447803] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7883268 kB) free!
[1563447838] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.7% (21474836 kB) free.
[1563447873] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563447923] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11962840 kB) free.
[1563447968] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 362 processes
[1563448088] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1563448130] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (1387812 kB) free!
[1563448208] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1563448250] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 53.5% (8692096 kB) free.
[1563448281] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1563448473] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563448523] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7055292 kB) free!
[1563448593] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563448643] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11969456 kB) free.
[1563448808] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1563448808] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1563448850] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (895892 kB) free!
[1563448970] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.5% (2196444 kB) free!
[1563449090] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.4% (2174160 kB) free!
[1563449193] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563449243] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (6990772 kB) free!
[1563449312] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1563449313] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563449364] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11967808 kB) free.
[1563449638] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (11386052 kB) free!
[1563449690] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (433904 kB) free!
[1563449758] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (15039504 kB) free.
[1563449913] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563449963] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7279740 kB) free!
[1563450033] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563450083] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11960620 kB) free.
[1563450634] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563450682] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 18.1% (1332624 kB) free!
[1563450683] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7183588 kB) free!
[1563450754] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563450802] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 16.5% (1211080 kB) free!
[1563450803] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11967300 kB) free.
[1563450922] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 15.2% (1116980 kB) free!
[1563451218] SERVICE FLAPPING ALERT: smscapp1;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1563451354] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563451403] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7055716 kB) free!
[1563451474] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563451522] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (544852 kB) free!
[1563451523] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11959692 kB) free.
[1563451737] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.75, 0.71, 17.69
[1563452151] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1563452159] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (11008172 kB) free!
[1563452279] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.8% (7091412 kB) free!
[1563452399] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (20317192 kB) free.
[1563452403] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 343 processes
[1563452480] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.5% (3173164 kB) free!
[1563452523] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563452600] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.7% (3044448 kB) free!
[1563452674] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563452720] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.2% (2953144 kB) free!
[1563452723] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7478504 kB) free!
[1563452794] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563452843] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11965832 kB) free.
[1563453323] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.0% (6838036 kB) free.
[1563453394] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563453443] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7390920 kB) free!
[1563453514] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563453537] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.43, 2.01, 4.40
[1563453563] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11963100 kB) free.
[1563453608] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1563453728] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1563453848] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1563453920] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1521060 kB) free!
[1563454115] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563454163] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7256988 kB) free!
[1563454199] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (10907420 kB) free!
[1563454235] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563454284] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11953564 kB) free.
[1563454319] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.7% (7048856 kB) free!
[1563454439] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (19613176 kB) free.
[1563454448] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1563454551] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1563454803] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.83, 4.31, 4.75
[1563454835] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1563455435] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563455555] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563455567] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 34.3% (8466184 kB) free.
[1563455639] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (12987680 kB) free!
[1563455759] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.8% (9043964 kB) free!
[1563455879] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (4594836 kB) free!
[1563456003] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.85, 6.32, 5.47
[1563456083] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8518260 kB) free!
[1563456124] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.80, 5.86, 5.41
[1563456155] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563456167] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193480 kB) free!
[1563456203] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11961440 kB) free.
[1563456244] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.16, 6.68, 5.79
[1563456275] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563456279] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1563456287] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194372 kB) free!
[1563456407] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182720 kB) free!
[1563456479] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (18959252 kB) free.
[1563456802] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563456803] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8666240 kB) free!
[1563456875] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563456923] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.2% (7011836 kB) free!
[1563456996] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563457043] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11959712 kB) free.
[1563457595] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563457643] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6928448 kB) free!
[1563457678] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563457715] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563457737] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.63, 5.63, 5.02
[1563457763] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11958020 kB) free.
[1563457857] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.01, 5.83, 5.17
[1563457977] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.15, 5.95, 5.29
[1563458279] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (8531368 kB) free!
[1563458315] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563458363] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6958940 kB) free!
[1563458400] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (4013600 kB) free!
[1563458435] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563458483] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11950460 kB) free.
[1563458520] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (18537848 kB) free.
[1563459035] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563459155] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563459683] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8505048 kB) free!
[1563459720] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (11528240 kB) free!
[1563459755] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563459803] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11959244 kB) free.
[1563459840] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.8% (7073948 kB) free!
[1563459876] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563459960] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (3087884 kB) free!
[1563460403] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8457656 kB) free!
[1563460475] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563460523] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11948272 kB) free.
[1563460560] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (17688516 kB) free.
[1563460595] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563461044] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.09, 3.63, 4.29
[1563461090] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563461123] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8296096 kB) free!
[1563461195] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563461243] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11958140 kB) free.
[1563461315] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563461760] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (11472656 kB) free!
[1563461843] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8323804 kB) free!
[1563461880] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.6% (7636700 kB) free!
[1563461915] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563461963] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11955988 kB) free.
[1563462000] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (4176912 kB) free!
[1563462035] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563462563] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8228876 kB) free!
[1563462600] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (17048900 kB) free.
[1563462635] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563462647] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1563462683] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11947832 kB) free.
[1563462756] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563462767] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1563463283] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8320596 kB) free!
[1563463355] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563463404] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11955304 kB) free.
[1563463475] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563463800] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (12095688 kB) free!
[1563463920] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.5% (8234144 kB) free!
[1563464003] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8282192 kB) free!
[1563464041] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (4520264 kB) free!
[1563464075] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563464123] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11946396 kB) free.
[1563464195] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563464641] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (16679692 kB) free.
[1563464720] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563464723] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8108196 kB) free!
[1563464795] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563464843] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11951884 kB) free.
[1563464915] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563465443] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (8002868 kB) free!
[1563465563] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11952476 kB) free.
[1563465841] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (12919568 kB) free!
[1563465848] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563465961] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.7% (9662668 kB) free!
[1563465979] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 64.3% (10458504 kB) free.
[1563466034] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 65.8% (16242396 kB) free.
[1563466066] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 43.4% (21384916 kB) free.
[1563466081] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (6264472 kB) free!
[1563466095] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 79.8% (19569392 kB) free.
[1563466116] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563466124] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 84.0% (13649456 kB) free.
[1563466156] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.2% (14504648 kB) free.
[1563466163] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7891340 kB) free!
[1563466186] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 48.5% (31914828 kB) free.
[1563466236] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563466283] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11943888 kB) free.
[1563466638] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (631728 kB) free!
[1563466666] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (214052 kB) free!
[1563466758] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187912 kB) free!
[1563466760] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (156548 kB) free!
[1563466786] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236400 kB) free!
[1563466836] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563466878] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170324 kB) free!
[1563466880] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (164256 kB) free!
[1563466883] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7743468 kB) free!
[1563466907] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (250008 kB) free!
[1563466956] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563467000] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (185344 kB) free!
[1563467003] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11951920 kB) free.
[1563467556] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563467603] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7971780 kB) free!
[1563467676] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563467724] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11941968 kB) free.
[1563468276] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563468323] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7817360 kB) free!
[1563468396] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563468443] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11951744 kB) free.
[1563468997] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563469043] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7552888 kB) free!
[1563469117] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563469163] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11952040 kB) free.
[1563469717] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563469763] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7319732 kB) free!
[1563469837] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563469883] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11944128 kB) free.
[1563470437] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563470483] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7052124 kB) free!
[1563470557] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563470603] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11950740 kB) free.
[1563471157] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563471203] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7350536 kB) free!
[1563471277] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563471323] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11944508 kB) free.
[1563471877] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563471923] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7616464 kB) free!
[1563471998] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563472044] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11949076 kB) free.
[1563472186] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (13161056 kB) free!
[1563472306] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.5% (23350540 kB) free.
[1563472597] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563472643] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7520988 kB) free!
[1563472717] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563472763] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11953024 kB) free.
[1563473317] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563473363] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7523256 kB) free!
[1563473437] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563473483] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11939700 kB) free.
[1563474037] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563474083] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7653960 kB) free!
[1563474157] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563474203] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11951364 kB) free.