[1559327454] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 87.9% (14292212 kB) free.
[1559327524] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 400 processes
[1559327660] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 2 warn out of 616 processes
[1559327720] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 616 processes
[1559327880] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47147 MB (20% inode=99%):
[1559328000] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47147 MB (20% inode=99%):
[1559328020] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 621 processes
[1559328080] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 617 processes
[1559328120] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47145 MB (20% inode=99%):
[1559328124] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 429 processes
[1559328980] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 612 processes
[1559329040] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 613 processes
[1559329100] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 0 warn out of 615 processes
[1559329160] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 610 processes
[1559329364] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1559329861] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1559330102] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (988820 kB) free!
[1559330360] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 617 processes
[1559330420] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 613 processes
[1559330520] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49370 MB (21% inode=99%):
[1559330814] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (2550596 kB) free!
[1559331347] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1559331902] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.8% (2728084 kB) free!
[1559331947] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 362 processes
[1559331962] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.3% (6043364 kB) free!
[1559332067] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1559332082] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13182080 kB) free.
[1559332614] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2452156 kB) free!
[1559332668] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559332682] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.1% (5489720 kB) free!
[1559332788] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559332802] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13164068 kB) free.
[1559333102] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (167388 kB) free!
[1559333388] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559333402] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7624152 kB) free!
[1559333508] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1559333522] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.7% (5292620 kB) free!
[1559333524] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 383 processes
[1559333628] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1559333642] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13178096 kB) free.
[1559333702] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.5% (2354468 kB) free!
[1559333788] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 466.09, 401.59, 220.07
[1559333873] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 363 processes
[1559333908] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 75.90, 279.95, 198.05
[1559333993] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 361 processes
[1559334028] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 11.58, 187.83, 174.22
[1559334124] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1559334124] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 417 processes
[1559334228] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559334242] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7279560 kB) free!
[1559334302] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.8% (3534952 kB) free.
[1559334348] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1559334362] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13178208 kB) free.
[1559334468] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1559334902] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (172932 kB) free!
[1559334962] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8956412 kB) free!
[1559335022] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (151408 kB) free!
[1559335069] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559335083] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7239596 kB) free!
[1559335143] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (170900 kB) free!
[1559335189] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559335203] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13169612 kB) free.
[1559335309] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1559335743] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.9% (2577356 kB) free!
[1559335803] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8762500 kB) free!
[1559335909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559335923] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.2% (6990688 kB) free!
[1559336029] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1559336043] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13179936 kB) free.
[1559336429] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.98, 1.93, 14.89
[1559336629] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1559336643] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9005624 kB) free!
[1559336749] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1559336763] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7380824 kB) free!
[1559336869] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1559336883] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.5% (5164072 kB) free!
[1559337124] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 360 processes
[1559337724] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 397 processes
[1559338083] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11995212 kB) free.
[1559338229] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.59, 1.72, 3.61
[1559338324] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1559338547] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 11.23, 56.59, 37.46
[1559338667] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.60, 37.91, 32.94
[1559338669] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1559338787] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.40, 25.42, 28.96
[1559338896] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.4% (2797260 kB) free!
[1559338924] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 341 processes
[1559339269] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559339283] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6821648 kB) free!
[1559339387] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.65, 3.83, 15.37
[1559339389] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1559339404] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13163184 kB) free.
[1559339989] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1559340003] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8547180 kB) free!
[1559340109] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1559340123] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.2% (6522184 kB) free!
[1559340230] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1559340243] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13175728 kB) free.
[1559340587] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.18, 0.38, 4.47
[1559340661] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 22.65, 180.58, 128.44
[1559340765] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.7% (2893204 kB) free!
[1559340781] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.33, 120.94, 112.90
[1559340829] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559340843] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8213060 kB) free!
[1559340901] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.73, 81.04, 99.25
[1559340949] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1559340963] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.3% (6042348 kB) free!
[1559341069] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559341083] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13167004 kB) free.
[1559341669] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1559341683] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7460960 kB) free!
[1559341789] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1559341803] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.6% (5219288 kB) free!
[1559341909] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1559341923] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13162504 kB) free.
[1559342343] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.1% (3421320 kB) free.
[1559342509] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559342523] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7287016 kB) free!
[1559342630] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559342643] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13168360 kB) free.
[1559342701] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.13, 0.43, 14.54
[1559342750] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1559342943] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (173828 kB) free!
[1559343064] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.2% (3444916 kB) free.
[1559343096] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2458932 kB) free!
[1559343124] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 394 processes
[1559343244] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8880212 kB) free!
[1559343350] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 363 processes
[1559343364] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.3% (7052300 kB) free!
[1559343470] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1559343484] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13161560 kB) free.
[1559343663] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (1269484 kB) free!
[1559343725] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 420 processes
[1559343783] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.4% (3473008 kB) free.
[1559343845] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 422 processes
[1559343901] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.32, 0.30, 4.19
[1559343965] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 405 processes
[1559344070] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559344084] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8690096 kB) free!
[1559344190] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559344204] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.7% (6729824 kB) free!
[1559344310] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559344324] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13160400 kB) free.
[1559344365] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (777956 kB) free!
[1559344501] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 14.13, 58.24, 39.11
[1559344565] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 378 processes
[1559344621] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.06, 39.04, 34.40
[1559344741] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1559344741] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.45, 26.17, 30.25
[1559344896] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.1% (2482400 kB) free!
[1559344910] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559344924] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8465068 kB) free!
[1559344965] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.7% (3376240 kB) free!
[1559345031] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559345044] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.1% (6437700 kB) free!
[1559345151] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 361 processes
[1559345164] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13157348 kB) free.
[1559345165] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 408 processes
[1559345285] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1559345342] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.26, 3.78, 15.99
[1559345496] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1559345496] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2428400 kB) free!
[1559345751] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1559345764] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8365592 kB) free!
[1559345871] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559345884] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.8% (6316940 kB) free!
[1559345991] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1559346004] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13149140 kB) free.
[1559346183] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.0% (2269804 kB) free!
[1559346303] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 26.0% (4218324 kB) free.
[1559346542] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.42, 0.38, 4.62
[1559346591] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559346605] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.2% (5518188 kB) free!
[1559346711] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559346725] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13150520 kB) free.
[1559347085] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 436 processes
[1559347205] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 435 processes
[1559347325] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 436 processes
[1559347911] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559347925] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 411 processes
[1559348031] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559348045] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 424 processes
[1559348151] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1559348646] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 415 processes
[1559348703] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.5% (3001080 kB) free!
[1559348766] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 426 processes
[1559348823] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1559348823] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.8% (4021712 kB) free.
[1559348886] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 416 processes
[1559349423] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (171808 kB) free!
[1559349543] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (963292 kB) free!
[1559349664] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (611108 kB) free!
[1559349720] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46869 MB (20% inode=99%):
[1559349840] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46740 MB (20% inode=99%):
[1559349961] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46708 MB (20% inode=99%):
[1559350686] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 417 processes
[1559350864] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.8% (3374700 kB) free.
[1559351286] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 402 processes
[1559351406] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 389 processes
[1559351429] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 128.95, 104.35, 54.19
[1559351526] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 386 processes
[1559351550] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 129.32, 112.60, 63.29
[1559351660] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4887008 kB) free!
[1559351670] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 19.28, 76.06, 55.88
[1559351720] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 19.9% (4881676 kB) free!
[1559351780] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 19.9% (4880932 kB) free!
[1559351814] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.6% (2848776 kB) free!
[1559351840] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 19.9% (4880848 kB) free!
[1559352870] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.98, 3.28, 16.78
[1559354527] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 419 processes
[1559354670] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.09, 2.07, 4.07
[1559354814] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2248560 kB) free!
[1559355727] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 397 processes
[1559355847] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 412 processes
[1559355967] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 420 processes
[1559356896] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559357343] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1559357525] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9754336 kB) free!
[1559357645] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8055756 kB) free!
[1559357670] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 477.86, 256.41, 104.92
[1559357765] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.4% (6114924 kB) free!
[1559357790] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 489.20, 334.00, 151.72
[1559357910] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1559357910] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 490.85, 385.90, 192.83
[1559358414] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.7% (2629512 kB) free!
[1559358665] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.5% (2352208 kB) free!
[1559358765] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2363460 kB) free!
[1559358785] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.7% (3195084 kB) free!
[1559358905] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.3% (3129384 kB) free!
[1559359014] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1559359014] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2238156 kB) free!
[1559360705] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1559360705] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (1380460 kB) free!
[1559360765] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4884176 kB) free!
[1559360767] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 436 processes
[1559360910] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.80, 2.12, 12.79
[1559361305] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.4% (2178764 kB) free!
[1559361367] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 434 processes
[1559361488] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 469 processes
[1559361513] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 18.7% (1377488 kB) free!
[1559361608] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 469 processes
[1559361633] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 16.4% (1205996 kB) free!
[1559361753] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 14.2% (1043200 kB) free!
[1559361860] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (1199608 kB) free!
[1559361905] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (626488 kB) free!
[1559361980] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (168496 kB) free!
[1559362100] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (182608 kB) free!
[1559362150] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.00, 5.19, 5.01
[1559362193] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 370 processes
[1559362270] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.50, 5.00, 4.97
[1559362313] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 1 warn out of 373 processes
[1559362353] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (134768 kB) free!
[1559362433] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 1 warn out of 373 processes
[1559362505] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.4% (2495072 kB) free!
[1559362710] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.08, 2.39, 3.91
[1559362808] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 428 processes
[1559363408] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 452 processes
[1559363470] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.60, 5.43, 5.08
[1559363528] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 455 processes
[1559363553] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.2% (6781176 kB) free.
[1559363590] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.90, 5.68, 5.22
[1559363648] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 454 processes
[1559363710] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.73, 5.73, 5.29
[1559363761] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 152163 MB (67% inode=99%):
[1559363814] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 58.9% (14446756 kB) free.
[1559364305] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (362612 kB) free!
[1559364365] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11828952 kB) free.
[1559364765] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 70.7% (17454704 kB) free.
[1559364831] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1559364905] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.7% (1902680 kB) free!
[1559364952] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1559364965] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (2742892 kB) free!
[1559365072] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1559365085] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10952888 kB) free.
[1559365672] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1559365820] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 610 processes
[1559365880] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 610 processes
[1559365896] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 72.0% (17658700 kB) free.
[1559366033] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 367 processes
[1559366105] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (764972 kB) free!
[1559366818] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1559366886] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 9 users currently logged in
[1559367305] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.3% (2154272 kB) free!
[1559367454] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.0% (2754192 kB) free!
[1559367848] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 448 processes
[1559367905] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174284 kB) free!
[1559368014] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4908448 kB) free!
[1559368054] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (162564 kB) free!
[1559368110] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 205.18, 115.02, 48.52
[1559368134] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175360 kB) free!
[1559368187] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 463.22, 267.01, 111.65
[1559368230] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 108.09, 120.42, 59.35
[1559368254] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (648596 kB) free!
[1559368307] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 469.11, 334.20, 155.11
[1559368350] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16.21, 81.15, 52.36
[1559368428] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 134.56, 274.86, 156.85
[1559368448] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 438 processes
[1559368505] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.9% (3230708 kB) free!
[1559369105] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (168504 kB) free!
[1559369550] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.82, 4.17, 16.17
[1559369705] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.3% (2974028 kB) free!
[1559369819] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 9 users currently logged in
[1559369840] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.1% (4934860 kB) free.
[1559370140] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4913728 kB) free!
[1559370153] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 14.5% (1066248 kB) free!
[1559370200] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 20.1% (4916452 kB) free.
[1559370273] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 12.1% (890920 kB) free!
[1559370305] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1562464 kB) free!
[1559370393] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1559370393] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (715832 kB) free!
[1559370828] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.40, 0.54, 12.33
[1559370994] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.2% (6853660 kB) free.
[1559371350] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.15, 2.96, 4.69
[1559371505] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.2% (3116532 kB) free!
[1559371654] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.9% (2746232 kB) free!
[1559371685] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (2211740 kB) free!
[1559371805] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10617208 kB) free.
[1559372028] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.28, 0.27, 3.61
[1559372105] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166228 kB) free!
[1559372705] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.2% (2783720 kB) free!
[1559373200] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4894076 kB) free!
[1559373260] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4892016 kB) free!
[1559373305] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.0% (3563628 kB) free.
[1559373320] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 19.9% (4881672 kB) free!
[1559373380] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 19.9% (4882768 kB) free!
[1559373454] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1559373454] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169440 kB) free!
[1559373472] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1559373592] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1559373712] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1559373905] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (1550588 kB) free!
[1559374026] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (164292 kB) free!
[1559374146] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (164700 kB) free!
[1559374912] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1559374912] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1559375346] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.8% (2071948 kB) free!
[1559376005] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.5% (6141752 kB) free!
[1559376054] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1559376125] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (4205096 kB) free!
[1559376245] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (2503540 kB) free!
[1559376454] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 20.0% (3251644 kB) free!
[1559376546] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (176332 kB) free!
[1559376712] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1559376765] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165104 kB) free!
[1559376832] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1559376885] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (2100744 kB) free!
[1559376953] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1559377005] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185136 kB) free!
[1559377054] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (153052 kB) free!
[1559377448] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 381 processes
[1559377594] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 11.9% (877280 kB) free!
[1559377714] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (720016 kB) free!
[1559377834] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (540472 kB) free!
[1559378048] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 398 processes
[1559378648] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 413 processes
[1559379248] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 387 processes
[1559379845] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.9% (5347420 kB) free!
[1559380147] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.1% (1794768 kB) free!
[1559380296] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (4599528 kB) free!
[1559380416] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.7% (4575636 kB) free!
[1559380448] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 419 processes
[1559380510] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.75, 3.98, 4.59
[1559380536] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.5% (4533224 kB) free!
[1559381048] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 379 processes
[1559381229] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45378 MB (20% inode=99%):
[1559381347] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169052 kB) free!
[1559381349] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45186 MB (20% inode=99%):
[1559381434] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.1% (6844036 kB) free.
[1559381469] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47598 MB (21% inode=99%):
[1559381550] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559381710] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.66, 6.31, 5.40
[1559381831] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.49, 6.43, 5.56
[1559381942] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 149.47, 107.84, 50.59
[1559381951] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.12, 5.64, 5.39
[1559382062] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 39.54, 86.57, 50.27
[1559382069] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46485 MB (20% inode=99%):
[1559382182] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 5.54, 58.00, 44.20
[1559382189] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46437 MB (20% inode=99%):
[1559382248] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 402 processes
[1559382309] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46086 MB (20% inode=99%):
[1559382551] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.79, 4.13, 4.68
[1559382848] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 400 processes
[1559382936] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2454824 kB) free!
[1559383147] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.4% (1848636 kB) free!
[1559383382] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.39, 1.51, 12.46
[1559383747] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (171156 kB) free!
[1559383751] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.31, 6.15, 5.30
[1559383871] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.57, 5.92, 5.32
[1559383982] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 455.51, 274.44, 122.88
[1559383991] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1559383991] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.37, 6.04, 5.46
[1559384048] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 405 processes
[1559384205] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.5% (2833772 kB) free!
[1559384648] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 382 processes
[1559384736] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (2532764 kB) free!
[1559384805] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2454020 kB) free!
[1559385336] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1559385336] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1955032 kB) free!
[1559385547] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.1% (2129228 kB) free!
[1559385791] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.78, 4.54, 4.74
[1559385848] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 407 processes
[1559386124] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1559386147] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (162852 kB) free!
[1559386244] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1559386364] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1559386382] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1559386382] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.34, 0.60, 15.56
[1559386391] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.48, 5.53, 5.20
[1559386448] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 378 processes
[1559386511] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.24, 5.60, 5.28
[1559386631] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.27, 5.54, 5.30
[1559386747] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.4% (2168108 kB) free!
[1559387153] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559387231] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.36, 4.48, 4.87
[1559387348] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169840 kB) free!
[1559387564] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1559387582] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.47, 0.46, 4.60
[1559387684] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1559387805] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1559388353] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 12 users currently logged in
[1559388405] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 12 users currently logged in
[1559388454] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (1954752 kB) free!
[1559388848] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 406 processes
[1559389054] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (175160 kB) free!
[1559389148] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.0% (2589740 kB) free!
[1559389428] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.07, 50.79, 49.20
[1559389445] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (429552 kB) free!
[1559389448] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 387 processes
[1559389548] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.59, 35.19, 43.67
[1559389553] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1559389668] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.25, 24.63, 38.78
[1559389748] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (154772 kB) free!
[1559389834] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 13.0% (955004 kB) free!
[1559389954] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 11.0% (806304 kB) free!
[1559390050] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.5% (4038372 kB) free!
[1559390074] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (707252 kB) free!
[1559390151] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1559390170] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (382660 kB) free!
[1559390290] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (1042484 kB) free!
[1559390304] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1559390424] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1559390544] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1559390582] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.44, 17.35, 15.84
[1559390702] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.46, 11.76, 13.97
[1559390823] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.45, 7.98, 12.32
[1559390868] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.32, 0.75, 10.91
[1559390948] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.8% (2562016 kB) free!
[1559391352] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1559391380] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 622 processes
[1559391441] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 622 processes
[1559391472] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559391548] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156740 kB) free!
[1559391592] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1559392023] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.74, 0.60, 3.67
[1559392054] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (2122440 kB) free!
[1559392068] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.32, 0.31, 3.21
[1559392148] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.3% (3300336 kB) free.
[1559392475] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.4% (6871516 kB) free.
[1559392655] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (481808 kB) free!
[1559392748] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (615532 kB) free!
[1559392868] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (161968 kB) free!
[1559392988] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (1368872 kB) free!
[1559394486] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1559394606] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1559394726] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1559394845] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.5% (5654984 kB) free!
[1559395193] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1559395326] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1559395421] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 63 users currently logged in
[1559395481] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 63 users currently logged in
[1559395541] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 57 users currently logged in
[1559395553] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1559396588] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.9% (3224864 kB) free!
[1559396736] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559396993] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559397113] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559397188] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156292 kB) free!
[1559397233] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1559397233] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559397353] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1559397353] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1559397490] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.9% (2675004 kB) free!
[1559397788] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.6% (2849092 kB) free!
[1559398055] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.3% (2964616 kB) free!
[1559398326] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1559398388] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (152752 kB) free!
[1559398446] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1559398566] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1559398631] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559398988] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.2% (2470144 kB) free!
[1559399153] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1559399255] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165848 kB) free!
[1559399588] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (162748 kB) free!
[1559399741] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1559399801] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1559399861] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1559399921] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1559400188] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.2% (2473736 kB) free!
[1559400789] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (167164 kB) free!
[1559400848] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559401421] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 60 users currently logged in
[1559402075] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1559402824] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1559403189] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.6% (3834104 kB) free.
[1559403367] SERVICE FLAPPING ALERT: db2;Current Users;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1559403367] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1559403789] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (158716 kB) free!
[1559403909] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.5% (2357024 kB) free!
[1559404029] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (168652 kB) free!
[1559404433] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559404445] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9905996 kB) free.
[1559404629] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.0% (2266156 kB) free!
[1559405229] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (160984 kB) free!
[1559406233] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559406245] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9823984 kB) free!
[1559406353] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1559406365] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7620460 kB) free!
[1559406473] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559406486] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.0% (5916656 kB) free!
[1559406848] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 385 processes
[1559407005] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.3% (2780996 kB) free!
[1559407023] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 27.20, 208.10, 148.21
[1559407029] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.1% (1795868 kB) free!
[1559407143] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.92, 139.35, 130.28
[1559407263] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.67, 93.29, 114.50
[1559407448] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 389 processes
[1559408753] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1559408868] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 113.39, 25.23, 8.44
[1559408988] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 456.69, 185.73, 69.30
[1559409063] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.38, 0.56, 16.77
[1559409108] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 502.94, 293.09, 122.75
[1559409429] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.2% (3280224 kB) free.
[1559409848] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 382 processes
[1559409937] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.6% (2605820 kB) free!
[1559410005] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2423172 kB) free!
[1559410029] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.5% (2020656 kB) free!
[1559410055] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1559410149] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (174708 kB) free!
[1559410263] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.36, 0.34, 4.87
[1559410269] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.0% (1784684 kB) free!
[1559410449] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1559410449] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 384 processes
[1559410567] SERVICE FLAPPING ALERT: db2;Current Users;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1559410869] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1413600 kB) free!
[1559411737] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2403224 kB) free!
[1559411890] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 20.4% (5008688 kB) free.
[1559412108] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.04, 0.25, 11.04
[1559412150] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 142.49, 96.20, 44.00
[1559412270] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 143.46, 111.82, 56.05
[1559412390] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 111.47, 115.13, 64.29
[1559413269] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.2% (2632844 kB) free!
[1559413308] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.28, 0.26, 3.21
[1559413590] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.79, 3.80, 18.93