[1577558326] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (1140428 kB) free!
[1577558446] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (1140932 kB) free!
[1577558488] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198220 kB) free!
[1577558548] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199440 kB) free!
[1577558566] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (1141224 kB) free!
[1577558571] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.62, 5.68, 5.17
[1577558610] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189356 kB) free!
[1577558668] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186076 kB) free!
[1577558691] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.24, 5.52, 5.17
[1577558725] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187488 kB) free!
[1577558730] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191552 kB) free!
[1577558788] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169896 kB) free!
[1577558811] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.23, 5.43, 5.18
[1577558846] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186692 kB) free!
[1577558966] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174364 kB) free!
[1577559411] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.76, 4.15, 4.63
[1577561366] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 44.8% (11049304 kB) free.
[1577561966] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.6% (3367868 kB) free!
[1577562086] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (1951324 kB) free!
[1577562206] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (595608 kB) free!
[1577563095] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 12.7% (6265252 kB) free!
[1577563215] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 10.9% (5349640 kB) free!
[1577563336] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4417312 kB) free!
[1577563611] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.61, 5.65, 5.12
[1577563731] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.31, 5.58, 5.17
[1577563851] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.54, 5.25, 5.10
[1577564451] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.62, 4.28, 4.69
[1577565767] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 43.5% (10653748 kB) free.
[1577566317] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.5% (8610832 kB) free!
[1577566437] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (1502628 kB) free!
[1577566558] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (215540 kB) free!
[1577567158] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 82.3% (40560884 kB) free.
[1577568500] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577568651] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.27, 5.31, 5.04
[1577568771] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.89, 5.20, 5.03
[1577568891] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.33, 5.02, 5.00
[1577569100] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(17) < 2019/12/29 03:06:40 [error] 20756#0: *20813563 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577569220] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1605) < 2019/12/29 03:10:16 [error] 20759#0: *20816456 no live upstreams while connecting to upstream, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "backend_3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577569239] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/12/29 03:10:24 [error] 28871#0: *28848499 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1577569245] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1468) < 49.50.64.122 [29/Dec/2019:03:10:16 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1577569341] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(531) < 2019/12/29 03:12:12 [error] 20757#0: *20816289 upstream timed out (110: Connection timed out) while connecting to upstream, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577569360] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577569365] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(280) < 172.105.41.64 [29/Dec/2019:03:11:22 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1577569485] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577569603] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1577569723] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1577573691] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.33, 5.57, 5.23
[1577573811] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.43, 5.46, 5.23
[1577573931] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.08, 4.93, 5.06
[1577574531] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.24, 4.06, 4.52
[1577574741] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577575341] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(20) < 2019/12/29 04:47:57 [error] 20748#0: *20836641 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577575461] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577576061] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(15) < 2019/12/29 05:00:45 [error] 20761#0: *20839147 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577576181] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2019/12/29 05:06:19 [error] 20756#0: *20840231 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577576301] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1577576301] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(3) < 2019/12/29 05:06:39 [error] 20758#0: *20840300 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577578731] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.67, 5.30, 5.15
[1577578851] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.31, 5.04, 5.08
[1577578972] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.99, 4.64, 4.93
[1577579567] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187448 kB) free!
[1577579687] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191208 kB) free!
[1577579807] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 28.5% (6984248 kB) free.
[1577579934] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 405 processes
[1577580054] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1577580331] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 88.2% (21616592 kB) free.
[1577580388] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.2% (21625468 kB) free.
[1577580408] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (230064 kB) free!
[1577580528] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (230708 kB) free!
[1577580648] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (234972 kB) free!
[1577580806] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 46.6% (11503952 kB) free.
[1577582901] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577583501] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2019/12/29 07:07:57 [error] 20753#0: *20864146 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577583621] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577583632] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1577583752] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1577583771] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.45, 5.47, 5.18
[1577583891] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.77, 5.22, 5.13
[1577584011] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.34, 4.89, 5.02
[1577584221] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2019/12/29 07:15:40 [error] 20752#0: *20865649 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577584341] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(6) < 2019/12/29 07:21:40 [error] 20760#0: *20866828 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577584462] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577584611] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.35, 3.91, 4.48
[1577585061] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(15) < 2019/12/29 07:32:57 [error] 20756#0: *20869061 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577585181] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577585781] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(11) < 2019/12/29 07:45:42 [error] 20757#0: *20871555 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577585901] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577586501] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/12/29 07:51:40 [error] 20752#0: *20872718 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577586621] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577587221] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(15) < 2019/12/29 08:02:57 [error] 20755#0: *20874951 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577587342] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577587941] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(21) < 2019/12/29 08:21:40 [error] 20754#0: *20878619 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577588061] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577588211] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.43, 5.96, 5.11
[1577588331] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.33, 6.12, 5.27
[1577588451] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.99, 6.02, 5.34
[1577588454] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1577588574] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1577588661] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(13) < 2019/12/29 08:31:19 [error] 20755#0: *20880521 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577588781] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577589051] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.84, 4.61, 5.00
[1577589381] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(15) < 2019/12/29 08:46:19 [error] 20753#0: *20883475 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577589501] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577590102] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(6) < 2019/12/29 08:56:40 [error] 20747#0: *20885486 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577590222] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1577590222] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577590822] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(18) < 2019/12/29 09:06:40 [error] 20753#0: *20887457 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577590942] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577591542] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(18) < 2019/12/29 09:21:40 [error] 20750#0: *20890439 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577591662] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2019/12/29 09:22:57 [error] 20755#0: *20890683 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577591782] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577592383] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(15) < 2019/12/29 09:31:40 [error] 20757#0: *20892170 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577592503] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2019/12/29 09:37:57 [error] 20748#0: *20893171 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577592623] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577592806] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194440 kB) free!
[1577592926] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192884 kB) free!
[1577593046] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167344 kB) free!
[1577593223] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2019/12/29 09:45:42 [error] 20759#0: *20894431 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577593252] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.23, 6.28, 5.27
[1577593343] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2019/12/29 09:51:19 [error] 20758#0: *20895328 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577593372] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.93, 6.75, 5.56
[1577593463] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577593492] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.27, 6.87, 5.75
[1577594063] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2019/12/29 10:02:57 [error] 20760#0: *20897208 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577594183] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577594239] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.9% (22190308 kB) free.
[1577594275] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.4% (38647532 kB) free.
[1577594306] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.3% (22372660 kB) free.
[1577594784] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2019/12/29 10:15:45 [error] 20751#0: *20899261 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577594904] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2019/12/29 10:16:40 [error] 20758#0: *20899411 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577595024] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577595624] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2019/12/29 10:27:57 [error] 20760#0: *20901206 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577595744] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(12) < 2019/12/29 10:30:43 [error] 20760#0: *20901674 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577595864] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577595939] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 494.98, 331.34, 143.88
[1577596059] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 72.64, 224.47, 127.49
[1577596179] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 18.97, 153.54, 113.25
[1577596464] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/12/29 10:42:57 [error] 20761#0: *20903606 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577596585] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(12) < 2019/12/29 10:45:43 [error] 20758#0: *20904080 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577596705] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(5) < 2019/12/29 10:47:57 [error] 20756#0: *20904432 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577597305] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577597845] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167388 kB) free!
[1577597905] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(18) < 2019/12/29 11:06:19 [error] 20759#0: *20907370 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577597965] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182380 kB) free!
[1577598025] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577598085] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183532 kB) free!
[1577598371] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1577598371] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.6% (22107440 kB) free.
[1577598625] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(9) < 2019/12/29 11:15:44 [error] 20757#0: *20908885 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577598745] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2019/12/29 11:21:41 [error] 20757#0: *20909836 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577598866] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577599179] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 9.09, 12.23, 29.48
[1577599465] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(15) < 2019/12/29 11:31:19 [error] 20759#0: *20911388 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577599585] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577599779] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 594.33, 288.39, 127.85
[1577599899] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 441.46, 362.93, 176.56
[1577600019] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 72.72, 248.10, 157.09
[1577600185] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2019/12/29 11:46:19 [error] 20759#0: *20913795 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577600305] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(3) < 2019/12/29 11:47:57 [error] 20758#0: *20914050 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577600425] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577601025] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(6) < 2019/12/29 11:51:41 [error] 20761#0: *20914649 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577601032] SERVICE ALERT: qservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 552 processes
[1577601146] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(12) < 2019/12/29 12:01:40 [error] 20756#0: *20916258 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577601152] SERVICE ALERT: qservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 549 processes
[1577601266] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577601272] SERVICE ALERT: qservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 549 processes
[1577601292] SERVICE FLAPPING ALERT: qservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1577601331] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (4510112 kB) free!
[1577601379] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170524 kB) free!
[1577601451] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4380040 kB) free!
[1577601475] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (250584 kB) free!
[1577601499] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185328 kB) free!
[1577601518] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172752 kB) free!
[1577601571] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (4300300 kB) free!
[1577601595] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 56.8% (28003356 kB) free.
[1577601619] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 39.6% (9768264 kB) free.
[1577601639] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192068 kB) free!
[1577601759] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188012 kB) free!
[1577601866] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(6) < 2019/12/29 12:07:57 [error] 20757#0: *20917260 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577601986] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(9) < 2019/12/29 12:15:44 [error] 20757#0: *20918519 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577602106] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577602706] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/12/29 12:22:57 [error] 20750#0: *20919673 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577602819] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (4489692 kB) free!
[1577602826] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577602939] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.2% (3261708 kB) free!
[1577603059] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.0% (3207568 kB) free!
[1577603427] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(15) < 2019/12/29 12:37:57 [error] 20758#0: *20922082 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577603547] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/12/29 12:41:19 [error] 20749#0: *20922618 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577603667] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(3) < 2019/12/29 12:42:57 [error] 20759#0: *20922874 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577603971] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2006800 kB) free!
[1577604859] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2119932 kB) free!
[1577608072] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.0% (22449244 kB) free.
[1577608096] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.0% (22313980 kB) free.
[1577608156] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.6% (4563920 kB) free!
[1577608168] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.6% (22935684 kB) free.
[1577608198] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.9% (22022972 kB) free.
[1577609619] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 8.56, 9.89, 32.02
[1577610219] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1577610219] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 12.64, 12.24, 22.34
[1577615276] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.1% (3222944 kB) free!
[1577615396] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.9% (3187328 kB) free!
[1577615516] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.2% (3010716 kB) free!
[1577615945] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.65, 4.65, 4.93
[1577616116] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (1347000 kB) free!
[1577617745] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.32, 5.51, 5.10
[1577617865] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.66, 5.93, 5.30
[1577617985] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.84, 5.96, 5.40
[1577619160] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.2% (20789984 kB) free.
[1577620385] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.49, 4.46, 4.89
[1577621619] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1577623385] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.36, 6.58, 5.70
[1577623505] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.21, 6.47, 5.77
[1577623625] SERVICE FLAPPING ALERT: qservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1577623625] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.21, 6.80, 5.98
[1577624395] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9669568 kB) free!
[1577624515] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (9299220 kB) free!
[1577624635] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.9% (9289664 kB) free!
[1577625425] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 5.49, 4.90, 4.96
[1577625508] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4666588 kB) free!
[1577625628] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (4587284 kB) free!
[1577625748] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (4539820 kB) free!
[1577626367] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1577627758] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9848652 kB) free!
[1577627878] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.9% (9785160 kB) free!
[1577627998] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.8% (9739540 kB) free!
[1577628152] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1577628272] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1577628425] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.35, 5.62, 5.23
[1577628545] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.28, 5.47, 5.22
[1577628665] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.50, 5.55, 5.29
[1577629265] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.10, 4.49, 4.91
[1577629348] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2426560 kB) free!
[1577630668] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577630788] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4909956 kB) free!
[1577630908] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4892276 kB) free!
[1577631028] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.9% (4867452 kB) free!
[1577631268] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(15) < 2019/12/29 20:15:43 [error] 20756#0: *20994056 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577631388] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577631579] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4875840 kB) free!
[1577631699] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (4796072 kB) free!
[1577631766] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.2% (3738968 kB) free!
[1577631819] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.3% (4729784 kB) free!
[1577631886] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.0% (3711900 kB) free!
[1577631989] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2019/12/29 20:36:19 [error] 20756#0: *20997341 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577632006] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.0% (3696020 kB) free!
[1577632109] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577632709] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(21) < 2019/12/29 20:46:40 [error] 20757#0: *20999005 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577632829] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1577632974] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1577633095] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1577633429] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(6) < 2019/12/29 20:57:57 [error] 20759#0: *21000778 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577633465] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.58, 5.83, 5.19
[1577633549] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(9) < 2019/12/29 21:00:38 [error] 20761#0: *21001210 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577633585] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.43, 5.69, 5.22
[1577633669] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(3) < 2019/12/29 21:02:57 [error] 20751#0: *21001606 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1577633705] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.18, 5.60, 5.25
[1577633811] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 75.6% (18525284 kB) free.
[1577634305] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.41, 4.59, 4.90
[1577635606] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2153364 kB) free!
[1577636203] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.2% (21754744 kB) free.
[1577638505] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.22, 5.50, 5.03
[1577638625] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.00, 5.35, 5.03
[1577638745] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.61, 5.08, 4.98
[1577640294] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1577640414] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1577642154] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (249292 kB) free!
[1577642274] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (245836 kB) free!
[1577642394] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (246344 kB) free!
[1577642619] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2392508 kB) free!