[1548009092] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548009202] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1548009212] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548009322] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548009442] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548010412] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548010511] SERVICE FLAPPING ALERT: webservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1548010532] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548011732] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548011852] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548012072] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.54, 5.63, 5.29
[1548012189] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (3136528 kB) free!
[1548012192] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.18, 5.71, 5.36
[1548012309] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (3002820 kB) free!
[1548012312] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.58, 5.94, 5.48
[1548012429] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (2885788 kB) free!
[1548013629] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (1453296 kB) free!
[1548015444] SERVICE FLAPPING ALERT: webservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1548015924] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2428428 kB) free!
[1548016652] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548016772] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1548017972] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548018092] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1548018460] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49049 MB (21% inode=99%):
[1548018692] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548018812] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1548020012] SERVICE FLAPPING ALERT: db1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1548020242] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1548020612] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548020732] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548021332] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548021452] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548021924] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1548022043] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548022163] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548022283] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1548022652] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548022772] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1548022883] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1548023483] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1548023603] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1548023723] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1548023972] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548024092] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1548025292] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548025412] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1548026123] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1548026612] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548026732] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548027932] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548028052] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548028161] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.7% (5285984 kB) free!
[1548029252] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548029361] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4849316 kB) free!
[1548029372] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548029961] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.1% (5488996 kB) free!
[1548031172] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548031293] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548031523] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548031643] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1548032492] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548032612] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1548033812] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548033932] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548035140] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 360 processes
[1548035243] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548035260] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1548035363] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1548035963] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1548036332] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548036452] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548038244] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1548038361] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (2645164 kB) free!
[1548038363] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1548038484] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1548039411] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.1% (4962060 kB) free!
[1548040161] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1548040161] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.2% (5509916 kB) free!
[1548050361] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1548051683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548051803] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1548051924] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1548052011] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4696460 kB) free!
[1548052524] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1548052611] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.4% (5140980 kB) free!
[1548053124] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548053244] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1548053252] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548053372] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548054444] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548054564] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1548054711] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1548054712] SERVICE ALERT: sig2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1548054712] SERVICE ALERT: sig2;HDD_Root status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Root status' on host 'sig2' was out of bounds)
[1548054727] SERVICE ALERT: sig2;CPU Load;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU Load' on host 'sig2' was out of bounds)
[1548054730] SERVICE ALERT: sig2;HDD_Home status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Home status' on host 'sig2' was out of bounds)
[1548054731] HOST ALERT: sig2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1548054795] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1548054823] SERVICE ALERT: sig2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 312 processes
[1548054832] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1548054832] SERVICE ALERT: sig2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38725 MB (75% inode=98%):
[1548054832] SERVICE ALERT: sig2;Total Processes;OK;HARD;1;PROCS OK: 313 processes
[1548054847] SERVICE ALERT: sig2;CPU Load;OK;HARD;1;OK - load average: 0.34, 0.17, 0.06
[1548054850] SERVICE ALERT: sig2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 210013 MB (93% inode=99%):
[1548054943] SERVICE ALERT: sig2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1548055164] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548055284] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1548055612] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4567448 kB) free!
[1548055761] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3695604 kB) free!
[1548056212] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1548056212] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.4% (5140464 kB) free!
[1548056732] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 17 users currently logged in
[1548056852] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 17 users currently logged in
[1548056972] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 17 users currently logged in
[1548058284] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548058404] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1548058512] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 33.08, 26.83, 22.14
[1548059961] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.2% (5037704 kB) free!
[1548059973] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548060093] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548061161] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4916200 kB) free!
[1548061761] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1548061761] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.3% (5078428 kB) free!
[1548061906] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 41 users currently logged in
[1548061966] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 41 users currently logged in
[1548062026] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 41 users currently logged in
[1548062086] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 41 users currently logged in
[1548062111] SERVICE ALERT: webservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1548062112] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 17.18, 16.57, 18.95
[1548062813] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4639076 kB) free!
[1548062972] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 22 users currently logged in
[1548063311] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1548063312] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 18.98, 20.23, 20.21
[1548063413] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.4% (5137096 kB) free!
[1548063911] SERVICE ALERT: webservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1548063912] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1548063912] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 20.30, 19.54, 19.82
[1548064761] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4520688 kB) free!
[1548065112] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 26.36, 23.55, 21.47
[1548065493] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548065613] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548065961] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.4% (5144388 kB) free!
[1548067404] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548067524] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1548068724] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548068844] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1548069214] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548069334] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1548069445] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548069561] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4445988 kB) free!
[1548069565] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548069685] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1548069913] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.40, 3.67, 12.60
[1548070013] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4922476 kB) free!
[1548070511] SERVICE FLAPPING ALERT: webservercluster2;Current Users;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1548070511] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1548070534] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548070654] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548070774] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548071064] SERVICE ALERT: sig4;CPU_Procs;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU_Procs' on host 'sig4' was out of bounds)
[1548071113] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.66, 2.03, 4.99
[1548071143] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1548071184] SERVICE ALERT: sig4;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1548071273] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1548071361] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.1% (5965984 kB) free!
[1548071393] SERVICE ALERT: sig2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 310 processes
[1548071993] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 319 processes
[1548071998] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 522 processes
[1548072058] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 531 processes
[1548072085] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1548072205] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548072319] SERVICE ALERT: sig5;CPU Load;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU Load' on host 'sig5' was out of bounds)
[1548072325] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548072337] HOST ALERT: sig5;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1548072401] HOST ALERT: sig5;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1548072439] SERVICE ALERT: sig5;CPU Load;OK;HARD;1;OK - load average: 0.61, 0.25, 0.09
[1548072575] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548072695] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548072815] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1548073772] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 18 users currently logged in
[1548074038] SERVICE ALERT: sig2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1548074040] SERVICE ALERT: sig2;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'sig2' was out of bounds)
[1548074040] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig2' was out of bounds)
[1548074040] HOST ALERT: sig2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.150)
[1548074040] SERVICE ALERT: sig2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig2' was out of bounds)
[1548074040] SERVICE ALERT: sig2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig2' was out of bounds)
[1548074040] SERVICE ALERT: sig2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig2' was out of bounds)
[1548074050] SERVICE ALERT: sig2;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sig2' was out of bounds)
[1548074050] SERVICE ALERT: sig2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig2' was out of bounds)
[1548074158] SERVICE ALERT: sig2;Sensors Status;OK;HARD;1;SENSORS OK
[1548074160] SERVICE ALERT: sig2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38559 MB (75% inode=98%):
[1548074160] SERVICE ALERT: sig2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1548074160] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1548074160] SERVICE ALERT: sig2;Total Processes;OK;HARD;1;PROCS OK: 331 processes
[1548074160] SERVICE ALERT: sig2;Memory;OK;HARD;1;OK - 93.8% (15238112 kB) free.
[1548074162] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1548074170] SERVICE ALERT: sig2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 208563 MB (92% inode=99%):
[1548074170] SERVICE ALERT: sig2;CPU Load;OK;HARD;1;OK - load average: 0.34, 0.19, 0.07
[1548074615] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548074708] SERVICE ALERT: sig1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1548074735] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1548074738] HOST ALERT: sig1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1548074802] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 1026.78 ms
[1548074828] SERVICE ALERT: sig1;Total Processes;OK;HARD;1;PROCS OK: 320 processes
[1548074874] SERVICE ALERT: sig1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1548074986] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 35 users currently logged in
[1548074994] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548075460] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47070 MB (20% inode=99%):
[1548075470] SERVICE ALERT: db1;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 10745 MB (20% inode=99%):
[1548075580] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46382 MB (20% inode=99%):
[1548075590] SERVICE ALERT: db1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 10986 MB (21% inode=99%):
[1548075700] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46189 MB (20% inode=99%):
[1548075935] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548076013] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1548076056] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548076176] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548076190] SERVICE ALERT: db1;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 10705 MB (20% inode=99%):
[1548076310] SERVICE ALERT: db1;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 10416 MB (20% inode=99%):
[1548076430] SERVICE ALERT: db1;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 10665 MB (20% inode=99%):
[1548076511] SERVICE FLAPPING ALERT: webservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1548076772] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 21 users currently logged in
[1548077372] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 20 users currently logged in
[1548077376] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1548077630] SERVICE ALERT: db1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 10780 MB (21% inode=99%):
[1548077725] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1548078311] SERVICE ALERT: webservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1548078326] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548078446] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1548078576] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548078696] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548079046] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1548079161] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4843984 kB) free!
[1548079166] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1548079286] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1548079296] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548079417] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548079537] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548080031] SERVICE ALERT: db1;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 10724 MB (20% inode=99%):
[1548080111] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1548080151] SERVICE ALERT: db1;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 10659 MB (20% inode=99%):
[1548080213] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.6% (7216548 kB) free!
[1548080271] SERVICE ALERT: db1;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 10659 MB (20% inode=99%):
[1548080291] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 379 processes
[1548080361] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.1% (4975076 kB) free!
[1548080486] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548080606] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548080612] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 6.78, 6.98, 9.02
[1548080727] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548080737] SERVICE FLAPPING ALERT: db1;CPU_Procs;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1548080737] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1548080961] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4842924 kB) free!
[1548081311] SERVICE FLAPPING ALERT: webservercluster2;Current Users;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1548081311] SERVICE ALERT: webservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1548081561] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6168392 kB) free!
[1548081913] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1548082761] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (3277616 kB) free!
[1548084327] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1548084327] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1548084572] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 14 users currently logged in
[1548085161] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6173852 kB) free!
[1548085511] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1548089711] SERVICE ALERT: webservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1548091537] SERVICE FLAPPING ALERT: db1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1548092727] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1548094560] SERVICE ALERT: sig2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3119740 kB) free!
[1548094680] SERVICE ALERT: sig2;Memory;WARNING;SOFT;2;WARNING - 18.8% (3050016 kB) free!
[1548094761] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1548094800] SERVICE ALERT: sig2;Memory;WARNING;HARD;3;WARNING - 18.3% (2981352 kB) free!
[1548095111] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in