[1595442762] SERVICE ALERT: sservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=3.71% system=7.92% iowait=0.00% idle=88.37%
[1595442811] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1595442811] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.00, 2.32, 2.48
[1595442832] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1595442888] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1595442901] SERVICE ALERT: sservercluster2;Total Processes;OK;HARD;1;PROCS OK: 347 processes
[1595442969] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1595442979] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 36.2% (8876044 kB) free.
[1595443075] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181436 kB) free!
[1595443083] SERVICE FLAPPING ALERT: sservercluster2;HDD_Root status;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1595443083] SERVICE ALERT: sservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 31392 MB (61% inode=98%):
[1595443179] SERVICE ALERT: sservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 163911 MB (74% inode=99%):
[1595443196] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165600 kB) free!
[1595443245] SERVICE ALERT: sservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1595443318] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187772 kB) free!
[1595443409] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197160 kB) free!
[1595443529] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175212 kB) free!
[1595443579] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190248 kB) free!
[1595443649] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1595443649] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188932 kB) free!
[1595443699] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167692 kB) free!
[1595443819] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193816 kB) free!
[1595445579] SERVICE ALERT: sservercluster2;HDD_Home status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595445655] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595445709] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595445772] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595445785] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595445821] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595445839] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595445842] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595445898] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595445902] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595445911] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595445915] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595445951] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595445972] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595445979] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595446028] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595446032] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595446041] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595446081] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595446093] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595446102] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595446109] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595446158] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1595446158] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595446171] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595446230] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595446239] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595446360] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595446552] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190168 kB) free!
[1595446612] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 71.5% (17535096 kB) free.
[1595447395] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 70.77, 46.22, 38.53
[1595447910] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (409632 kB) free!
[1595447995] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 40.82, 46.63, 43.33
[1595448007] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 222892 MB (20% inode=99%):
[1595448030] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (382196 kB) free!
[1595448098] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (238000 kB) free!
[1595448127] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 215534 MB (20% inode=99%):
[1595448150] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400200 kB) free!
[1595448219] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (243584 kB) free!
[1595448247] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 207983 MB (19% inode=99%):
[1595448339] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (242404 kB) free!
[1595449649] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 44.5% (10906332 kB) free.
[1595450139] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 39.4% (19396964 kB) free.
[1595450849] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.2% (3238236 kB) free!
[1595450969] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182708 kB) free!
[1595451089] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183480 kB) free!
[1595451320] SERVICE ALERT: sservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1595451437] SERVICE ALERT: sservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=3.54% system=7.42% iowait=0.00% idle=89.04%
[1595451481] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.52, 1.72, 1.85
[1595451502] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1595451559] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1595451571] SERVICE ALERT: sservercluster2;Total Processes;OK;HARD;1;PROCS OK: 340 processes
[1595451639] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1595451760] SERVICE ALERT: sservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 31393 MB (61% inode=98%):
[1595451839] SERVICE ALERT: sservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 164580 MB (74% inode=99%):
[1595451921] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 66.2% (32607904 kB) free.
[1595452159] SERVICE ALERT: sservercluster2;CPU_Procs;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595452171] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1595452220] SERVICE ALERT: sservercluster2;Memory;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595452249] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452289] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452301] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452367] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 65.5% (32283488 kB) free.
[1595452370] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452379] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452409] SERVICE ALERT: sservercluster2;CPU_Procs;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595452431] SERVICE FLAPPING ALERT: sservercluster2;Total Processes;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1595452431] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452449] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452500] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452509] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1595452509] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452530] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452569] SERVICE ALERT: sservercluster2;HDD_Home status;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595452630] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452647] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452660] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452691] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452699] SERVICE FLAPPING ALERT: sservercluster2;HDD_Home status;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1595452699] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452712] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452777] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452790] SERVICE FLAPPING ALERT: sservercluster2;Sensors Status;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1595452790] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452821] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595452830] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452842] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452907] SERVICE FLAPPING ALERT: sservercluster2;CPU-STATISTICS;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1595452907] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452951] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452964] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595452972] SERVICE FLAPPING ALERT: sservercluster2;LocalService-NTPD SERVICE;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1595452972] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595453019] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595453026] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595453040] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595453089] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595453096] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595453146] SERVICE ALERT: pservercluster2;Total Processes;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595453160] SERVICE ALERT: pservercluster2;Sensors Status;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595453188] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595453210] SERVICE ALERT: pservercluster2;CPU_Procs;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595453210] SERVICE ALERT: pservercluster2;CPU Load;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595453221] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595453228] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595453276] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595453280] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595453289] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595453318] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595453328] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595453340] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595453348] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595453353] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595453358] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595453419] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595453449] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595453458] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595453470] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595453479] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595453488] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595453512] SERVICE ALERT: sservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=3.67% system=8.28% iowait=0.00% idle=88.04%
[1595453539] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1595453575] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2732028 kB) free!
[1595453583] SERVICE ALERT: pservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=5.60% system=11.11% iowait=0.08% idle=83.20%
[1595453599] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 11008 open files (0% of max 2424156)
[1595453619] SERVICE ALERT: sservercluster2;CPU_Procs;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595453810] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1595453876] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 346 processes
[1595453880] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1595453948] SERVICE ALERT: pservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38300 MB (74% inode=98%):
[1595453949] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(22) < 172.105.59.173 [23/Jul/2020:03:03:48 +0530] TCP 502 0 0 120.000 "10.147.212.154:13001, 10.147.212.143:13001, backend_13001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1595454049] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1595454069] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1595454070] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.04, 2.22, 2.13
[1595454088] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 157038 MB (69% inode=99%):
[1595454122] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595454229] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595454252] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595454382] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595455289] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 60.2% (14754824 kB) free.
[1595455375] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2443220 kB) free!
[1595456395] SERVICE ALERT: mrq4;CPU Load;OK;HARD;1;OK - load average: 29.57, 26.89, 29.61
[1595456995] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;1;WARNING - load average: 24.74, 30.82, 31.52
[1595457115] SERVICE ALERT: mrq4;CPU Load;WARNING;SOFT;2;WARNING - load average: 33.40, 34.36, 32.85
[1595457235] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 47.32, 37.63, 34.12
[1595457689] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194808 kB) free!
[1595457810] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173140 kB) free!
[1595457930] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196592 kB) free!
[1595458530] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 59.3% (14538452 kB) free.
[1595459470] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 11908.57, 3798.54, 1335.94
[1595459590] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15583.62, 7887.64, 3132.99
[1595459635] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 55.47, 42.26, 37.10
[1595459710] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16078.94, 10623.48, 4712.05
[1595460235] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 18.21, 27.47, 32.61
[1595460330] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173740 kB) free!
[1595460450] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (202944 kB) free!
[1595460570] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202112 kB) free!
[1595461170] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 21.2% (5198848 kB) free.
[1595461375] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 59.8% (14753676 kB) free.
[1595462710] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1595462710] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.78, 6.10, 677.25
[1595462970] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (2351676 kB) free!
[1595463090] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (2359084 kB) free!
[1595463210] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2357372 kB) free!
[1595463310] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.51, 2.84, 355.99
[1595463600] SERVICE FLAPPING ALERT: sservercluster2;Sensors Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1595463761] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1595463782] SERVICE FLAPPING ALERT: sservercluster2;LocalService-NTPD SERVICE;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1595463841] SERVICE FLAPPING ALERT: sservercluster2;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595463919] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595464040] SERVICE FLAPPING ALERT: sservercluster2;HDD_Root status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595464109] SERVICE FLAPPING ALERT: sservercluster2;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595464240] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595464918] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 85.8% (21018932 kB) free.
[1595465639] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595465792] SERVICE FLAPPING ALERT: sservercluster2;CPU-STATISTICS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595466750] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 41.4% (54558764 kB) free.
[1595466847] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 241373 MB (22% inode=99%):
[1595467727] SERVICE ALERT: sig4-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 87104 MB (67% inode=99%):
[1595468035] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 66.47, 61.52, 51.64
[1595469775] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1595469829] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1595469831] SERVICE ALERT: sservercluster2;Total Processes;OK;HARD;1;PROCS OK: 335 processes
[1595469909] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1595469987] SERVICE ALERT: sservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=3.67% system=7.79% iowait=0.00% idle=88.53%
[1595470030] SERVICE ALERT: sservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 31392 MB (61% inode=98%):
[1595470099] SERVICE ALERT: sservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 164418 MB (74% inode=99%):
[1595470190] SERVICE ALERT: sservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1595470351] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.00, 46.50, 55.07
[1595472751] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15148.21, 12747.81, 6777.26
[1595472830] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 776 processes
[1595472831] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 776 processes
[1595472871] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15150.61, 13542.79, 7792.50
[1595472950] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 775 processes
[1595472952] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 775 processes
[1595472992] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 9982.46, 12858.70, 8282.31
[1595473072] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 757 processes
[1595473410] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.8% (22996120 kB) free.
[1595473672] SERVICE ALERT: sservercluster2;Total Processes;OK;HARD;1;PROCS OK: 337 processes
[1595474035] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1595474710] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595475175] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (2280220 kB) free!
[1595475295] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (1517116 kB) free!
[1595475392] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.39, 6.12, 626.12
[1595475415] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (789156 kB) free!
[1595475950] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1595475992] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1595475992] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.72, 3.09, 329.27
[1595476070] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1595476190] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1595476230] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 71.1% (17415336 kB) free.
[1595476912] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (2196596 kB) free!
[1595476972] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168084 kB) free!
[1595477032] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (190616 kB) free!
[1595477094] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (166140 kB) free!
[1595477217] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.7% (20913080 kB) free.
[1595477394] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 47.8% (11704640 kB) free.
[1595479410] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1595479509] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1595479629] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1595479749] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1595480280] SERVICE ALERT: pservercluster2;Sensors Status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595480281] SERVICE ALERT: pservercluster2;Total Processes;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595480349] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1595480353] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480401] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480405] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480454] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480478] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480493] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480521] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480523] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480530] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480545] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480579] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480584] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480600] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480603] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480618] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480648] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480670] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480704] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480709] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480715] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480725] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480743] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480773] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480795] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480815] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480834] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480840] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480850] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480940] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595480965] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595481065] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595482107] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.3% (3888252 kB) free!
[1595482227] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.1% (3618580 kB) free!
[1595482347] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.8% (3316244 kB) free!
[1595482543] SERVICE ALERT: pservercluster2;HDD_Home status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595482765] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.06, 2.40, 2.53
[1595482865] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1595482921] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 350 processes
[1595482930] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1595482947] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1821216 kB) free!
[1595483003] SERVICE ALERT: pservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 36605 MB (71% inode=98%):
[1595483104] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1595483143] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 175788 MB (78% inode=99%):
[1595483173] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 81.8% (20044252 kB) free.
[1595483195] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1595483239] SERVICE ALERT: pservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=5.80% system=11.19% iowait=0.01% idle=83.00%
[1595483250] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 11520 open files (0% of max 2424156)
[1595483949] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1595484069] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1595484077] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1916) < 182.18.184.244 [23/Jul/2020:11:31:17 +0530] TCP 502 0 0 0.002 "10.147.212.61:5500, 10.147.212.49:5500, 10.147.212.38:5500, 10.147.212.37:5500" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 0.000, 0.000, 0.001"
[1595484189] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1595484198] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(479) < 182.18.184.244 [23/Jul/2020:11:33:17 +0530] TCP 502 0 0 0.000 "backend_5000" "0" "0" "0.000"
[1595484318] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(485) < 182.18.184.244 [23/Jul/2020:11:35:18 +0530] TCP 502 0 0 0.000 "backend_5002" "0" "0" "0.000"
[1595484578] SERVICE ALERT: smscapp1;Open-Files;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1595484698] SERVICE ALERT: smscapp1;Open-Files;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1595484789] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1595484789] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1595484810] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (526876 kB) free!
[1595484818] SERVICE ALERT: smscapp1;Open-Files;UNKNOWN;HARD;3;NRPE: Unable to read output
[1595484930] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 71.1% (17435916 kB) free.
[1595485189] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1595485230] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (1644920 kB) free!
[1595485350] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184412 kB) free!
[1595485389] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1595485471] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182312 kB) free!
[1595485509] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1595485530] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187908 kB) free!
[1595485629] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1595485650] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190736 kB) free!
[1595485765] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 10505.02, 3104.83, 1075.09
[1595485770] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190160 kB) free!
[1595485886] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15141.37, 7326.68, 2868.30
[1595485948] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 52.0% (11675484 kB) free.
[1595486006] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15765.79, 10151.19, 4443.92
[1595486229] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1595486370] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 55.3% (13565604 kB) free.
[1595486494] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (34.2% change >= 20.0% threshold)
[1595486774] SERVICE ALERT: webserver2;Open-Files;WARNING;HARD;4;NRPE: Unable to read output
[1595487032] SERVICE ALERT: sig1-vm2;Memory;WARNING;SOFT;1;WARNING - 18.5% (3028228 kB) free!
[1595487152] SERVICE ALERT: sig1-vm2;Memory;WARNING;SOFT;2;WARNING - 18.5% (3026388 kB) free!
[1595487272] SERVICE ALERT: sig1-vm2;Memory;WARNING;HARD;3;WARNING - 18.5% (3025512 kB) free!
[1595488218] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (31.8% change >= 20.0% threshold)
[1595488767] SERVICE ALERT: smscapp1;Open-Files;OK;HARD;1;OK: 32928 open files (0% of max 4878254)
[1595488879] SERVICE ALERT: webserver2;Open-Files;OK;HARD;1;OK: 9408 open files (0% of max 2423406)
[1595488967] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9433188 kB) free!
[1595489006] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1595489006] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.55, 6.44, 678.52
[1595489087] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.8% (9258736 kB) free!
[1595489207] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.4% (9063568 kB) free!
[1595489606] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.69, 3.02, 356.68
[1595489671] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (3831612 kB) free!
[1595489999] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1595490119] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1595490239] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1595490271] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186840 kB) free!
[1595490840] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1595491471] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1595491471] SERVICE ALERT: sservercluster2;Memory;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595491597] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595491597] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1595491639] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595491640] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595491682] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595491709] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595491717] SERVICE ALERT: sservercluster2;CPU-STATISTICS;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595491717] SERVICE ALERT: sservercluster2;CPU Load;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595491769] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595491770] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595491809] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595491812] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595491839] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595491847] SERVICE ALERT: sservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595491847] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595491899] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595491900] SERVICE ALERT: sservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595491939] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595491942] SERVICE ALERT: sservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595491969] SERVICE ALERT: sservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595491985] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595492020] SERVICE ALERT: sservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595492050] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595492069] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595492081] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595492115] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595492150] SERVICE ALERT: sservercluster2;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595492175] SERVICE ALERT: sservercluster2;CPU_Procs;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595492245] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595492280] SERVICE ALERT: sservercluster2;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595492305] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595494007] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4567960 kB) free!
[1595494170] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163172 kB) free!
[1595494290] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184092 kB) free!
[1595494410] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195624 kB) free!
[1595494948] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4396704 kB) free!
[1595495011] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1595495011] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.3% (22134808 kB) free.
[1595495068] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.8% (4225604 kB) free!
[1595495188] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.8% (4005464 kB) free!
[1595495594] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 19.4% (6335664 kB) free!
[1595495714] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 19.2% (6297316 kB) free!
[1595495834] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 19.1% (6244364 kB) free!
[1595495845] SERVICE ALERT: sservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1595495880] SERVICE ALERT: sservercluster2;Open-Files;OK;HARD;1;OK: 14592 open files (0% of max 2424092)
[1595495905] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1595495905] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1595496047] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.83, 2.12, 2.30
[1595496052] SERVICE ALERT: sservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=5.24% system=9.55% iowait=0.07% idle=85.14%
[1595496099] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1595496100] SERVICE ALERT: sservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 30879 MB (60% inode=98%):
[1595496142] SERVICE ALERT: sservercluster2;Total Processes;OK;HARD;1;PROCS OK: 342 processes
[1595496169] SERVICE ALERT: sservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 187270 MB (84% inode=99%):
[1595496269] SERVICE ALERT: sservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1595496281] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 83.5% (20457380 kB) free.
[1595496699] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1595496819] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1595496882] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181932 kB) free!
[1595496939] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1595496970] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1595496970] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (29.2% change >= 20.0% threshold)
[1595496989] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 51.9% (11655164 kB) free.
[1595497003] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193164 kB) free!
[1595497105] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1595497123] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (35.9% change >= 20.0% threshold)
[1595497123] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170660 kB) free!
[1595497130] SERVICE ALERT: pservercluster1;Secure Log Status;CRITICAL;SOFT;1;(2) < Jul 23 15:08:50 PSERVERCLUSTER1 sudo:  nagios : TTY=unknown : PWD=/ : USER=root : COMMAND=/usr/local/nagios/libexec/check_log -F /var/log/secure -O /tmp/secure.log -q Cannot allocate memory
[1595497226] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1595497250] SERVICE ALERT: pservercluster1;Secure Log Status;CRITICAL;SOFT;2;(1) < Jul 23 15:10:50 PSERVERCLUSTER1 sudo:  nagios : TTY=unknown : PWD=/ : USER=root : COMMAND=/usr/local/nagios/libexec/check_log -F /var/log/secure -O /tmp/secure.log -q Cannot allocate memory
[1595497319] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 24.0% (11828152 kB) free.
[1595497371] SERVICE ALERT: pservercluster1;Secure Log Status;CRITICAL;HARD;3;(1) < Jul 23 15:12:51 PSERVERCLUSTER1 sudo:  nagios : TTY=unknown : PWD=/ : USER=root : COMMAND=/usr/local/nagios/libexec/check_log -F /var/log/secure -O /tmp/secure.log -q Cannot allocate memory
[1595497663] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1595497663] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (24.8% change >= 20.0% threshold)
[1595497663] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (35.9% change >= 20.0% threshold)
[1595497762] SERVICE ALERT: pservercluster2;Secure Log Status;CRITICAL;SOFT;1;(2) < Jul 23 15:19:22 pservercluster2 sudo:  nagios : TTY=unknown : PWD=/ : USER=root : COMMAND=/usr/local/nagios/libexec/check_log -F /var/log/secure -O /tmp/secure.log -q Cannot allocate memory
[1595497882] SERVICE ALERT: pservercluster2;Secure Log Status;CRITICAL;SOFT;2;(1) < Jul 23 15:21:22 pservercluster2 sudo:  nagios : TTY=unknown : PWD=/ : USER=root : COMMAND=/usr/local/nagios/libexec/check_log -F /var/log/secure -O /tmp/secure.log -q Cannot allocate memory
[1595497919] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8655884 kB) free!
[1595498002] SERVICE ALERT: pservercluster2;Secure Log Status;CRITICAL;HARD;3;(1) < Jul 23 15:23:22 pservercluster2 sudo:  nagios : TTY=unknown : PWD=/ : USER=root : COMMAND=/usr/local/nagios/libexec/check_log -F /var/log/secure -O /tmp/secure.log -q Cannot allocate memory
[1595498021] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1595498021] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (24.8% change >= 20.0% threshold)
[1595498021] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (31.3% change >= 20.0% threshold)
[1595498039] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.5% (8641096 kB) free!
[1595498120] SERVICE ALERT: sservercluster1;Secure Log Status;CRITICAL;SOFT;1;(2) < Jul 23 15:25:20 sservercluster1 sudo:  nagios : TTY=unknown : PWD=/ : USER=root : COMMAND=/usr/local/nagios/libexec/check_log -F /var/log/secure -O /tmp/secure.log -q Cannot allocate memory
[1595498159] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.5% (8605340 kB) free!
[1595498240] SERVICE ALERT: sservercluster1;Secure Log Status;CRITICAL;SOFT;2;(1) < Jul 23 15:27:20 sservercluster1 sudo:  nagios : TTY=unknown : PWD=/ : USER=root : COMMAND=/usr/local/nagios/libexec/check_log -F /var/log/secure -O /tmp/secure.log -q Cannot allocate memory
[1595498360] SERVICE ALERT: sservercluster1;Secure Log Status;CRITICAL;HARD;3;(1) < Jul 23 15:29:20 sservercluster1 sudo:  nagios : TTY=unknown : PWD=/ : USER=root : COMMAND=/usr/local/nagios/libexec/check_log -F /var/log/secure -O /tmp/secure.log -q Cannot allocate memory
[1595498418] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1595498418] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1595498418] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (30.5% change >= 20.0% threshold)
[1595498516] SERVICE ALERT: sservercluster2;Secure Log Status;CRITICAL;SOFT;1;(2) < Jul 23 15:31:56 sservercluster2 sudo:  nagios : TTY=unknown : PWD=/ : USER=root : COMMAND=/usr/local/nagios/libexec/check_log -F /var/log/secure -O /tmp/secure.log -q Cannot allocate memory
[1595498636] SERVICE ALERT: sservercluster2;Secure Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1595499028] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1595499107] SERVICE ALERT: sservercluster1;Secure Log Status;UNKNOWN;HARD;3;Log check error: Log file /var/log/secure is not readable!
[1595499148] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1595499236] SERVICE ALERT: sservercluster2;Secure Log Status;UNKNOWN;SOFT;1;Log check error: Log file /var/log/secure is not readable!
[1595499356] SERVICE ALERT: sservercluster2;Secure Log Status;UNKNOWN;SOFT;2;Log check error: Log file /var/log/secure is not readable!
[1595499451] SERVICE ALERT: sservercluster1;Secure Log Status;CRITICAL;HARD;3;(2) < Jul 23 15:47:31 sservercluster1 sudo:  nagios : TTY=unknown : PWD=/ : USER=root : COMMAND=/usr/local/nagios/libexec/check_log -F /var/log/secure -O /tmp/secure.log -q Cannot allocate memory
[1595499476] SERVICE ALERT: sservercluster2;Secure Log Status;UNKNOWN;HARD;3;Log check error: Log file /var/log/secure is not readable!
[1595499523] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.8% (4858128 kB) free!
[1595499748] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1595499868] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1595499988] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1595500123] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180948 kB) free!
[1595501759] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4947116 kB) free!
[1595501788] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1595501962] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15900.92, 11610.69, 5463.86
[1595502082] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15921.94, 13038.25, 6732.29
[1595502202] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15925.09, 13993.59, 7846.97
[1595502357] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 12.7% (2070268 kB) free!
[1595502939] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1595504188] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1595504309] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1595504908] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1595505028] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1595505389] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4289068 kB) free!
[1595505509] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.6% (4172552 kB) free!
[1595505630] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 18.0% (4048848 kB) free!
[1595505802] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.05, 2.30, 213.68
[1595506521] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9615892 kB) free!
[1595506536] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1595506641] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (9549800 kB) free!
[1595506656] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1595506723] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 42.2% (10350600 kB) free.
[1595506762] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (9490880 kB) free!
[1595507234] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (3092264 kB) free!
[1595507323] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191848 kB) free!
[1595507430] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 51.7% (11617400 kB) free.
[1595507443] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 59.3% (14536064 kB) free.
[1595507834] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 10.7% (3491084 kB) free!
[1595508044] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170152 kB) free!
[1595508164] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199232 kB) free!
[1595508284] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179452 kB) free!
[1595508435] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (3169092 kB) free!
[1595508628] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1595508748] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1595508939] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1595509059] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1595509179] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1595509348] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1595509468] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1595509484] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (4284424 kB) free!
[1595509557] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1276532 kB) free!
[1595509589] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1595509779] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1595510084] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167368 kB) free!
[1595510189] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1595510379] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1595510499] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1595510619] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1595510684] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.3% (21409064 kB) free.
[1595514685] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 6324.34, 1535.92, 512.51
[1595514805] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 14736.69, 6337.78, 2396.24
[1595514925] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15870.51, 9550.39, 4051.42
[1595515762] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4926980 kB) free!
[1595516072] SERVICE ALERT: sig1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (1505824 kB) free!
[1595517630] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.3% (3891240 kB) free!
[1595517750] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.7% (3742008 kB) free!
[1595517870] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.1% (3621852 kB) free!
[1595518471] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 51.7% (11604264 kB) free.
[1595518526] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.62, 3.11, 364.22
[1595518635] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 44.4% (14528368 kB) free.
[1595520202] SERVICE ALERT: pservercluster2;Secure Log Status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595520203] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1595520206] SERVICE ALERT: pservercluster2;CPU Load;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1595520323] SERVICE ALERT: pservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 36438 MB (71% inode=98%):
[1595520326] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.76, 2.03, 2.12
[1595520380] SERVICE ALERT: pservercluster2;Memory;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1595520395] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;NRPE: Call to popen() failed
[1595520449] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520460] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520500] SERVICE ALERT: pservercluster2;Memory;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1595520522] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520525] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520579] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520590] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520630] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520652] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520655] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520675] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520709] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520720] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520731] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520740] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520782] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520805] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520812] SERVICE ALERT: pservercluster2;Secure Log Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520861] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520870] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520914] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520933] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520935] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520936] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520953] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595520991] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595521000] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595521044] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595521063] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595521066] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595521083] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595521174] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595521193] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595521196] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595521213] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595521589] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595521774] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1595521793] SERVICE ALERT: pservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 36438 MB (71% inode=98%):
[1595521796] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.17, 2.07, 2.01
[1595521813] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 143045 MB (63% inode=99%):
[1595521830] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 80.7% (19781648 kB) free.
[1595521855] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1595521914] SERVICE ALERT: pservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=4.12% system=7.71% iowait=0.00% idle=88.17%
[1595521920] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 10368 open files (0% of max 2424156)
[1595522019] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595522035] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 25.32, 32.57, 39.72
[1595522085] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1595522384] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522403] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522406] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522423] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522440] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522465] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522514] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522524] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522530] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522533] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522536] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522553] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522570] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522595] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522644] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522654] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522660] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522663] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522666] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522683] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522700] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522725] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522784] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595522790] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595524291] SERVICE ALERT: webserver2;Open-Files;WARNING;SOFT;1;NRPE: Unable to read output
[1595524353] SERVICE ALERT: webserver2;Open-Files;WARNING;SOFT;2;NRPE: Unable to read output
[1595524416] SERVICE ALERT: webserver2;Open-Files;WARNING;SOFT;3;NRPE: Unable to read output
[1595524478] SERVICE ALERT: webserver2;Open-Files;WARNING;HARD;4;NRPE: Unable to read output
[1595524535] SERVICE FLAPPING ALERT: pservercluster2;LocalService-NTPD SERVICE;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1595524535] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595524584] SERVICE FLAPPING ALERT: pservercluster2;CPU-STATISTICS;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1595524584] SERVICE ALERT: pservercluster2;CPU-STATISTICS;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1595525145] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1595525194] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1595528635] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 55.37, 35.28, 33.30