[1555181814] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45502 MB (20% inode=99%):
[1555181934] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57277 MB (25% inode=99%):
[1555183734] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45148 MB (20% inode=99%):
[1555183854] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57276 MB (25% inode=99%):
[1555185654] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45500 MB (20% inode=99%):
[1555185774] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57275 MB (25% inode=99%):
[1555189374] SERVICE FLAPPING ALERT: db2;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1555189970] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 478.00, 313.52, 138.89
[1555190090] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 480.50, 368.93, 180.37
[1555190210] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 83.82, 263.59, 165.44
[1555190484] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2727748 kB) free!
[1555190574] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45061 MB (20% inode=99%):
[1555190694] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57254 MB (25% inode=99%):
[1555191217] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1555192235] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (733700 kB) free!
[1555192494] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44962 MB (20% inode=99%):
[1555192610] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.99, 0.93, 13.22
[1555192614] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57252 MB (25% inode=99%):
[1555192884] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2465556 kB) free!
[1555193274] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (670836 kB) free!
[1555193794] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (633548 kB) free!
[1555193810] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.98, 0.84, 4.20
[1555193995] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 427.29, 167.86, 62.87
[1555194115] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 479.36, 273.67, 114.38
[1555194235] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 486.15, 344.33, 159.59
[1555194414] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46099 MB (20% inode=99%):
[1555194447] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 71.57, 245.30, 151.78
[1555194534] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57234 MB (25% inode=99%):
[1555194567] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 9.84, 164.22, 133.41
[1555194687] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.45, 109.93, 117.25
[1555194892] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (2757432 kB) free!
[1555195418] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1555196018] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555196138] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1555196258] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1555196487] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.26, 0.50, 17.08
[1555197235] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.56, 1.51, 13.69
[1555197552] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 459.83, 311.44, 140.28
[1555197672] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 236.60, 303.25, 159.77
[1555197687] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.13, 0.25, 4.87
[1555197693] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.9% (2921288 kB) free!
[1555197792] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 32.41, 203.07, 140.45
[1555198135] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45946 MB (20% inode=99%):
[1555198255] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57146 MB (25% inode=99%):
[1555198435] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.26, 1.44, 4.85
[1555198855] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45105 MB (20% inode=99%):
[1555198975] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45105 MB (20% inode=99%):
[1555199095] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45112 MB (20% inode=99%):
[1555199695] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49017 MB (21% inode=99%):
[1555200192] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.58, 0.68, 11.17
[1555201387] SERVICE ALERT: sservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46331 MB (20% inode=99%):
[1555201392] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.91, 0.94, 3.69
[1555201507] SERVICE ALERT: sservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46330 MB (20% inode=99%):
[1555201627] SERVICE ALERT: sservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46330 MB (20% inode=99%):
[1555203875] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;1;WARNING - 20.0% (1470560 kB) free!
[1555203893] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2324840 kB) free!
[1555203895] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44936 MB (19% inode=99%):
[1555203996] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;2;WARNING - 20.0% (1468544 kB) free!
[1555204015] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57213 MB (25% inode=99%):
[1555204116] SERVICE ALERT: sig4-vm1;Memory;WARNING;HARD;3;WARNING - 19.9% (1463360 kB) free!
[1555205493] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2459920 kB) free!
[1555205815] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47164 MB (20% inode=99%):
[1555205936] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45115 MB (20% inode=99%):
[1555206056] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56699 MB (25% inode=99%):
[1555207499] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41034 MB (18% inode=99%):
[1555209000] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555209056] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45409 MB (20% inode=99%):
[1555209120] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555209145] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 55.4% (9001316 kB) free.
[1555209176] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56928 MB (25% inode=99%):
[1555209240] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555209299] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 72637 MB (32% inode=99%):
[1555209299] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555209419] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555209519] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1747092 kB) free!
[1555209539] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555210739] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555211040] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555212428] SERVICE ALERT: sservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52654 MB (23% inode=99%):
[1555212459] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1555212840] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555212960] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555213059] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555213080] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555213139] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555213179] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555213259] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555213299] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1555213376] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45288 MB (20% inode=99%):
[1555213379] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555213496] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56932 MB (25% inode=99%):
[1555214280] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1555214280] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555214579] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1555214579] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555215296] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45216 MB (20% inode=99%):
[1555215416] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57055 MB (25% inode=99%):
[1555216010] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 101.90, 283.86, 176.15
[1555216130] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 14.54, 190.23, 154.90
[1555216250] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.88, 127.71, 136.27
[1555217316] SERVICE ALERT: sig4-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (621524 kB) free!
[1555218323] SERVICE ALERT: sig4-vm2;Memory;WARNING;SOFT;1;WARNING - 19.2% (1410868 kB) free!
[1555218416] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46931 MB (20% inode=99%):
[1555218443] SERVICE ALERT: sig4-vm2;Memory;WARNING;SOFT;2;WARNING - 19.0% (1400212 kB) free!
[1555218536] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44686 MB (19% inode=99%):
[1555218563] SERVICE ALERT: sig4-vm2;Memory;WARNING;HARD;3;WARNING - 18.9% (1391772 kB) free!
[1555218650] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.80, 1.05, 11.33
[1555218656] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56197 MB (25% inode=99%):
[1555219288] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 23.87, 194.82, 140.30
[1555219408] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.63, 130.50, 123.34
[1555219493] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.4% (2805252 kB) free!
[1555219528] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.98, 87.49, 108.45
[1555219850] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.58, 0.88, 3.82
[1555219945] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.7% (3195252 kB) free!
[1555220065] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.1% (3111192 kB) free!
[1555220185] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.6% (3021400 kB) free!
[1555220456] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46358 MB (20% inode=99%):
[1555220576] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46358 MB (20% inode=99%):
[1555220696] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46358 MB (20% inode=99%):
[1555221294] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2457252 kB) free!
[1555221296] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48941 MB (21% inode=99%):
[1555221328] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.51, 0.59, 15.90
[1555221894] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.8% (2651036 kB) free!
[1555221896] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43631 MB (19% inode=99%):
[1555221959] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 325 processes
[1555222016] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56870 MB (25% inode=99%):
[1555222528] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.16, 0.18, 4.50
[1555222585] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1314200 kB) free!
[1555223094] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1555223094] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2448912 kB) free!
[1555225080] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1555225163] SERVICE ALERT: sig4-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (725668 kB) free!
[1555225299] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 315 processes
[1555225379] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1555225899] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1555226019] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1555226140] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555226161] SERVICE ALERT: sig5-vm1;Memory;WARNING;SOFT;1;WARNING - 19.6% (1441520 kB) free!
[1555226216] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46674 MB (20% inode=99%):
[1555226281] SERVICE ALERT: sig5-vm1;Memory;WARNING;SOFT;2;WARNING - 19.5% (1433336 kB) free!
[1555226336] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44625 MB (19% inode=99%):
[1555226401] SERVICE ALERT: sig5-vm1;Memory;WARNING;HARD;3;WARNING - 19.4% (1425764 kB) free!
[1555226456] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56651 MB (25% inode=99%):
[1555227294] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 24.1% (5918260 kB) free.
[1555227519] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (3887804 kB) free!
[1555227650] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 336 processes
[1555227693] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 41.4% (10152052 kB) free.
[1555227959] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 1 warn out of 326 processes
[1555228119] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 39.4% (9648768 kB) free.
[1555228250] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 337 processes
[1555228256] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46662 MB (20% inode=99%):
[1555228376] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44612 MB (19% inode=99%):
[1555228496] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56129 MB (24% inode=99%):
[1555228634] SERVICE ALERT: sig6-vm1;Memory;WARNING;SOFT;1;WARNING - 19.9% (1466640 kB) free!
[1555228754] SERVICE ALERT: sig6-vm1;Memory;WARNING;SOFT;2;WARNING - 19.9% (1461928 kB) free!
[1555228874] SERVICE ALERT: sig6-vm1;Memory;WARNING;HARD;3;WARNING - 19.8% (1456224 kB) free!
[1555229096] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46847 MB (20% inode=99%):
[1555229216] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46847 MB (20% inode=99%):
[1555229336] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43212 MB (19% inode=99%):
[1555229936] SERVICE FLAPPING ALERT: db2;HDD_Home status;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1555229936] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48521 MB (21% inode=99%):
[1555230536] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43959 MB (19% inode=99%):
[1555230656] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56692 MB (25% inode=99%):
[1555230928] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.33, 5.97, 5.33
[1555231048] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.36, 4.11, 4.73
[1555232456] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44394 MB (19% inode=99%):
[1555232576] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56687 MB (25% inode=99%):
[1555233176] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44949 MB (20% inode=99%):
[1555233296] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44949 MB (20% inode=99%):
[1555233417] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44956 MB (20% inode=99%):
[1555234017] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48477 MB (21% inode=99%):
[1555234541] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1555235094] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1555235141] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555235217] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46494 MB (20% inode=99%):
[1555235261] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1555235337] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46494 MB (20% inode=99%):
[1555235381] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1555235381] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1555235457] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42910 MB (19% inode=99%):
[1555236057] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48445 MB (21% inode=99%):
[1555237257] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46907 MB (20% inode=99%):
[1555237377] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44859 MB (19% inode=99%):
[1555237497] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56570 MB (25% inode=99%):
[1555239193] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (1468444 kB) free!
[1555239298] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46905 MB (20% inode=99%):
[1555239313] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 19.9% (1465764 kB) free!
[1555239418] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46905 MB (20% inode=99%):
[1555239433] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 19.9% (1463260 kB) free!
[1555239538] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44857 MB (19% inode=99%):
[1555240138] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48439 MB (21% inode=99%):
[1555240738] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44343 MB (19% inode=99%):
[1555240781] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1555240858] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56117 MB (24% inode=99%):
[1555241381] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1555241501] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1555241622] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1555242658] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44268 MB (19% inode=99%):
[1555242778] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44275 MB (19% inode=99%):
[1555242899] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56561 MB (25% inode=99%):
[1555244694] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4880096 kB) free!
[1555244698] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44974 MB (20% inode=99%):
[1555244814] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (4868160 kB) free!
[1555244818] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56558 MB (25% inode=99%):
[1555244934] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.8% (4853712 kB) free!
[1555246618] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44259 MB (19% inode=99%):
[1555246738] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44266 MB (19% inode=99%):
[1555246858] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56552 MB (25% inode=99%):
[1555247401] SERVICE ALERT: sig5-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (699504 kB) free!
[1555247592] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 408.08, 143.47, 51.91
[1555247635] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 493.30, 257.73, 104.37
[1555247712] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 489.83, 262.31, 106.56
[1555247755] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 506.81, 340.88, 153.44
[1555247832] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 500.71, 341.74, 154.55
[1555247871] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.9% (1460968 kB) free!
[1555247875] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 508.82, 396.57, 196.57
[1555247991] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 19.8% (1455628 kB) free!
[1555248111] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 19.7% (1451420 kB) free!
[1555248658] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43675 MB (19% inode=99%):
[1555248778] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56416 MB (25% inode=99%):
[1555250232] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.91, 0.97, 17.49
[1555250474] SERVICE ALERT: sig6-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (723764 kB) free!
[1555250578] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44116 MB (19% inode=99%):
[1555250698] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44123 MB (19% inode=99%):
[1555250818] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56410 MB (25% inode=99%):
[1555250875] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.77, 1.84, 12.79
[1555252033] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.64, 0.64, 3.06
[1555252075] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.61, 1.65, 4.75
[1555252618] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 39721 MB (17% inode=99%):
[1555252738] SERVICE FLAPPING ALERT: db2;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1555252738] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56250 MB (25% inode=99%):
[1555253022] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1555254538] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43958 MB (19% inode=99%):
[1555254658] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43958 MB (19% inode=99%):
[1555254779] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43958 MB (19% inode=99%):
[1555255134] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2366328 kB) free!
[1555255379] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48050 MB (21% inode=99%):
[1555257779] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46719 MB (20% inode=99%):
[1555257899] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44670 MB (19% inode=99%):
[1555258019] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56445 MB (25% inode=99%):
[1555259819] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44658 MB (19% inode=99%):
[1555259939] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56418 MB (25% inode=99%):
[1555262450] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 441.63, 239.60, 97.94
[1555262570] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 450.34, 309.75, 140.82
[1555262690] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 451.63, 356.68, 178.49
[1555262848] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 389.64, 158.52, 59.19
[1555262939] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44099 MB (19% inode=99%):
[1555262968] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 429.12, 250.13, 104.81
[1555263059] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56387 MB (25% inode=99%):
[1555263084] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (2760080 kB) free!
[1555263088] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 434.63, 311.46, 144.90
[1555263534] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.8% (2890748 kB) free!
[1555263601] SERVICE ALERT: sig5-vm1;Memory;OK;HARD;1;OK - 96.0% (7061292 kB) free.
[1555263636] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 96.2% (7077324 kB) free.
[1555263675] SERVICE ALERT: sig6-vm1;Memory;OK;HARD;1;OK - 96.0% (7059884 kB) free.
[1555263712] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 96.0% (7062032 kB) free.
[1555263995] SERVICE ALERT: sig2-vm1;Memory;OK;HARD;1;OK - 96.0% (7062436 kB) free.
[1555264034] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 95.9% (7054496 kB) free.
[1555264075] SERVICE ALERT: sig3-vm1;Memory;OK;HARD;1;OK - 96.0% (7062832 kB) free.
[1555264103] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 87.1% (14158396 kB) free.
[1555264116] SERVICE ALERT: sig4-vm1;Memory;OK;HARD;1;OK - 96.1% (7065276 kB) free.
[1555264163] SERVICE ALERT: sig4-vm2;Memory;OK;HARD;1;OK - 96.1% (7064904 kB) free.
[1555264859] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44097 MB (19% inode=99%):
[1555264979] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56386 MB (25% inode=99%):
[1555265023] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1555265090] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.57, 0.82, 14.93
[1555265488] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.35, 0.44, 12.88
[1555265623] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1555265743] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1555265863] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1555266291] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.63, 0.81, 4.68