[1586975846] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6284748 kB) free!
[1586975966] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29828104 kB) free.
[1586976016] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (202160 kB) free!
[1586976136] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186136 kB) free!
[1586976256] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186148 kB) free!
[1586976566] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8434920 kB) free!
[1586976686] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.2% (20100176 kB) free!
[1586976806] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26894236 kB) free.
[1586977239] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1586977406] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6158116 kB) free!
[1586977526] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28156752 kB) free.
[1586977543] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180024 kB) free!
[1586977604] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (201004 kB) free!
[1586977664] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (195956 kB) free!
[1586977725] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (177184 kB) free!
[1586977994] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1613812 kB) free!
[1586978025] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 71.3% (17479100 kB) free.
[1586978126] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7325028 kB) free!
[1586978246] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29271188 kB) free.
[1586978846] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7876344 kB) free!
[1586978966] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29490116 kB) free.
[1586979567] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8107048 kB) free!
[1586979687] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29414968 kB) free.
[1586980239] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.3% (20204764 kB) free!
[1586980287] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8568988 kB) free!
[1586980359] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (9523392 kB) free!
[1586980407] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29336988 kB) free.
[1586980479] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (457328 kB) free!
[1586980799] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (232256 kB) free!
[1586980919] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (240724 kB) free!
[1586981007] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8787388 kB) free!
[1586981039] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237240 kB) free!
[1586981127] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29274012 kB) free.
[1586981725] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1586981727] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8264344 kB) free!
[1586981847] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29164780 kB) free.
[1586982447] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7727208 kB) free!
[1586982567] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29092716 kB) free.
[1586982839] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.4% (8096976 kB) free!
[1586983167] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6918512 kB) free!
[1586983287] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28538456 kB) free.
[1586983554] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 28.4% (6951044 kB) free.
[1586983887] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6891128 kB) free!
[1586984007] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26695328 kB) free.
[1586984154] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199612 kB) free!
[1586984275] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191140 kB) free!
[1586984394] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 68.5% (33760740 kB) free.
[1586984395] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185740 kB) free!
[1586984594] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.1% (2126108 kB) free!
[1586984935] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.7% (42748312 kB) free.
[1586985194] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.4% (4774900 kB) free.
[1586985207] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9625340 kB) free!
[1586985328] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28921024 kB) free.
[1586985927] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8061484 kB) free!
[1586986047] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26487500 kB) free.
[1586986260] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(255) < 103.16.101.72 [16/Apr/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1586986281] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1838) < 103.13.97.241 [16/Apr/2020:03:00:33 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1586986380] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586986394] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(306) < 2020/04/16 03:00:22 [error] 22120#0: *84359708 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1586986401] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(63) < 203.212.70.170 [16/Apr/2020:03:02:32 +0530] TCP 502 0 0 120.000 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1586986440] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.5% (44577508 kB) free.
[1586986514] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586986521] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586987247] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9724444 kB) free!
[1586987367] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28525792 kB) free.
[1586988567] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10430904 kB) free!
[1586988687] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27440376 kB) free.
[1586992287] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11286720 kB) free!
[1586992408] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27133260 kB) free.
[1586994207] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (26005708 kB) free!
[1586994256] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 80.5% (19719260 kB) free.
[1586994327] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.3% (25462160 kB) free!
[1586994447] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20565336 kB) free!
[1586994591] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 81.7% (20024164 kB) free.
[1586994702] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 59.6% (14706284 kB) free.
[1586996594] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1586998902] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (934912 kB) free!
[1586999022] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (887332 kB) free!
[1586999079] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 36.9% (48642256 kB) free.
[1586999142] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (858660 kB) free!
[1587003447] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28855336 kB) free.
[1587005247] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (26030716 kB) free!
[1587005367] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.5% (21762160 kB) free!
[1587005487] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17476596 kB) free!
[1587006567] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.62, 6.46, 6.16
[1587006687] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.14, 6.11, 6.07
[1587006687] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28642844 kB) free.
[1587006807] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.36, 5.89, 6.00
[1587007887] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26246680 kB) free!
[1587008007] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42246232 kB) free.
[1587008391] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (4501952 kB) free!
[1587008511] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (4359904 kB) free!
[1587008631] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (4235944 kB) free!
[1587009914] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(360) < 2020/04/16 09:35:14 [error] 22135#0: *84399105 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1587010034] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(43) < 2020/04/16 09:35:56 [error] 22131#0: *84399337 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1587010154] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587010380] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(301) < 172.105.57.57 [16/Apr/2020:09:35:47 +0530] TCP 502 0 0 60.027 "10.147.212.38:4000, backend3" "0, 0" "0, 0" "60.027, 0.000"
[1587010500] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587010521] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(185) < 103.16.101.72 [16/Apr/2020:09:37:58 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1587010641] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587011007] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8890860 kB) free!
[1587011031] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1926276 kB) free!
[1587011127] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32392160 kB) free.
[1587012327] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (14962824 kB) free!
[1587012447] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33871548 kB) free.
[1587013647] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24226156 kB) free!
[1587013767] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30360516 kB) free.
[1587014967] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8035640 kB) free!
[1587015012] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.3% (2986752 kB) free!
[1587015087] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28185568 kB) free.
[1587015132] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.2% (2511832 kB) free!
[1587015252] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2040220 kB) free!
[1587015442] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(326) < 103.255.146.164 [16/Apr/2020:11:01:50 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587015562] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587015688] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3275900 kB) free!
[1587015808] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34494344 kB) free.
[1587017008] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5161324 kB) free!
[1587017128] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30517624 kB) free.
[1587017594] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3046896 kB) free!
[1587017714] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (2956960 kB) free!
[1587017728] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (4806900 kB) free!
[1587017834] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3525200 kB) free.
[1587017848] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (382748 kB) free!
[1587017968] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33184512 kB) free.
[1587018279] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.5% (23083860 kB) free!
[1587018399] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.2% (31863996 kB) free.
[1587018434] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3182428 kB) free!
[1587018554] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (3122048 kB) free!
[1587018568] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4061492 kB) free!
[1587018674] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.3% (2328268 kB) free!
[1587018688] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33431876 kB) free.
[1587018817] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.7% (4838260 kB) free!
[1587018852] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.4% (16261904 kB) free.
[1587018937] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.6% (4798948 kB) free!
[1587019058] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.4% (4760340 kB) free!
[1587019288] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8546212 kB) free!
[1587019408] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4361904 kB) free!
[1587019528] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29784280 kB) free.
[1587019599] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.0% (25009900 kB) free!
[1587019719] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.7% (29987764 kB) free.
[1587020056] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (4535404 kB) free!
[1587020128] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5291900 kB) free!
[1587020176] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4462572 kB) free!
[1587020248] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30189260 kB) free.
[1587020296] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (4379252 kB) free!
[1587020849] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4621368 kB) free!
[1587020962] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(321) < 103.255.146.164 [16/Apr/2020:12:31:25 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587020969] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28941512 kB) free.
[1587020994] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8894236 kB) free!
[1587021082] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587021114] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (8683884 kB) free!
[1587021234] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (8433240 kB) free!
[1587021519] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.7% (24623024 kB) free!
[1587021569] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (3973392 kB) free!
[1587021639] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.0% (30304192 kB) free.
[1587021682] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(147) < 103.255.146.164 [16/Apr/2020:12:49:45 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587021689] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (395552 kB) free!
[1587021803] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587021809] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26843880 kB) free.
[1587022135] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9817080 kB) free!
[1587022255] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.7% (9709092 kB) free!
[1587022375] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.1% (9432276 kB) free!
[1587022409] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (2354384 kB) free!
[1587022529] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27913708 kB) free.
[1587022696] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1898152 kB) free!
[1587022839] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.2% (25297344 kB) free!
[1587022959] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.4% (30846580 kB) free.
[1587023129] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (3495396 kB) free!
[1587023249] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (386236 kB) free!
[1587023369] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 20.0% (26303288 kB) free!
[1587023559] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.2% (22683644 kB) free!
[1587023602] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(467) < 103.255.146.164 [16/Apr/2020:13:15:34 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587023634] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4271752 kB) free!
[1587023679] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.3% (29397384 kB) free.
[1587023722] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587023969] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (1763032 kB) free!
[1587024279] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.3% (24113388 kB) free!
[1587024399] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.5% (29599400 kB) free.
[1587024775] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4357224 kB) free!
[1587024999] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.0% (22360792 kB) free!
[1587025119] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.5% (28370452 kB) free.
[1587025169] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14574932 kB) free!
[1587025452] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.1% (4074652 kB) free!
[1587025522] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(223) < 103.255.146.164 [16/Apr/2020:13:49:15 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587025572] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.9% (3567448 kB) free!
[1587025642] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587025692] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.4% (3228016 kB) free!
[1587025719] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.2% (22730288 kB) free!
[1587025770] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1587025770] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32438668 kB) free.
[1587025839] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 12.7% (16730652 kB) free!
[1587025959] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.3% (20135144 kB) free!
[1587026292] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1587026292] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (1041452 kB) free!
[1587026370] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7685488 kB) free!
[1587026490] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30221028 kB) free.
[1587027090] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7060948 kB) free!
[1587027159] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.7% (28648132 kB) free.
[1587027210] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (3144932 kB) free!
[1587027330] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28978740 kB) free.
[1587027442] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(128) < 103.255.146.164 [16/Apr/2020:14:20:17 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587027563] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587027759] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.9% (26246428 kB) free!
[1587027879] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.0% (19830296 kB) free!
[1587027930] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (5021112 kB) free!
[1587027999] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.5% (27081868 kB) free.
[1587028050] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (714900 kB) free!
[1587028170] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29234624 kB) free.
[1587028274] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (617952 kB) free!
[1587028599] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.5% (17855152 kB) free!
[1587028719] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (10691656 kB) free!
[1587028770] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (2820716 kB) free!
[1587028839] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.8% (24827748 kB) free!
[1587028890] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (378428 kB) free!
[1587029010] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32118212 kB) free.
[1587029474] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.1% (1972000 kB) free!
[1587029610] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7896808 kB) free!
[1587029730] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3580668 kB) free!
[1587029850] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386380 kB) free!
[1587029893] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.8% (16347104 kB) free.
[1587029962] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(144) < 185.255.8.58 [16/Apr/2020:15:00:05 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1587030074] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1587030074] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3271616 kB) free.
[1587030082] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587030639] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1587030639] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (10243772 kB) free!
[1587030675] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.8% (2074392 kB) free!
[1587030683] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(66) < 103.255.146.164 [16/Apr/2020:15:13:43 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587030795] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.6% (2376700 kB) free!
[1587030803] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587030915] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.1% (2613252 kB) free!
[1587031036] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1587031050] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19467548 kB) free!
[1587031156] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 334 processes
[1587031276] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1587031403] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(110) < 103.255.146.164 [16/Apr/2020:15:23:50 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587031523] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587031650] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30188084 kB) free.
[1587031839] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.6% (15262064 kB) free!
[1587032123] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(139) < 103.255.146.164 [16/Apr/2020:15:41:18 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587032243] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587032250] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5868728 kB) free!
[1587032370] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1825428 kB) free!
[1587032490] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30014384 kB) free.
[1587032742] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.3% (3036568 kB) free!
[1587033090] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6960648 kB) free!
[1587033173] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.9% (6857664 kB) free!
[1587033210] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2718084 kB) free!
[1587033330] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29518880 kB) free.
[1587033342] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2154808 kB) free!
[1587033443] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(196) < 103.255.146.164 [16/Apr/2020:16:00:48 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587033563] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587033773] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (1342040 kB) free!
[1587033930] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6852464 kB) free!
[1587034050] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (2507004 kB) free!
[1587034170] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29357648 kB) free.
[1587034763] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(259) < 103.255.146.164 [16/Apr/2020:16:21:19 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587034770] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5820272 kB) free!
[1587034883] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587034890] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1674384 kB) free!
[1587035010] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30244216 kB) free.
[1587035115] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (717260 kB) free!
[1587035611] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6043264 kB) free!
[1587035715] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.5% (2518832 kB) free!
[1587035731] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1809492 kB) free!
[1587035851] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399596 kB) free!
[1587035858] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.9% (2417484 kB) free!
[1587036315] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1212304 kB) free!
[1587036915] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (3130892 kB) free!
[1587037051] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19829896 kB) free!
[1587037651] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30859628 kB) free.
[1587038115] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (1368104 kB) free!
[1587038251] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7750376 kB) free!
[1587038293] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1587038371] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3567296 kB) free!
[1587038491] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30270664 kB) free.
[1587038715] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (2586176 kB) free!
[1587039091] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7243024 kB) free!
[1587039211] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3284144 kB) free!
[1587039331] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381092 kB) free!
[1587039639] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (12385672 kB) free!
[1587039684] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(137) < 103.255.146.164 [16/Apr/2020:17:43:46 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587039804] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587039807] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.75, 3.63, 4.33
[1587040093] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.8% (3771768 kB) free!
[1587040214] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.6% (16299340 kB) free.
[1587040239] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.2% (26566996 kB) free.
[1587040407] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.50, 7.60, 6.04
[1587040527] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.27, 7.51, 6.21
[1587040531] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19234868 kB) free!
[1587040648] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.07, 7.16, 6.25
[1587040839] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.8% (22085000 kB) free!
[1587040959] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 10.9% (14350948 kB) free!
[1587041004] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(562) < 103.255.146.164 [16/Apr/2020:18:11:34 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587041079] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.4% (25530136 kB) free!
[1587041124] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587041131] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381836 kB) free!
[1587041715] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1269956 kB) free!
[1587042076] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1587042196] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1587042315] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (2786032 kB) free!
[1587042331] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15525108 kB) free!
[1587043531] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1175616 kB) free!
[1587043648] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.46, 3.04, 4.30
[1587044115] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1025268 kB) free!
[1587044131] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16488868 kB) free!
[1587044248] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.35, 6.68, 5.38
[1587044368] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.00, 6.41, 5.46
[1587044488] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1587044488] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.10, 5.92, 5.39
[1587044731] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33162284 kB) free.
[1587045279] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.3% (28070916 kB) free.
[1587045315] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (3054880 kB) free!
[1587045331] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7351536 kB) free!
[1587045451] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (3142744 kB) free!
[1587045571] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33583032 kB) free.
[1587045688] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.75, 4.42, 4.92
[1587045879] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.1% (25180316 kB) free!
[1587045999] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.8% (18211168 kB) free!
[1587046119] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.4% (25602040 kB) free!
[1587046171] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9886580 kB) free!
[1587046291] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (5790640 kB) free!
[1587046411] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (1375300 kB) free!
[1587047611] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.4% (22931108 kB) free!
[1587048211] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (374616 kB) free!
[1587048315] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1033184 kB) free!
[1587048519] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.4% (33458720 kB) free.
[1587048916] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (2953696 kB) free!
[1587049411] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.4% (25530108 kB) free!
[1587049720] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.8% (22208884 kB) free!
[1587049840] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.8% (27376452 kB) free.
[1587050011] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396124 kB) free!
[1587050716] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3291744 kB) free.
[1587051211] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.4% (25573092 kB) free!
[1587051811] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3399032 kB) free!
[1587052411] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14277148 kB) free!
[1587053116] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (2855976 kB) free!
[1587053236] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (3091692 kB) free!
[1587053356] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3366808 kB) free.
[1587053611] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (3331404 kB) free!
[1587054325] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1659 processes
[1587054554] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/04/16 21:55:53 [error] 22131#0: *84482491 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1587054674] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587054811] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.2% (22735040 kB) free!
[1587054925] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1808 processes
[1587055411] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (664452 kB) free!
[1587056125] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1679 processes
[1587056488] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1587056611] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20817252 kB) free!
[1587056725] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1587056725] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1824 processes
[1587057211] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32457320 kB) free.
[1587057811] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7587796 kB) free!
[1587057932] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3666508 kB) free!
[1587058052] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32182960 kB) free.
[1587058240] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1587058652] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8030984 kB) free!
[1587058772] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4241708 kB) free!
[1587058892] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31999776 kB) free.
[1587059492] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9301072 kB) free!
[1587059612] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (5678244 kB) free!
[1587059732] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31909232 kB) free.
[1587059956] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.1% (2456504 kB) free!
[1587060076] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.6% (3828356 kB) free.
[1587060332] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9587528 kB) free!
[1587060452] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (5949432 kB) free!
[1587060572] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31238564 kB) free.
[1587061172] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9469336 kB) free!
[1587061276] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1587061292] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (5903844 kB) free!
[1587061412] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29862300 kB) free.