[1559586605] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.4% (6867200 kB) free.
[1559586606] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559586726] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559586846] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559587013] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (163832 kB) free!
[1559587524] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559587613] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.6% (3995528 kB) free.
[1559588813] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.0% (2111640 kB) free!
[1559588839] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.1% (10391452 kB) free.
[1559588933] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.1% (3428112 kB) free.
[1559589265] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 380 processes
[1559589534] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (1377796 kB) free!
[1559589654] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.5% (3330508 kB) free.
[1559589865] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 404 processes
[1559589985] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 396 processes
[1559590197] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559590254] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (1117324 kB) free!
[1559590374] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.5% (3331320 kB) free.
[1559590585] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 396 processes
[1559590705] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 388 processes
[1559590825] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 392 processes
[1559590974] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (1387512 kB) free!
[1559591094] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.7% (3365532 kB) free.
[1559591239] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9284972 kB) free!
[1559591359] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.9% (7331448 kB) free!
[1559591469] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559591479] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.5% (5655808 kB) free!
[1559591694] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (1163416 kB) free!
[1559591814] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.2% (3438436 kB) free.
[1559592026] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1559592414] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (814380 kB) free!
[1559592534] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.9% (2900400 kB) free!
[1559592626] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 374 processes
[1559592654] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.9% (3547232 kB) free.
[1559592746] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 374 processes
[1559592846] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559593254] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.3% (2644104 kB) free!
[1559593346] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 398 processes
[1559593374] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.1% (3579856 kB) free.
[1559593466] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 394 processes
[1559593879] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4586660 kB) free!
[1559593974] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (176272 kB) free!
[1559594066] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 374 processes
[1559594095] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (175152 kB) free!
[1559594186] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 398 processes
[1559594192] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 64.19, 19.63, 7.80
[1559594215] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1526076 kB) free!
[1559594306] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 397 processes
[1559594312] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 96.08, 46.55, 19.09
[1559594432] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 100.30, 64.57, 29.03
[1559594646] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559594674] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5126 MB (10% inode=97%):
[1559594766] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559594815] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.6% (1881512 kB) free!
[1559594886] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1559594886] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559594906] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 382 processes
[1559595026] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 384 processes
[1559595275] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6970 MB (13% inode=98%):
[1559595448] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4890492 kB) free!
[1559595568] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (4877204 kB) free!
[1559595626] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 366 processes
[1559595632] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.56, 5.20, 17.94
[1559595688] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (4867184 kB) free!
[1559595746] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 366 processes
[1559595867] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 371 processes
[1559596015] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1207932 kB) free!
[1559596232] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 80.71, 19.11, 15.71
[1559596279] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.8% (5792220 kB) free!
[1559596615] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.4% (3313900 kB) free.
[1559596686] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559597013] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.4% (2799776 kB) free!
[1559597215] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (1131576 kB) free!
[1559597335] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.7% (2871380 kB) free!
[1559597455] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.0% (3728172 kB) free.
[1559597887] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559598007] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559598055] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (747812 kB) free!
[1559598127] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559598175] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.6% (2205620 kB) free!
[1559598295] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.1% (3906256 kB) free.
[1559598679] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10310936 kB) free.
[1559598727] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1559598813] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2461288 kB) free!
[1559598867] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 384 processes
[1559599232] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1559599232] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.83, 1.93, 19.83
[1559599279] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (2967568 kB) free!
[1559599399] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10318292 kB) free.
[1559599467] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 387 processes
[1559599495] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.8% (3043568 kB) free!
[1559599587] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 387 processes
[1559599615] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.6% (3994512 kB) free.
[1559599707] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 382 processes
[1559599927] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559600047] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559600167] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559600307] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 390 processes
[1559600427] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 374 processes
[1559600547] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 376 processes
[1559600997] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 139.03, 57.81, 21.82
[1559601032] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.73, 1.93, 4.56
[1559601117] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 152.18, 89.73, 37.88
[1559601199] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9633480 kB) free!
[1559601237] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 154.05, 111.09, 52.00
[1559601319] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.7% (7733900 kB) free!
[1559601322] SERVICE ALERT: webserver1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 354 processes
[1559601439] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6154520 kB) free!
[1559601442] SERVICE ALERT: webserver1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 352 processes
[1559601562] SERVICE ALERT: webserver1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 353 processes
[1559601688] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 20.5% (5030296 kB) free.
[1559602015] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 10.1% (1646604 kB) free!
[1559602135] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.4% (2342372 kB) free!
[1559602255] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.0% (3083928 kB) free!
[1559602347] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 385 processes
[1559602437] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.19, 2.77, 17.03
[1559602855] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.7% (3350272 kB) free.
[1559602947] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 395 processes
[1559603067] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1559603216] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2461520 kB) free!
[1559603488] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4898552 kB) free!
[1559603608] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (4888556 kB) free!
[1559603637] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.08, 0.24, 4.85
[1559603728] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (4878924 kB) free!
[1559603839] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4826212 kB) free!
[1559604267] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 391 processes
[1559604387] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 385 processes
[1559604507] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 376 processes
[1559604655] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.7% (2716332 kB) free!
[1559604775] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 16.7% (2714084 kB) free!
[1559604895] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.7% (2717240 kB) free!
[1559605270] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 10.41, 55.83, 37.20
[1559605390] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.64, 37.49, 32.74
[1559605495] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (754176 kB) free!
[1559605510] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.71, 25.24, 28.83
[1559605616] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.1% (2993428 kB) free!
[1559605832] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.11, 5.91, 5.09
[1559605952] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.25, 4.75, 4.77
[1559606013] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.1% (3943176 kB) free!
[1559606095] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.9% (3068672 kB) free!
[1559606110] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.67, 3.82, 15.31
[1559606239] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10300180 kB) free.
[1559606308] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1559606695] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.2% (3756456 kB) free.
[1559606908] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 397 processes
[1559607028] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 398 processes
[1559607148] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1559607310] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.46, 0.51, 4.55
[1559607748] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 380 processes
[1559607868] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 386 processes
[1559607988] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 382 processes
[1559608495] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.6% (3185768 kB) free!
[1559608567] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1559608615] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (175904 kB) free!
[1559608735] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 27.3% (4432524 kB) free.
[1559609014] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (652300 kB) free!
[1559609128] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (764632 kB) free!
[1559609188] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 380 processes
[1559609216] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183096 kB) free!
[1559609239] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9622100 kB) free!
[1559609360] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7652440 kB) free!
[1559609480] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.1% (5964120 kB) free!
[1559609767] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559609788] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 411 processes
[1559609887] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559609908] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1559610008] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559610508] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 407 processes
[1559610629] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 403 processes
[1559610749] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 397 processes
[1559610847] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 397 processes
[1559610923] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 627 processes
[1559610967] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 395 processes
[1559610983] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 623 processes
[1559611280] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4634232 kB) free!
[1559611352] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559611437] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 53.77, 242.98, 157.96
[1559611528] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1559611528] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.4% (2788728 kB) free!
[1559611557] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 7.58, 162.69, 138.84
[1559611677] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.15, 108.91, 122.03
[1559613477] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1559613477] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.24, 0.56, 18.15
[1559614135] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.5% (2996696 kB) free!
[1559614255] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 27.4% (4439044 kB) free.
[1559614280] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (10009136 kB) free.
[1559615277] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.26, 0.27, 2.82
[1559615408] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559615480] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9430428 kB) free!
[1559615549] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 386 processes
[1559615601] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.1% (7420540 kB) free!
[1559615721] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.8% (5802044 kB) free!
[1559616749] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1559616869] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1559617208] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559617328] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559617448] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559617470] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 377 processes
[1559617528] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2417884 kB) free!
[1559617590] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 372 processes
[1559617855] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.7% (2063052 kB) free!
[1559617975] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.4% (3635292 kB) free.
[1559618121] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4312444 kB) free!
[1559618190] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 409 processes
[1559618310] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 415 processes
[1559618430] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 411 processes
[1559618575] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.3% (2806120 kB) free!
[1559618695] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.0% (2114052 kB) free!
[1559618815] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.6% (2688544 kB) free!
[1559619910] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 41.65, 81.70, 44.45
[1559620015] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (632152 kB) free!
[1559620030] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 5.87, 54.81, 39.12
[1559620150] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.73, 37.04, 34.51
[1559620230] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 430 processes
[1559620567] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 399 processes
[1559620615] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.9% (3073232 kB) free!
[1559620687] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 397 processes
[1559620750] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.50, 6.18, 18.72
[1559620830] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 442 processes
[1559620950] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 429 processes
[1559621070] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 426 processes
[1559621121] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.8% (9769072 kB) free!
[1559621123] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.46, 7.20, 5.93
[1559621243] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.92, 7.42, 6.17
[1559621287] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 5 crit, 0 warn out of 398 processes
[1559621363] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 9.14, 7.99, 6.52
[1559621405] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 14.0% (1027312 kB) free!
[1559621407] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 396 processes
[1559621525] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 10.7% (783916 kB) free!
[1559621527] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 397 processes
[1559621645] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (550540 kB) free!
[1559621671] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 436 processes
[1559621791] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 426 processes
[1559621911] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 439 processes
[1559622511] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 443 processes
[1559622550] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.83, 1.28, 4.43
[1559622631] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 435 processes
[1559622727] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 395 processes
[1559622751] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 468 processes
[1559622845] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.5% (6873492 kB) free.
[1559623016] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (189160 kB) free!
[1559623327] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 6 crit, 0 warn out of 400 processes
[1559623447] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 6 crit, 0 warn out of 400 processes
[1559623449] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1559623529] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1559623567] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 4 crit, 0 warn out of 397 processes
[1559623952] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 21.0% (3411628 kB) free.
[1559624216] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.9% (3063016 kB) free!
[1559624551] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 418 processes
[1559624552] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 12.2% (1976568 kB) free!
[1559624672] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 10.3% (1680324 kB) free!
[1559624792] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (1371540 kB) free!
[1559624816] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (811748 kB) free!
[1559625309] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46168 MB (20% inode=99%):
[1559625416] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.8% (3216920 kB) free!
[1559625429] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45319 MB (20% inode=99%):
[1559625477] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559625549] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45319 MB (20% inode=99%):
[1559625598] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.42, 4.93, 5.03
[1559625718] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.68, 4.85, 4.99
[1559625751] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 452 processes
[1559625871] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 448 processes
[1559625967] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1559625967] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 6 warn out of 402 processes
[1559625991] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 389 processes
[1559626318] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.47, 7.60, 6.27
[1559626438] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.75, 7.68, 6.47
[1559626558] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.10, 7.45, 6.53
[1559626567] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 7 crit, 0 warn out of 399 processes
[1559626616] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1557096 kB) free!
[1559627191] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 429 processes
[1559627791] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 409 processes
[1559627816] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.9% (3881284 kB) free.
[1559627911] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 404 processes
[1559628031] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 413 processes
[1559628249] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559628369] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559628489] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559628550] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 146.66, 296.02, 167.76
[1559628631] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 401 processes
[1559628670] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 20.23, 198.25, 147.49
[1559628752] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 403 processes
[1559628790] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1559628790] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.99, 132.76, 129.64
[1559628845] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (715816 kB) free!
[1559628872] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 415 processes
[1559628966] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (474776 kB) free!
[1559629016] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (151168 kB) free!
[1559629086] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (243012 kB) free!
[1559629137] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (173044 kB) free!
[1559629257] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (841176 kB) free!
[1559629472] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 404 processes
[1559629592] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 416 processes
[1559629712] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 411 processes
[1559629857] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.6% (2045320 kB) free!
[1559630122] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1559630167] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 7 warn out of 410 processes
[1559630202] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 62 users currently logged in
[1559630262] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 62 users currently logged in
[1559630286] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1559630286] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.3% (6786088 kB) free.
[1559630322] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 62 users currently logged in
[1559630382] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 62 users currently logged in
[1559630457] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1090932 kB) free!
[1559630590] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.80, 1.35, 19.68
[1559631367] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 7 warn out of 412 processes
[1559631657] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.2% (2468388 kB) free!
[1559631752] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 503.91, 284.44, 118.63
[1559631872] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 511.56, 359.98, 166.41
[1559631992] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 512.29, 410.44, 208.37
[1559632257] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174380 kB) free!
[1559632390] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.61, 0.68, 3.61
[1559633082] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 58 users currently logged in
[1559633121] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4417036 kB) free!
[1559633167] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 394 processes
[1559634057] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.4% (3632172 kB) free.
[1559634358] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.03, 1.16, 3.57
[1559634657] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (318764 kB) free!
[1559634777] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (155544 kB) free!
[1559634898] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (729840 kB) free!
[1559634992] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.17, 2.44, 13.29
[1559635498] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.3% (3780252 kB) free.
[1559635603] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 12.5% (2031944 kB) free!
[1559635686] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 19.6% (1438892 kB) free!
[1559635689] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1559635749] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24359 MB (10% inode=99%):
[1559635806] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 16.0% (1178072 kB) free!
[1559635926] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 12.3% (905784 kB) free!
[1559636098] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (148452 kB) free!
[1559636167] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 393 processes
[1559636218] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.8% (2564012 kB) free!
[1559636288] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 393 processes
[1559636338] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (192928 kB) free!
[1559636526] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (123096 kB) free!
[1559636792] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.46, 2.33, 3.88
[1559636803] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (163296 kB) free!
[1559636887] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 396 processes
[1559636938] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.8% (3046868 kB) free!
[1559637007] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 395 processes
[1559637127] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 397 processes
[1559637392] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 143.97, 53.04, 21.40
[1559637512] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 169.68, 92.93, 39.85
[1559637538] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (158784 kB) free!
[1559637632] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1559637632] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 173.74, 119.94, 56.19
[1559638327] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 390 processes
[1559638738] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.3% (3936324 kB) free.
[1559639277] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 114.81, 75.29, 33.49
[1559639338] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (149740 kB) free!
[1559639397] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 82.62, 81.08, 40.96
[1559639432] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.32, 2.91, 17.12
[1559639458] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 16.5% (2681328 kB) free!
[1559639517] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 11.26, 54.31, 36.01
[1559639578] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161648 kB) free!
[1559639728] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.6% (2597404 kB) free!
[1559639913] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559640117] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.20, 7.57, 19.03
[1559640127] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 1 warn out of 389 processes
[1559640248] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 390 processes
[1559640328] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2000028 kB) free!
[1559640368] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 388 processes
[1559640726] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.4% (6866572 kB) free.
[1559640778] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.7% (2392388 kB) free!
[1559641112] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 370 processes
[1559641232] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.67, 2.44, 4.59
[1559641378] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (177524 kB) free!
[1559641390] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 61.67, 38.51, 17.25
[1559641510] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 61.62, 46.15, 22.63
[1559641630] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15.10, 36.08, 21.99
[1559641712] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 388 processes
[1559641917] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.24, 0.35, 3.13
[1559642230] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.29, 5.62, 11.91
[1559642454] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.1% (3112064 kB) free!
[1559642574] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.4% (2983916 kB) free!
[1559642578] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.3% (2644364 kB) free!
[1559642695] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.3% (2811528 kB) free!
[1559642768] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 384 processes
[1559643178] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (152940 kB) free!
[1559643321] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.1% (4985824 kB) free!
[1559643368] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 384 processes
[1559643430] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.85, 1.65, 4.24
[1559643488] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 384 processes
[1559643512] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 372 processes
[1559643608] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 384 processes
[1559643895] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (1519548 kB) free!
[1559643921] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4872524 kB) free!
[1559644003] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 14.8% (2411404 kB) free!
[1559644112] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1559644112] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 384 processes
[1559644209] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 387 processes
[1559644329] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 391 processes
[1559644449] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 390 processes
[1559644482] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1559644542] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1559644592] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 39.9% (6480124 kB) free.
[1559644602] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1559644603] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1559644603] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (157832 kB) free!
[1559644662] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1559644749] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26316 MB (11% inode=99%):
[1559644962] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 56 users currently logged in
[1559645049] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 386 processes
[1559645722] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.1% (5490984 kB) free!
[1559645830] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.45, 16.36, 8.11
[1559645949] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 23570 MB (10% inode=99%):
[1559645950] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.00, 11.21, 7.23
[1559646070] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.65, 7.67, 6.41
[1559646249] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 388 processes
[1559646369] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 4 crit, 1 warn out of 390 processes
[1559646392] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (171656 kB) free!
[1559646489] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 387 processes
[1559646512] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 393 processes
[1559646512] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (164432 kB) free!
[1559646632] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (148000 kB) free!
[1559646670] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.59, 1.57, 3.64
[1559647112] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 377 processes
[1559647326] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 19.8% (1455864 kB) free!
[1559647446] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 17.2% (1267872 kB) free!
[1559647566] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 14.5% (1064388 kB) free!
[1559648166] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (271516 kB) free!
[1559648363] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.79, 3.63, 4.34
[1559648373] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1559648493] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1559648578] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.3% (1667136 kB) free!
[1559648613] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1559648912] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 372 processes
[1559649117] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 506.46, 326.04, 142.88
[1559649178] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155152 kB) free!
[1559649237] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 364.74, 353.86, 176.52
[1559649328] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.7% (2876652 kB) free!
[1559649357] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 49.54, 236.90, 155.16
[1559649513] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 383 processes
[1559650089] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 383 processes
[1559650362] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 62 users currently logged in
[1559650422] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 62 users currently logged in
[1559650482] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 62 users currently logged in
[1559650542] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 62 users currently logged in
[1559650689] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 393 processes
[1559650809] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 391 processes
[1559650978] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.4% (2817048 kB) free!
[1559651128] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2459268 kB) free!
[1559651313] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 402 processes
[1559651432] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559651695] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 82.1% (13340624 kB) free.
[1559651757] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.12, 0.29, 11.91
[1559651767] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.0% (6836928 kB) free.
[1559651913] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 390 processes
[1559652179] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (160424 kB) free!
[1559652957] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1559652957] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.13, 0.25, 3.46
[1559653210] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 396 processes
[1559653330] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 394 processes
[1559653414] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 40.5% (9931968 kB) free.
[1559653439] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.7% (7242608 kB) free!
[1559653560] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 12.1% (5966580 kB) free!
[1559653580] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 16.9% (4144544 kB) free!
[1559653584] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 672 processes
[1559653641] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 16.3% (3995292 kB) free!
[1559653644] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 669 processes
[1559653680] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 12.1% (5962660 kB) free!
[1559653701] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 16.3% (3991588 kB) free!
[1559653761] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 16.3% (3989712 kB) free!
[1559654142] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 60 users currently logged in
[1559654313] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 390 processes
[1559654914] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 383 processes
[1559655480] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (220804 kB) free!
[1559656004] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559656149] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 59298 MB (26% inode=99%):
[1559656522] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.1% (13368112 kB) free.
[1559656832] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 29.79, 13.33, 6.21
[1559656952] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.08, 9.74, 5.76
[1559657072] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.17, 7.51, 5.43
[1559657314] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 364 processes
[1559657672] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.84, 3.16, 3.96
[1559657914] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 365 processes
[1559658070] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559658103] SERVICE ALERT: sigqscluster1;HDD_Root status;CRITICAL;SOFT;1;DISK CRITICAL - free space: / 2257 MB (4% inode=95%):
[1559658179] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.5% (3817908 kB) free.
[1559658223] SERVICE ALERT: sigqscluster1;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 6765 MB (13% inode=98%):
[1559658343] SERVICE ALERT: sigqscluster1;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6765 MB (13% inode=98%):
[1559658730] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1559658779] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (151568 kB) free!
[1559658899] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (154932 kB) free!
[1559658943] SERVICE ALERT: sigqscluster1;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 3484 MB (6% inode=96%):
[1559659019] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (175112 kB) free!
[1559659522] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3557232 kB) free!
[1559659642] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.7% (12684376 kB) free.
[1559660167] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1559660314] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 382 processes
[1559660914] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 361 processes
[1559661489] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1559662019] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.2% (2303356 kB) free!
[1559662089] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559662210] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559662270] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 115.41, 56.87, 22.64
[1559662330] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559662390] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 119.49, 77.80, 34.46
[1559662510] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 86.31, 84.15, 42.33
[1559662616] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.5% (2826668 kB) free!
[1559662619] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (160856 kB) free!
[1559663014] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4903288 kB) free!
[1559663134] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (4866884 kB) free!
[1559663219] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.6% (3179668 kB) free!
[1559663254] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.8% (4851440 kB) free!
[1559663314] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 373 processes
[1559663710] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.38, 1.92, 11.91
[1559663757] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1559663819] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (164280 kB) free!
[1559663914] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 380 processes
[1559664416] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197360 kB) free!
[1559664454] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189832 kB) free!
[1559664910] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.53, 0.48, 3.61
[1559665642] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (4371868 kB) free!
[1559665762] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.7% (12668084 kB) free.
[1559666219] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.9% (2098512 kB) free!
[1559666314] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 383 processes
[1559666672] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 512.17, 410.11, 207.92
[1559666792] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 97.81, 298.60, 192.67
[1559666912] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14.63, 200.33, 169.50
[1559666914] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 386 processes
[1559668019] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (158480 kB) free!
[1559668714] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 359 processes
[1559669312] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.21, 1.91, 14.37
[1559669314] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 366 processes
[1559669989] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1559670109] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1559670229] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1559670910] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 14.04, 22.04, 10.28
[1559671031] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.21, 14.91, 9.10
[1559671112] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.48, 1.78, 3.62
[1559671114] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 393 processes
[1559671151] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1559671151] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.65, 10.14, 8.06
[1559671751] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.20, 1.64, 4.38
[1559672314] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 365 processes
[1559672820] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.3% (2000464 kB) free!