[1550255400] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11796916 kB) free.
[1550255920] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550255931] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46761 MB (20% inode=99%):
[1550256000] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8520696 kB) free!
[1550256051] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50857 MB (22% inode=99%):
[1550256120] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11865552 kB) free.
[1550256255] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550256316] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1550256436] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1550256651] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46761 MB (20% inode=99%):
[1550256720] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8720988 kB) free!
[1550256771] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50857 MB (22% inode=99%):
[1550256840] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11942972 kB) free.
[1550257371] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46761 MB (20% inode=99%):
[1550257440] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8951332 kB) free!
[1550257491] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50857 MB (22% inode=99%):
[1550257561] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12000212 kB) free.
[1550259175] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43658 MB (20% inode=99%):
[1550259295] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41610 MB (19% inode=99%):
[1550259360] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9316828 kB) free!
[1550259415] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 45453 MB (21% inode=99%):
[1550259480] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12389040 kB) free.
[1550259891] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46760 MB (20% inode=99%):
[1550260011] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50342 MB (22% inode=99%):
[1550260611] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46760 MB (20% inode=99%):
[1550260680] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7795480 kB) free!
[1550260731] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50856 MB (22% inode=99%):
[1550260795] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550260800] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12382440 kB) free.
[1550261815] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43708 MB (20% inode=99%):
[1550261935] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 37552 MB (18% inode=99%):
[1550261995] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550262055] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 45256 MB (21% inode=99%):
[1550262116] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550262236] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550262836] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550262956] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550264455] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43561 MB (20% inode=99%):
[1550264575] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39473 MB (18% inode=99%):
[1550264695] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46645 MB (22% inode=99%):
[1550264803] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 354 processes
[1550264924] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 354 processes
[1550265044] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 360 processes
[1550265895] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43568 MB (20% inode=99%):
[1550266016] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47670 MB (22% inode=99%):
[1550266800] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9781072 kB) free!
[1550266920] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12355044 kB) free.
[1550269615] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43312 MB (20% inode=99%):
[1550269735] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47417 MB (22% inode=99%):
[1550270520] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7729196 kB) free!
[1550270640] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12342480 kB) free.
[1550271535] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43355 MB (20% inode=99%):
[1550271655] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47454 MB (22% inode=99%):
[1550274044] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1550274055] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43455 MB (20% inode=99%):
[1550274175] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47561 MB (22% inode=99%):
[1550276156] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550276276] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550276331] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46878 MB (20% inode=99%):
[1550276396] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550276451] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50974 MB (22% inode=99%):
[1550277051] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46869 MB (20% inode=99%):
[1550277171] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50932 MB (22% inode=99%):
[1550277175] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43421 MB (20% inode=99%):
[1550277296] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47506 MB (22% inode=99%):
[1550277771] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46832 MB (20% inode=99%):
[1550277840] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8475792 kB) free!
[1550277891] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50928 MB (22% inode=99%):
[1550277960] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11815716 kB) free.
[1550278491] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46832 MB (20% inode=99%):
[1550278560] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8847720 kB) free!
[1550278611] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50928 MB (22% inode=99%):
[1550278680] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12030132 kB) free.
[1550279211] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46832 MB (20% inode=99%):
[1550279280] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8944520 kB) free!
[1550279331] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46446 MB (20% inode=99%):
[1550279397] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550279400] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12270324 kB) free.
[1550279451] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49901 MB (22% inode=99%):
[1550279695] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42858 MB (20% inode=99%):
[1550279815] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46452 MB (22% inode=99%):
[1550279997] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550280000] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8857264 kB) free!
[1550280051] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46447 MB (20% inode=99%):
[1550280117] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550280120] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12068384 kB) free.
[1550280171] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49901 MB (22% inode=99%):
[1550280237] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550280720] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8975188 kB) free!
[1550280771] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46702 MB (20% inode=99%):
[1550280837] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1550280837] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550280841] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12252564 kB) free.
[1550280891] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49901 MB (22% inode=99%):
[1550281015] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42886 MB (20% inode=99%):
[1550281135] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46985 MB (22% inode=99%):
[1550281437] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550281440] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9188664 kB) free!
[1550281491] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46826 MB (20% inode=99%):
[1550281557] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550281560] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12243868 kB) free.
[1550281611] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50408 MB (22% inode=99%):
[1550282157] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550282160] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9305644 kB) free!
[1550282211] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46761 MB (20% inode=99%):
[1550282277] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550282280] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12239420 kB) free.
[1550282331] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50344 MB (22% inode=99%):
[1550282878] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550282880] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9459228 kB) free!
[1550282932] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46760 MB (20% inode=99%):
[1550282998] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550283000] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12231296 kB) free.
[1550283052] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50600 MB (22% inode=99%):
[1550283535] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42665 MB (20% inode=99%):
[1550283598] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550283600] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9566212 kB) free!
[1550283652] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46759 MB (20% inode=99%):
[1550283655] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46761 MB (22% inode=99%):
[1550283718] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1550283720] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12222132 kB) free.
[1550283772] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50599 MB (22% inode=99%):
[1550284318] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550284320] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9542116 kB) free!
[1550284372] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46751 MB (20% inode=99%):
[1550284438] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1550284440] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12200156 kB) free.
[1550284492] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50847 MB (22% inode=99%):
[1550285038] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550285040] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9798476 kB) free!
[1550285092] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46749 MB (20% inode=99%):
[1550285158] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550285161] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7553508 kB) free!
[1550285212] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50845 MB (22% inode=99%):
[1550285279] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550285281] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11371792 kB) free.
[1550285455] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42647 MB (20% inode=99%):
[1550285575] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46740 MB (22% inode=99%):
[1550285812] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46748 MB (20% inode=99%):
[1550285881] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7735536 kB) free!
[1550285932] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50844 MB (22% inode=99%):
[1550286001] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11469276 kB) free.
[1550286532] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46747 MB (20% inode=99%):
[1550286601] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8101804 kB) free!
[1550286652] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50842 MB (22% inode=99%):
[1550286721] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11617056 kB) free.
[1550287252] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46745 MB (20% inode=99%):
[1550287321] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8475464 kB) free!
[1550287372] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50841 MB (22% inode=99%):
[1550287441] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11679844 kB) free.
[1550287972] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46565 MB (20% inode=99%):
[1550288041] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8671500 kB) free!
[1550288092] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46344 MB (20% inode=99%):
[1550288161] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12091840 kB) free.
[1550288212] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49664 MB (22% inode=99%):
[1550288761] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8844508 kB) free!
[1550288812] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46591 MB (20% inode=99%):
[1550288879] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550288881] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12070684 kB) free.
[1550288932] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50175 MB (22% inode=99%):
[1550289176] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43433 MB (20% inode=99%):
[1550289296] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41389 MB (19% inode=99%):
[1550289416] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44455 MB (21% inode=99%):
[1550289479] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550289482] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9055572 kB) free!
[1550289532] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46524 MB (20% inode=99%):
[1550289599] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550289602] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12032948 kB) free.
[1550289652] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50105 MB (22% inode=99%):
[1550290200] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550290202] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9272328 kB) free!
[1550290253] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46580 MB (20% inode=99%):
[1550290320] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550290322] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12015912 kB) free.
[1550290373] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50675 MB (22% inode=99%):
[1550290920] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550290922] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9495840 kB) free!
[1550290973] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46578 MB (20% inode=99%):
[1550291040] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550291042] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.6% (7205996 kB) free!
[1550291093] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50674 MB (22% inode=99%):
[1550291160] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550291162] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11114892 kB) free.
[1550291693] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46561 MB (20% inode=99%):
[1550291762] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8448792 kB) free!
[1550291813] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46185 MB (20% inode=99%):
[1550291882] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11873460 kB) free.
[1550291933] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49637 MB (22% inode=99%):
[1550292482] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8793828 kB) free!
[1550292533] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46501 MB (20% inode=99%):
[1550292602] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11879968 kB) free.
[1550292653] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50080 MB (22% inode=99%):
[1550292960] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1550293203] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8888548 kB) free!
[1550293253] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46330 MB (20% inode=99%):
[1550293323] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11835036 kB) free.
[1550293373] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50367 MB (22% inode=99%):
[1550293561] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550293681] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1550293923] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9068768 kB) free!
[1550293973] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46477 MB (20% inode=99%):
[1550294043] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11784860 kB) free.
[1550294093] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50572 MB (22% inode=99%):
[1550294281] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550294401] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550294643] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9171720 kB) free!
[1550294693] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46473 MB (20% inode=99%):
[1550294763] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.0% (6901368 kB) free!
[1550294813] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50568 MB (22% inode=99%):
[1550294883] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10860868 kB) free.
[1550295001] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1550295121] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550295241] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550295413] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46347 MB (20% inode=99%):
[1550295455] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.3% (2813096 kB) free!
[1550295483] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7100920 kB) free!
[1550295533] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50440 MB (22% inode=99%):
[1550295575] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.7% (2714832 kB) free!
[1550295603] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10895204 kB) free.
[1550295695] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.8% (2567224 kB) free!
[1550295960] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24569 MB (10% inode=99%):
[1550296016] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40387 MB (19% inode=99%):
[1550296133] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46333 MB (20% inode=99%):
[1550296136] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48008 MB (23% inode=99%):
[1550296253] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50299 MB (22% inode=99%):
[1550296803] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8925396 kB) free!
[1550296853] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46184 MB (20% inode=99%):
[1550296895] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1181300 kB) free!
[1550296923] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.9% (8812284 kB) free!
[1550296973] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50400 MB (22% inode=99%):
[1550297044] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10570532 kB) free.
[1550297573] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46029 MB (20% inode=99%):
[1550297641] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550297693] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50107 MB (22% inode=99%):
[1550297937] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.1% (2779692 kB) free!
[1550298057] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.0% (2596324 kB) free!
[1550298177] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.4% (2346120 kB) free!
[1550298241] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550298243] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8583208 kB) free!
[1550298293] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45763 MB (20% inode=99%):
[1550298361] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1550298363] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11316424 kB) free.
[1550298413] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49857 MB (22% inode=99%):
[1550298557] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 8 users currently logged in
[1550298657] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9325116 kB) free!
[1550298677] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1550298777] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.0% (13803488 kB) free.
[1550298797] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1550298961] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550298963] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8703864 kB) free!
[1550299013] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45735 MB (20% inode=99%):
[1550299081] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550299083] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.8% (6302764 kB) free!
[1550299133] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50079 MB (22% inode=99%):
[1550299201] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550299203] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10225784 kB) free.
[1550299455] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550299575] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550299695] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550299733] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45705 MB (20% inode=99%):
[1550299803] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.5% (6148636 kB) free!
[1550299853] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49796 MB (22% inode=99%):
[1550299923] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10106296 kB) free.
[1550299977] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (817896 kB) free!
[1550300453] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45673 MB (20% inode=99%):
[1550300523] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6298412 kB) free!
[1550300573] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49749 MB (22% inode=99%):
[1550300643] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10194624 kB) free.
[1550301173] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45361 MB (20% inode=99%):
[1550301177] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3585904 kB) free!
[1550301243] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6113228 kB) free!
[1550301293] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49436 MB (21% inode=99%):
[1550301297] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9207016 kB) free!
[1550301364] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10082236 kB) free.
[1550301417] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.9% (5846664 kB) free!
[1550301893] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45263 MB (20% inode=99%):
[1550301960] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 25228 MB (11% inode=99%):
[1550301963] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6650048 kB) free!
[1550302013] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49356 MB (21% inode=99%):
[1550302083] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10192356 kB) free.
[1550302136] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 36452 MB (17% inode=99%):
[1550302256] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44350 MB (21% inode=99%):
[1550302377] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 12.8% (2087700 kB) free!
[1550302613] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45235 MB (20% inode=99%):
[1550302683] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6855484 kB) free!
[1550302733] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45034 MB (20% inode=99%):
[1550302801] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1550302803] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10541464 kB) free.
[1550302853] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48231 MB (21% inode=99%):
[1550302977] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1550302977] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1521232 kB) free!
[1550303160] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24636 MB (10% inode=99%):
[1550303401] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550303403] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6908872 kB) free!
[1550303453] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45127 MB (20% inode=99%):
[1550303456] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 36554 MB (17% inode=99%):
[1550303521] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1550303523] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10444748 kB) free.
[1550303573] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48061 MB (21% inode=99%):
[1550303576] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43693 MB (20% inode=99%):
[1550303696] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 31909 MB (15% inode=99%):
[1550304121] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550304123] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7655848 kB) free!
[1550304173] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45985 MB (20% inode=99%):
[1550304241] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1550304243] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11402208 kB) free.
[1550304293] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49422 MB (21% inode=99%):
[1550304417] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (225384 kB) free!
[1550304777] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 70.1% (11393016 kB) free.
[1550304841] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550304843] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7748388 kB) free!
[1550304893] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46034 MB (20% inode=99%):
[1550304962] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1550304963] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11263792 kB) free.
[1550305014] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49616 MB (22% inode=99%):
[1550305561] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550305564] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7622364 kB) free!
[1550305613] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45939 MB (20% inode=99%):
[1550305681] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550305684] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11142728 kB) free.
[1550305733] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49509 MB (22% inode=99%):
[1550306096] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46388 MB (22% inode=99%):
[1550306217] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.0% (8392692 kB) free!
[1550306281] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550306284] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7670908 kB) free!
[1550306333] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46148 MB (20% inode=99%):
[1550306401] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550306404] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11016140 kB) free.
[1550306453] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49472 MB (22% inode=99%):
[1550306817] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1550306817] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.1% (9916464 kB) free.
[1550307001] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550307004] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7776648 kB) free!
[1550307053] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45736 MB (20% inode=99%):
[1550307121] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1550307124] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10918800 kB) free.
[1550307173] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50063 MB (22% inode=99%):
[1550307296] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 38149 MB (18% inode=99%):
[1550307416] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46561 MB (22% inode=99%):
[1550307721] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550307724] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7889500 kB) free!
[1550307773] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45948 MB (20% inode=99%):
[1550307841] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550307844] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10774372 kB) free.
[1550307893] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50042 MB (22% inode=99%):
[1550307962] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550308377] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.5% (3008620 kB) free!
[1550308444] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8136620 kB) free!
[1550308493] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45870 MB (20% inode=99%):
[1550308497] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.4% (2833200 kB) free!
[1550308564] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.9% (5874196 kB) free!
[1550308613] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49965 MB (22% inode=99%):
[1550308616] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 38516 MB (18% inode=99%):
[1550308617] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4277764 kB) free!
[1550308617] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.3% (2649724 kB) free!
[1550308684] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 20.0% (9870752 kB) free!
[1550308736] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 46684 MB (22% inode=99%):
[1550308737] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.4% (10550360 kB) free.
[1550309213] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45401 MB (20% inode=99%):
[1550309277] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4777976 kB) free!
[1550309333] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49487 MB (22% inode=99%):
[1550309336] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 35884 MB (17% inode=99%):
[1550309337] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (2584572 kB) free!
[1550309397] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4652132 kB) free!
[1550309456] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44027 MB (21% inode=99%):
[1550309457] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9049696 kB) free!
[1550309517] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.5% (4562176 kB) free!
[1550309577] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.4% (5602292 kB) free!
[1550309762] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1550309933] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45378 MB (20% inode=99%):
[1550310053] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45378 MB (20% inode=99%):
[1550310056] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 38394 MB (18% inode=99%):
[1550310173] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48954 MB (21% inode=99%):
[1550310176] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42844 MB (20% inode=99%):
[1550310296] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 36688 MB (17% inode=99%):
[1550310362] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550310436] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 323 processes
[1550310482] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550310484] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10504828 kB) free.
[1550310557] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 322 processes
[1550310773] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45359 MB (20% inode=99%):
[1550310893] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48939 MB (21% inode=99%):
[1550311017] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (984352 kB) free!
[1550311083] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550311084] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7231036 kB) free!
[1550311120] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1550311203] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550311204] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.1% (10410740 kB) free.
[1550311240] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1550311323] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550311360] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1550311377] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (1860008 kB) free!
[1550311493] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45339 MB (20% inode=99%):
[1550311496] SERVICE FLAPPING ALERT: db1;HDD_Home status;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1550311496] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44536 MB (21% inode=99%):
[1550311613] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49431 MB (21% inode=99%):
[1550311804] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7572220 kB) free!
[1550311924] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.0% (5439140 kB) free!
[1550311977] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.1% (4962016 kB) free!
[1550312044] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.2% (9459400 kB) free!
[1550312096] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 35835 MB (17% inode=99%):
[1550312213] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44936 MB (19% inode=99%):
[1550312216] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 44378 MB (21% inode=99%):
[1550312334] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49150 MB (21% inode=99%):
[1550312554] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 51 users currently logged in
[1550312614] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 47 users currently logged in
[1550312816] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 35774 MB (17% inode=99%):
[1550312914] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 51 users currently logged in
[1550312933] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47092 MB (20% inode=99%):
[1550312936] SERVICE ALERT: db1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42924 MB (20% inode=99%):
[1550312974] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 51 users currently logged in
[1550313034] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 51 users currently logged in
[1550313053] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44787 MB (19% inode=99%):
[1550313056] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 39342 MB (18% inode=99%):
[1550313072] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 594 processes
[1550313094] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 53 users currently logged in
[1550313117] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (826572 kB) free!
[1550313123] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1550313132] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 602 processes
[1550313173] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48616 MB (21% inode=99%):
[1550313177] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (238644 kB) free!
[1550313417] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 87.4% (14202740 kB) free.
[1550313556] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1550313676] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1550313723] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1550313773] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45023 MB (20% inode=99%):
[1550313843] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1550313844] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1550313844] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4693240 kB) free!
[1550313893] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49099 MB (21% inode=99%):
[1550313960] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 31628 MB (14% inode=99%):
[1550313963] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1550314377] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.3% (6539300 kB) free!
[1550314444] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.4% (7606404 kB) free!
[1550314493] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45002 MB (20% inode=99%):
[1550314613] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44107 MB (19% inode=99%):
[1550314733] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48038 MB (21% inode=99%):
[1550314762] SERVICE ALERT: sig1;Memory;WARNING;SOFT;1;WARNING - 19.6% (3190704 kB) free!
[1550314882] SERVICE ALERT: sig1;Memory;WARNING;SOFT;2;WARNING - 19.4% (3160680 kB) free!
[1550315002] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 19.2% (3122744 kB) free!
[1550315045] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9971408 kB) free.
[1550315333] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44982 MB (20% inode=99%):
[1550315453] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48562 MB (21% inode=99%):
[1550315577] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (2185584 kB) free!
[1550315645] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6694568 kB) free!
[1550315763] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1550315765] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9896192 kB) free.
[1550315794] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 44 users currently logged in
[1550315832] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 569 processes
[1550315892] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 570 processes
[1550315962] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 54.6% (26908388 kB) free.
[1550316053] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44969 MB (20% inode=99%):
[1550316093] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.5% (22823804 kB) free.
[1550316160] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550316173] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49065 MB (21% inode=99%):
[1550316195] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 93.8% (15242200 kB) free.
[1550316197] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550316336] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.9% (15106788 kB) free.
[1550316363] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550316365] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7199552 kB) free!
[1550316483] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1550316485] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.3% (5077216 kB) free!
[1550316527] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 47.6% (23448192 kB) free.
[1550316603] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1550316773] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44894 MB (19% inode=99%):
[1550316893] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49046 MB (21% inode=99%):
[1550317493] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44940 MB (19% inode=99%):
[1550317613] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44940 MB (19% inode=99%):
[1550317733] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48516 MB (21% inode=99%):
[1550318333] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44564 MB (19% inode=99%):
[1550318404] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1550318453] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48661 MB (21% inode=99%):
[1550318895] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550319004] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1550319054] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44557 MB (19% inode=99%):
[1550319124] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550319174] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48652 MB (21% inode=99%):
[1550319244] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550319774] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44645 MB (19% inode=99%):
[1550319894] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44516 MB (19% inode=99%):
[1550320014] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48228 MB (21% inode=99%):
[1550320095] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 8 users currently logged in
[1550320109] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550320215] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1550320229] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550320335] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1550320349] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550320360] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1550320397] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 8 users currently logged in
[1550320480] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1550320517] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1550320600] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1550320614] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44130 MB (19% inode=99%):
[1550320638] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1550320719] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184892 kB) free!
[1550320734] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48095 MB (21% inode=99%):
[1550320777] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (1383480 kB) free!
[1550320839] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.5% (3317160 kB) free!
[1550320897] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (235264 kB) free!
[1550320959] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185412 kB) free!
[1550321018] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (218116 kB) free!
[1550321044] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550321334] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44620 MB (19% inode=99%):
[1550321454] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48698 MB (21% inode=99%):
[1550321549] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550321644] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550321764] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550321884] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550322054] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44601 MB (19% inode=99%):
[1550322174] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48694 MB (21% inode=99%):
[1550322774] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44589 MB (19% inode=99%):
[1550322894] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44589 MB (19% inode=99%):
[1550323014] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48168 MB (21% inode=99%):
[1550323017] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1550323419] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.7% (9224096 kB) free!
[1550323614] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44582 MB (19% inode=99%):
[1550323684] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550323734] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48645 MB (21% inode=99%):
[1550323856] SERVICE FLAPPING ALERT: db1;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1550323894] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 52 users currently logged in
[1550323954] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 52 users currently logged in
[1550324014] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 52 users currently logged in
[1550324019] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4504068 kB) free!
[1550324074] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 52 users currently logged in
[1550324284] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550324334] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44539 MB (19% inode=99%):
[1550324404] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550324454] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48634 MB (21% inode=99%):
[1550324456] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 71576 MB (34% inode=99%):
[1550324524] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550324619] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 35.3% (17408240 kB) free.
[1550325054] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44537 MB (19% inode=99%):
[1550325174] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48632 MB (21% inode=99%):
[1550325644] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 379 processes
[1550325764] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 377 processes
[1550325774] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43822 MB (19% inode=99%):
[1550325884] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 376 processes
[1550325895] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43316 MB (19% inode=99%):
[1550326015] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47412 MB (21% inode=99%):
[1550326136] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1550326174] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 50 users currently logged in
[1550326325] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1550326419] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.2% (8463208 kB) free!
[1550326474] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 54 users currently logged in
[1550326534] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 54 users currently logged in
[1550326539] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 33.4% (16476724 kB) free.
[1550326594] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 54 users currently logged in
[1550326615] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43828 MB (19% inode=99%):
[1550326654] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 54 users currently logged in
[1550326735] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47924 MB (21% inode=99%):
[1550326925] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550327045] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550327139] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9063032 kB) free!
[1550327165] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1550327259] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 34.4% (16964444 kB) free.
[1550327335] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43827 MB (19% inode=99%):
[1550327455] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47923 MB (21% inode=99%):
[1550327859] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9467004 kB) free!
[1550327979] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 34.1% (16806776 kB) free.
[1550328055] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43809 MB (19% inode=99%):
[1550328175] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43689 MB (19% inode=99%):
[1550328295] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46886 MB (20% inode=99%):
[1550328579] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9848072 kB) free!
[1550328700] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 33.8% (16676484 kB) free.
[1550329495] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47794 MB (21% inode=99%):
[1550330095] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44723 MB (19% inode=99%):
[1550330166] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1550330215] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48757 MB (21% inode=99%):
[1550330499] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.6% (7670460 kB) free!
[1550330501] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.84, 10.02, 6.26
[1550330619] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 32.2% (15887796 kB) free.
[1550330621] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.40, 7.05, 5.62
[1550330684] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1550330741] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.10, 5.07, 5.07
[1550330766] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1550330815] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44637 MB (19% inode=99%):
[1550330886] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550330935] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48734 MB (21% inode=99%):
[1550331006] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1550331219] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.9% (7826684 kB) free!
[1550331339] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 32.7% (16121872 kB) free.
[1550331341] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.40, 1.79, 3.26
[1550331535] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44628 MB (19% inode=99%):
[1550331655] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43602 MB (19% inode=99%):
[1550331775] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47699 MB (21% inode=99%):
[1550331939] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.3% (8544232 kB) free!
[1550332059] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 32.9% (16191424 kB) free.
[1550332375] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44612 MB (19% inode=99%):
[1550332495] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48193 MB (21% inode=99%):
[1550332659] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8775140 kB) free!
[1550332779] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 32.6% (16088056 kB) free.
[1550332806] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1550333095] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44609 MB (19% inode=99%):
[1550333216] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48193 MB (21% inode=99%):
[1550333379] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9492480 kB) free!
[1550333406] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1550333499] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 32.5% (16012840 kB) free.
[1550333527] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1550333800] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550333815] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44608 MB (19% inode=99%):
[1550333838] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550333854] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 36 users currently logged in
[1550333935] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48705 MB (21% inode=99%):
[1550334100] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9713616 kB) free!
[1550334126] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550334135] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550334220] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 32.3% (15915056 kB) free.
[1550334246] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550334366] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1550334535] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44606 MB (19% inode=99%):
[1550334655] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48704 MB (21% inode=99%):
[1550335255] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44566 MB (19% inode=99%):
[1550335375] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44575 MB (19% inode=99%):
[1550335420] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.5% (5657968 kB) free!
[1550335495] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48181 MB (21% inode=99%):
[1550335540] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.7% (14144952 kB) free.
[1550335935] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550336055] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550336095] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44596 MB (19% inode=99%):
[1550336140] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.2% (7962176 kB) free!
[1550336166] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1550336175] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550336215] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48686 MB (21% inode=99%):
[1550336260] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 31.9% (15717872 kB) free.
[1550336766] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550336815] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44595 MB (19% inode=99%):
[1550336860] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.1% (7910420 kB) free!
[1550336887] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550336935] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48693 MB (21% inode=99%):
[1550336980] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 31.8% (15683508 kB) free.
[1550337007] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550337535] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44594 MB (19% inode=99%):
[1550337580] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.7% (7740856 kB) free!
[1550337655] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48692 MB (21% inode=99%):
[1550337700] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1550337700] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 31.8% (15659116 kB) free.
[1550338255] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44529 MB (19% inode=99%):
[1550338300] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.5% (8135752 kB) free!
[1550338375] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44531 MB (19% inode=99%):
[1550338420] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 31.7% (15610968 kB) free.
[1550338495] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48115 MB (21% inode=99%):
[1550338807] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550338959] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.3% (3740844 kB) free!
[1550339020] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9139052 kB) free!
[1550339095] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44463 MB (19% inode=99%):
[1550339141] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 31.6% (15570836 kB) free.
[1550339215] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48561 MB (21% inode=99%):
[1550339407] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550339527] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550339647] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550339740] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.8% (8271016 kB) free!
[1550339815] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44462 MB (19% inode=99%):
[1550339860] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 31.5% (15527140 kB) free.
[1550339935] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48560 MB (21% inode=99%):
[1550340400] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1550340438] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1550340460] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8059592 kB) free!
[1550340520] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1550340535] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46512 MB (20% inode=99%):
[1550340558] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1550340580] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 31.4% (15497768 kB) free.
[1550340640] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1550340656] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44462 MB (19% inode=99%):
[1550340678] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1550340776] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48048 MB (21% inode=99%):
[1550341180] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.8% (8266696 kB) free!
[1550341300] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 31.4% (15467904 kB) free.
[1550341376] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44461 MB (19% inode=99%):
[1550341496] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48358 MB (21% inode=99%):
[1550341544] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.6% (3185088 kB) free!
[1550341664] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.5% (3166760 kB) free!
[1550341784] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.4% (3153192 kB) free!