[1587839696] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (3203176 kB) free!
[1587840124] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.8% (4032456 kB) free.
[1587840896] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21224540 kB) free!
[1587841496] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30705840 kB) free.
[1587841511] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174624 kB) free!
[1587841571] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195136 kB) free!
[1587841631] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (184196 kB) free!
[1587841691] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (173264 kB) free!
[1587841924] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (155972 kB) free!
[1587841991] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 73.1% (17909364 kB) free.
[1587842044] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (1276304 kB) free!
[1587842096] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9244736 kB) free!
[1587842164] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1109644 kB) free!
[1587842216] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (5741764 kB) free!
[1587842336] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30726148 kB) free.
[1587842936] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (9989972 kB) free!
[1587843056] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6390852 kB) free!
[1587843176] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30728512 kB) free.
[1587843776] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10659464 kB) free!
[1587843896] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (7099516 kB) free!
[1587844013] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1587844016] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30690568 kB) free.
[1587844613] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (618880 kB) free!
[1587844617] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11347936 kB) free!
[1587844733] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (2519552 kB) free!
[1587844737] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7791512 kB) free!
[1587844853] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (2662648 kB) free!
[1587844857] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30702060 kB) free.
[1587844931] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (231696 kB) free!
[1587845051] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (230512 kB) free!
[1587845171] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (230816 kB) free!
[1587845457] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11928440 kB) free!
[1587845577] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8319240 kB) free!
[1587845697] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (4840712 kB) free!
[1587846897] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19653256 kB) free!
[1587846972] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 13.4% (6620012 kB) free!
[1587847064] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.32, 5.53, 5.05
[1587847184] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.68, 5.89, 5.23
[1587847304] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.19, 5.97, 5.34
[1587847497] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28828188 kB) free.
[1587848097] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8751788 kB) free!
[1587848217] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (5321592 kB) free!
[1587848337] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28334912 kB) free.
[1587848504] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.92, 4.40, 4.74
[1587848523] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 77.3% (38091420 kB) free.
[1587848937] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8254900 kB) free!
[1587848938] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.0% (43864624 kB) free.
[1587849057] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30605864 kB) free.
[1587849657] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10163672 kB) free!
[1587849728] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 39.3% (9641216 kB) free.
[1587849777] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6394548 kB) free!
[1587849898] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28509380 kB) free.
[1587850207] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1087) < 182.18.144.246 [26/Apr/2020:03:00:07 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1587850327] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1114) < 49.50.86.152 [26/Apr/2020:03:01:47 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1587850394] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(378) < 103.16.101.72 [26/Apr/2020:03:01:32 +0530] TCP 502 0 0 60.001 "10.147.212.38:4000, backend3" "0, 0" "0, 0" "60.001, 0.000"
[1587850447] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587850462] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.2% (2740472 kB) free!
[1587850497] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8482160 kB) free!
[1587850515] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587850573] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.4% (44531728 kB) free.
[1587850617] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (13121808 kB) free!
[1587850640] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(500) < 2020/04/26 03:01:32 [error] 22131#0: *85772511 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1587850737] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27234984 kB) free.
[1587850761] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587850929] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (298616 kB) free!
[1587851049] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (299212 kB) free!
[1587851169] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (294752 kB) free!
[1587851337] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6161964 kB) free!
[1587851457] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30368728 kB) free.
[1587851764] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.1% (1640912 kB) free!
[1587852057] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8117948 kB) free!
[1587852104] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.30, 5.48, 5.03
[1587852177] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4402464 kB) free!
[1587852224] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.76, 5.55, 5.12
[1587852297] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28957472 kB) free.
[1587852344] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.37, 5.18, 5.04
[1587852897] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6869584 kB) free!
[1587852944] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1587852944] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.64, 3.64, 4.26
[1587853017] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30294704 kB) free.
[1587853564] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3669908 kB) free.
[1587853569] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.9% (10262664 kB) free.
[1587853617] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9194020 kB) free!
[1587853737] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (5566524 kB) free!
[1587853857] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27850612 kB) free.
[1587854062] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1620884 kB) free!
[1587854169] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (299584 kB) free!
[1587854289] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198044 kB) free!
[1587854409] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1587854409] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (306116 kB) free!
[1587854457] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7492836 kB) free!
[1587854577] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30382316 kB) free.
[1587855178] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9324180 kB) free!
[1587855298] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30453208 kB) free.
[1587855609] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 38.6% (9457180 kB) free.
[1587855898] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10064412 kB) free!
[1587856018] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30454388 kB) free.
[1587856618] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5210848 kB) free!
[1587856738] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30337084 kB) free.
[1587857338] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8449588 kB) free!
[1587857458] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30446416 kB) free.
[1587858009] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (151992 kB) free!
[1587858058] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (12992060 kB) free!
[1587858129] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.1% (3209112 kB) free!
[1587858178] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31545052 kB) free.
[1587858249] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194636 kB) free!
[1587858262] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 76.3% (18694264 kB) free.
[1587858597] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 56.8% (13913020 kB) free.
[1587858778] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14659608 kB) free!
[1587858898] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32628384 kB) free.
[1587859071] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 31.1% (7671708 kB) free.
[1587859498] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15263176 kB) free!
[1587859618] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31327268 kB) free.
[1587860819] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13714844 kB) free!
[1587860939] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31346148 kB) free.
[1587861539] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1587862139] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14070096 kB) free!
[1587862259] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31575448 kB) free.
[1587862671] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (1331356 kB) free!
[1587862791] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167772 kB) free!
[1587862853] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 34.1% (45014556 kB) free.
[1587862911] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195316 kB) free!
[1587863459] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15677660 kB) free!
[1587863579] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31869760 kB) free.
[1587863764] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1587864344] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1587864779] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15397432 kB) free!
[1587864899] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31928808 kB) free.
[1587865499] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14635992 kB) free!
[1587865619] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31127164 kB) free.
[1587866819] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15688560 kB) free!
[1587866939] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31774040 kB) free.
[1587867540] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14130616 kB) free!
[1587867660] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30611216 kB) free.
[1587867849] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1587868860] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13334620 kB) free!
[1587868980] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30410348 kB) free.
[1587870180] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11702692 kB) free!
[1587870300] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29839044 kB) free.
[1587871500] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11303520 kB) free!
[1587871620] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29292548 kB) free.
[1587872144] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.36, 7.66, 5.89
[1587872264] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.97, 7.88, 6.18
[1587872384] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.62, 6.66, 5.94
[1587872820] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9261108 kB) free!
[1587872910] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1587872940] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28506436 kB) free.
[1587873030] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 350 processes
[1587873150] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1587873540] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7127928 kB) free!
[1587873584] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.48, 4.46, 4.94
[1587873661] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26773824 kB) free.
[1587874184] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.91, 5.57, 5.31
[1587874305] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.69, 5.18, 5.19
[1587874425] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.43, 5.77, 5.41
[1587874860] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6987452 kB) free!
[1587874980] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26517232 kB) free.
[1587875580] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4223116 kB) free!
[1587875700] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.3% (25476344 kB) free!
[1587875820] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20450060 kB) free!
[1587876150] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 348 processes
[1587876262] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192704 kB) free!
[1587876382] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193476 kB) free!
[1587876502] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193484 kB) free!
[1587876523] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1587876643] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1587876750] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1587877020] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (438436 kB) free!
[1587877197] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (4552328 kB) free!
[1587877317] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (1365164 kB) free!
[1587877437] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (1256368 kB) free!
[1587879364] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3257024 kB) free!
[1587879420] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.5% (21796916 kB) free!
[1587879484] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (3223684 kB) free!
[1587879604] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (3191564 kB) free!
[1587879697] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 63.3% (15526004 kB) free.
[1587879750] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1587879750] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 350 processes
[1587880020] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1587880020] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31335968 kB) free.
[1587880620] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7462140 kB) free!
[1587880741] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30931308 kB) free.
[1587880960] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.9% (3118880 kB) free!
[1587881080] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.3% (2750776 kB) free!
[1587881200] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.6% (2374552 kB) free!
[1587881340] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6435928 kB) free!
[1587881405] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (1014460 kB) free!
[1587881460] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28963420 kB) free.
[1587881801] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (201440 kB) free!
[1587882005] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.7% (2222508 kB) free!
[1587882060] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (2600080 kB) free!
[1587882180] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (385760 kB) free!
[1587882300] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35642004 kB) free.
[1587882750] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 358 processes
[1587883001] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.4% (16258220 kB) free.
[1587883500] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15312644 kB) free!
[1587883620] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11126156 kB) free!
[1587883740] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35324056 kB) free.
[1587884340] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (11973572 kB) free!
[1587884460] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (36001540 kB) free.
[1587884647] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(679) < 103.255.146.164 [26/Apr/2020:12:27:46 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587884767] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587885060] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11547184 kB) free!
[1587885180] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35101448 kB) free.
[1587885781] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11370060 kB) free!
[1587885901] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6847932 kB) free!
[1587886021] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32656448 kB) free.
[1587886621] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8118108 kB) free!
[1587886702] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 30.9% (7581472 kB) free.
[1587886741] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32881048 kB) free.
[1587886805] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1587886805] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1169976 kB) free!
[1587887341] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8851876 kB) free!
[1587887405] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.7% (1741256 kB) free!
[1587887461] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4116968 kB) free!
[1587887581] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30052728 kB) free.
[1587888181] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4476792 kB) free!
[1587888301] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30452516 kB) free.
[1587888901] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7400068 kB) free!
[1587889021] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (2992516 kB) free!
[1587889141] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29662112 kB) free.
[1587889297] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.4% (4507276 kB) free!
[1587889417] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.0% (4407804 kB) free!
[1587889538] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.7% (4348596 kB) free!
[1587889741] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (5055368 kB) free!
[1587889861] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29068700 kB) free.
[1587890124] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 684 processes
[1587890168] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(223) < 103.255.146.164 [26/Apr/2020:14:01:44 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587890244] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 359 processes
[1587890288] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587890406] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1613644 kB) free!
[1587890462] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6004672 kB) free!
[1587890582] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (1926648 kB) free!
[1587890702] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28897316 kB) free.
[1587890902] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4749340 kB) free!
[1587891006] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.1% (2291708 kB) free!
[1587891022] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174268 kB) free!
[1587891142] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202936 kB) free!
[1587891225] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.40, 3.72, 4.58
[1587891302] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5429200 kB) free!
[1587891422] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.6% (25893232 kB) free!
[1587891542] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.1% (25214948 kB) free!
[1587891581] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 40.6% (9944140 kB) free.
[1587891606] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1227936 kB) free!
[1587891750] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1587891825] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.72, 8.09, 6.02
[1587891945] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 12.24, 9.30, 6.70
[1587892065] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 14.08, 11.03, 7.65
[1587892088] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(95) < 103.255.146.164 [26/Apr/2020:14:28:21 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587892142] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (376020 kB) free!
[1587892208] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587892538] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1965720 kB) free!
[1587892601] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.2% (2960136 kB) free!
[1587892721] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.5% (2578364 kB) free!
[1587892842] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2182908 kB) free!
[1587893342] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.0% (23717552 kB) free!
[1587893409] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(144) < 185.255.8.58 [26/Apr/2020:15:00:05 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1587893442] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1587893442] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.3% (16239304 kB) free.
[1587893529] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(90) < 103.255.146.164 [26/Apr/2020:15:01:17 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587893649] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587893938] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9660516 kB) free!
[1587893942] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32581648 kB) free.
[1587894058] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.4% (9544060 kB) free!
[1587894178] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.1% (9429940 kB) free!
[1587894542] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6392380 kB) free!
[1587894606] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.7% (2387656 kB) free!
[1587894662] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11678748 kB) free!
[1587894782] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29629104 kB) free.
[1587895206] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1301424 kB) free!
[1587895382] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4368652 kB) free!
[1587895502] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32036748 kB) free.
[1587895806] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3523152 kB) free.
[1587896103] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6242536 kB) free!
[1587896223] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1816620 kB) free!
[1587896343] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30333760 kB) free.
[1587896943] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4895848 kB) free!
[1587897006] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3256780 kB) free!
[1587897063] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7409704 kB) free!
[1587897126] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (3057968 kB) free!
[1587897183] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27814452 kB) free.
[1587897246] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.8% (2730844 kB) free!
[1587897783] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (2489240 kB) free!
[1587897903] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31189132 kB) free.
[1587898378] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4913844 kB) free!
[1587898503] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5595664 kB) free!
[1587898623] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1234424 kB) free!
[1587898665] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.72, 3.57, 4.65
[1587898743] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30638652 kB) free.
[1587899343] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (4976436 kB) free!
[1587899463] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (413540 kB) free!
[1587899583] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24680144 kB) free!
[1587899840] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1587899960] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 328 processes
[1587900080] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1587900183] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384184 kB) free!
[1587900246] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3424908 kB) free.
[1587900249] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(228) < 103.255.146.164 [26/Apr/2020:16:48:27 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587900369] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587900784] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13598428 kB) free!
[1587900969] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(630) < 103.255.146.164 [26/Apr/2020:17:01:11 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587901089] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587901446] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.1% (2290280 kB) free!
[1587901566] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.8% (2081604 kB) free!
[1587901665] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.38, 6.09, 5.01
[1587901686] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.6% (1892476 kB) free!
[1587901785] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.98, 5.68, 4.99
[1587901984] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28639760 kB) free.
[1587902985] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.73, 7.03, 6.27
[1587903105] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.12, 5.92, 5.95
[1587903184] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9652164 kB) free!
[1587903225] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.56, 5.93, 5.95
[1587903304] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4829764 kB) free!
[1587903424] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34328300 kB) free.
[1587904024] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (9982840 kB) free!
[1587904144] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (4994996 kB) free!
[1587904253] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.5% (23024240 kB) free!
[1587904264] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32919456 kB) free.
[1587904373] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.6% (28439392 kB) free.
[1587904686] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (1388964 kB) free!
[1587904842] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1587904864] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6153624 kB) free!
[1587904984] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (1994180 kB) free!
[1587905104] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30871588 kB) free.
[1587905704] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5640944 kB) free!
[1587905824] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32641436 kB) free.
[1587906124] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9625028 kB) free!
[1587906225] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.76, 4.61, 4.77
[1587906244] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (9561692 kB) free!
[1587906364] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (9489940 kB) free!
[1587906425] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7844172 kB) free!
[1587906486] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (3205264 kB) free!
[1587906545] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2909232 kB) free!
[1587906665] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31995668 kB) free.
[1587906681] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1587906773] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.4% (25576184 kB) free!
[1587906801] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1587906825] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.17, 6.75, 5.91
[1587906893] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.2% (29245648 kB) free.
[1587906945] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.03, 6.51, 5.93
[1587907065] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1587907065] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.38, 6.12, 5.85
[1587907265] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6164028 kB) free!
[1587907385] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1856416 kB) free!
[1587907493] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.2% (23975724 kB) free!
[1587907505] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29796700 kB) free.
[1587907613] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.8% (30045532 kB) free.
[1587908105] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4110556 kB) free!
[1587908225] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31096360 kB) free.
[1587908265] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.30, 3.01, 4.18
[1587908825] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7282356 kB) free!
[1587908945] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2868404 kB) free!
[1587909065] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30761424 kB) free.
[1587909665] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2874644 kB) free!
[1587909785] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393412 kB) free!
[1587909905] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (31025248 kB) free.
[1587910065] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.16, 5.88, 5.14
[1587910086] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (976960 kB) free!
[1587910185] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.11, 6.18, 5.33
[1587910305] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.96, 6.06, 5.39
[1587910505] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (3169340 kB) free!
[1587910626] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33523636 kB) free.
[1587910686] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.3% (2325640 kB) free!
[1587911225] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6314728 kB) free!
[1587911286] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1277360 kB) free!
[1587911345] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (1423932 kB) free!
[1587911465] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34585096 kB) free.
[1587911886] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3568984 kB) free.
[1587912065] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6105564 kB) free!
[1587912185] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (1401080 kB) free!
[1587912289] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < webmachine error: path="/api/whoami"
[1587912305] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34140812 kB) free.
[1587912409] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1587912486] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (3001512 kB) free!
[1587912606] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.2% (2791712 kB) free!
[1587912705] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.09, 4.28, 4.78
[1587912726] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (2535800 kB) free!
[1587912905] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (3970116 kB) free!
[1587913025] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (384816 kB) free!
[1587913145] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33189612 kB) free.
[1587913306] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.53, 6.70, 5.65
[1587913326] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (1441880 kB) free!
[1587913426] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.45, 5.90, 5.49
[1587913546] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.41, 5.40, 5.35
[1587913565] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4921036 kB) free!
[1587913745] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4325448 kB) free!
[1587913865] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (385908 kB) free!
[1587913926] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (3020620 kB) free!
[1587913985] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33566540 kB) free.
[1587914585] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (4794860 kB) free!
[1587914705] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (402408 kB) free!
[1587914746] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.62, 3.97, 4.60
[1587914813] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.3% (22739824 kB) free!
[1587914825] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32166652 kB) free.
[1587914933] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.4% (28230176 kB) free.
[1587915346] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.69, 6.07, 5.27
[1587915426] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4259212 kB) free!
[1587915466] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.69, 5.89, 5.31
[1587915533] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 20.0% (26370944 kB) free!
[1587915546] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34269576 kB) free.
[1587915586] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.55, 6.27, 5.53
[1587915653] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.0% (30283304 kB) free.
[1587915726] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1031632 kB) free!
[1587916146] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10029940 kB) free!
[1587916266] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (5888536 kB) free!
[1587916326] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.4% (3808228 kB) free.
[1587916386] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1587916386] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33797564 kB) free.
[1587916986] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9088796 kB) free!
[1587917106] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4885384 kB) free!
[1587917226] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33222744 kB) free.
[1587917526] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.3% (2648480 kB) free!
[1587917646] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.7% (2382232 kB) free!
[1587917766] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.1% (2129128 kB) free!
[1587917826] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10282072 kB) free!
[1587917946] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6427120 kB) free!
[1587917986] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.85, 3.30, 4.37
[1587918066] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31210744 kB) free.
[1587918366] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3637068 kB) free.
[1587918666] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8977572 kB) free!
[1587918786] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33218264 kB) free.
[1587918966] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.6% (2541280 kB) free!
[1587919086] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.8% (2244904 kB) free!
[1587919206] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3590532 kB) free.
[1587919386] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10797700 kB) free!
[1587919506] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (6784376 kB) free!
[1587919626] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31518236 kB) free.
[1587920227] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9479772 kB) free!
[1587920347] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33029748 kB) free.
[1587920947] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11190520 kB) free!
[1587921006] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (3135964 kB) free!
[1587921067] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7223188 kB) free!
[1587921126] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.1% (2778864 kB) free!
[1587921187] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30727532 kB) free.
[1587921247] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.1% (2296960 kB) free!
[1587921787] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8392964 kB) free!
[1587921907] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32825532 kB) free.
[1587922447] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.2% (3438488 kB) free.
[1587922507] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10045004 kB) free!
[1587922627] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32777880 kB) free.
[1587923227] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12623236 kB) free!
[1587923347] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8877860 kB) free!
[1587923467] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31403356 kB) free.
[1587923647] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.2% (2312272 kB) free!
[1587923767] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3436760 kB) free.
[1587924067] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10103048 kB) free!
[1587924187] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32737036 kB) free.
[1587924787] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11294424 kB) free!
[1587924908] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32712196 kB) free.
[1587925507] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13358988 kB) free!
[1587925627] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9721476 kB) free!
[1587925747] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31565384 kB) free.