[1581878033] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.7% (29405612 kB) free.
[1581878578] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183740 kB) free!
[1581878698] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180448 kB) free!
[1581878818] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171940 kB) free!
[1581879755] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195584 kB) free!
[1581879815] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185804 kB) free!
[1581879875] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (181048 kB) free!
[1581879935] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 71.0% (17394948 kB) free.
[1581881756] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 48.0% (11835168 kB) free.
[1581882356] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.7% (3868656 kB) free!
[1581882476] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (1641168 kB) free!
[1581882596] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1581882596] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181232 kB) free!
[1581882755] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9213608 kB) free!
[1581882875] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.5% (8122472 kB) free!
[1581882995] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.4% (7109920 kB) free!
[1581883169] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1581883289] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1581883595] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (2775448 kB) free!
[1581886576] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (248060 kB) free!
[1581886696] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (235916 kB) free!
[1581886816] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233248 kB) free!
[1581887218] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 49.9% (12236416 kB) free.
[1581887416] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 80.5% (39647200 kB) free.
[1581888711] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(226) < 2020/02/17 03:00:22 [error] 19831#0: *55075195 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581888760] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(174) < 103.16.101.72 [17/Feb/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1581888831] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581888881] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581889085] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1645) < 203.212.70.170 [17/Feb/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1581889205] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581891089] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1581891209] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1581893996] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581895805] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(529) < 172.105.59.173 [17/Feb/2020:04:59:07 +0530] TCP 502 0 0 180.002 "10.147.212.155:3001, 10.147.212.154:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.001, 60.000, 60.001, 0.000"
[1581895925] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1) < 172.105.59.173 [17/Feb/2020:05:00:07 +0530] TCP 502 0 0 240.000 "10.147.212.154:3001, 10.147.212.143:3001, 10.147.212.142:3001, 10.147.212.155:3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 60.000"
[1581896032] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(61) < 2020/02/17 04:56:56 [error] 19830#0: *55104748 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581896045] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581896152] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581896219] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176548 kB) free!
[1581896339] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177260 kB) free!
[1581896459] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192228 kB) free!
[1581896996] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 45.1% (11122568 kB) free.
[1581897018] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.0% (20822588 kB) free.
[1581899609] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 406 processes
[1581899729] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 406 processes
[1581900293] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 88.8% (117021616 kB) free.
[1581900328] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.5% (39650772 kB) free.
[1581900372] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.1% (21593888 kB) free.
[1581900503] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 238364 MB (22% inode=99%):
[1581901447] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(135) < 92.42.104.166 [17/Feb/2020:06:30:18 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1581901567] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581901796] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (1780864 kB) free!
[1581901916] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (1758428 kB) free!
[1581902036] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1734460 kB) free!
[1581903318] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.0% (21712804 kB) free.
[1581906368] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(121) < 92.42.104.166 [17/Feb/2020:07:50:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1581906488] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581906929] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1581906952] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(68) < 2020/02/17 08:02:04 [error] 19830#0: *55167848 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581907049] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1581907072] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581907088] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 203.212.70.170 [17/Feb/2020:08:03:28 +0530] TCP 502 0 0 60.001 "10.147.212.143:3001, 10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 0.001, 60.000, 0.000"
[1581907208] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581907672] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(28) < 2020/02/17 08:16:46 [error] 19841#0: *55173764 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581907793] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581908249] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1581908370] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1581914105] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.5% (4527732 kB) free!
[1581914225] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.8% (4366252 kB) free!
[1581914345] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.4% (4267040 kB) free!
[1581915698] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1581915819] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1581916145] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2331592 kB) free!
[1581916599] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 341 processes
[1581916719] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 340 processes
[1581916839] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1581917618] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1581917726] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.8% (4143980 kB) free!
[1581917738] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1581917846] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.0% (3952620 kB) free!
[1581917967] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.2% (3745876 kB) free!
[1581917983] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.3% (21653876 kB) free.
[1581918567] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (1397168 kB) free!
[1581918629] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4446472 kB) free!
[1581918640] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1581918749] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.1% (4055392 kB) free!
[1581918869] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.4% (3691028 kB) free!
[1581918938] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1581919058] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1581919098] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (18986772 kB) free!
[1581919179] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1581919218] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41920584 kB) free.
[1581919469] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (1984984 kB) free!
[1581919779] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 696 processes
[1581919792] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/02/17 11:37:01 [error] 19835#0: *55264678 open() "/usr/share/nginx/html/js/tmpl/404.ejs" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /js/tmpl/404.ejs?0.8049999510579628 HTTP/1.1", host: "qservercluster1.com:8091", referrer: "http://qservercluster1.com:8091/"
[1581919818] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (16969180 kB) free!
[1581919912] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581919938] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (18868444 kB) free!
[1581920033] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9743244 kB) free!
[1581920058] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32607108 kB) free.
[1581920153] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.3% (9034688 kB) free!
[1581920273] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.9% (8329276 kB) free!
[1581920658] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8229340 kB) free!
[1581920691] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1581920708] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 17.8% (4398092 kB) free!
[1581920778] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.0% (43456076 kB) free.
[1581920811] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1581920828] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 16.0% (3960784 kB) free!
[1581920948] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 14.4% (3552252 kB) free!
[1581920979] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 713 processes
[1581921099] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 713 processes
[1581921219] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 708 processes
[1581921269] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.7% (15877616 kB) free.
[1581921378] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19821736 kB) free!
[1581921411] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 332 processes
[1581921473] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.2% (30180000 kB) free.
[1581921499] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.3% (43905412 kB) free.
[1581921531] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1581921548] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1279800 kB) free!
[1581921867] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (185408 kB) free!
[1581921981] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176072 kB) free!
[1581921987] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (226836 kB) free!
[1581922098] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19115928 kB) free!
[1581922101] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189724 kB) free!
[1581922107] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (171444 kB) free!
[1581922148] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.6% (22868156 kB) free.
[1581922218] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.4% (45287252 kB) free.
[1581922221] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171404 kB) free!
[1581922272] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.3% (21800224 kB) free.
[1581922311] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.0% (21322248 kB) free.
[1581922356] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.8% (14108564 kB) free.
[1581922818] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (18949916 kB) free!
[1581922938] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.3% (14852328 kB) free!
[1581923058] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41081520 kB) free.
[1581923136] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9179660 kB) free!
[1581923256] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.6% (9173708 kB) free!
[1581923376] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.6% (9172688 kB) free!
[1581923416] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9523504 kB) free!
[1581923418] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.4% (4272928 kB) free!
[1581923537] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.3% (9493504 kB) free!
[1581923538] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.0% (4162100 kB) free!
[1581923657] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.5% (9117384 kB) free!
[1581923658] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15392936 kB) free!
[1581923659] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.4% (4017044 kB) free!
[1581923778] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.3% (41223784 kB) free.
[1581924219] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1581924219] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 695 processes
[1581924378] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (16960764 kB) free!
[1581924498] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (12529976 kB) free!
[1581924576] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4833476 kB) free!
[1581924618] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38212256 kB) free.
[1581924678] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180592 kB) free!
[1581924712] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/17 12:56:21 [error] 19836#0: *55308637 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1581924765] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (174852 kB) free!
[1581924798] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171224 kB) free!
[1581924832] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581924885] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (200156 kB) free!
[1581924918] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1581924918] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182228 kB) free!
[1581925006] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (276688 kB) free!
[1581925079] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.9% (21682576 kB) free.
[1581925121] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.4% (39640860 kB) free.
[1581925152] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1581925152] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.6% (14232864 kB) free.
[1581925187] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192660 kB) free!
[1581925218] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12178260 kB) free!
[1581925307] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193528 kB) free!
[1581925338] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (8043844 kB) free!
[1581925427] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186120 kB) free!
[1581925457] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4507056 kB) free!
[1581925458] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35157936 kB) free.
[1581925459] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193412 kB) free!
[1581925501] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1581925621] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1581925688] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4907660 kB) free!
[1581925808] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179660 kB) free!
[1581925929] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173276 kB) free!
[1581926059] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9779960 kB) free!
[1581926179] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36293116 kB) free.
[1581926355] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.1% (2298996 kB) free!
[1581926475] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.1% (2292420 kB) free!
[1581926595] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 14.1% (2290968 kB) free!
[1581926627] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1581926627] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 75.8% (18572140 kB) free.
[1581926779] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10616048 kB) free!
[1581926899] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6317576 kB) free!
[1581927019] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34507536 kB) free.
[1581927118] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (4789284 kB) free!
[1581927238] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4621208 kB) free!
[1581927359] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.2% (4462184 kB) free!
[1581927619] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8877236 kB) free!
[1581927739] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4527420 kB) free!
[1581927859] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32658000 kB) free.
[1581928459] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5762168 kB) free!
[1581928558] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1581928558] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2181992 kB) free!
[1581928579] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1566820 kB) free!
[1581928699] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29398120 kB) free.
[1581929069] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.9% (4010860 kB) free!
[1581929189] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.4% (3674336 kB) free!
[1581929299] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (3211456 kB) free!
[1581929309] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.7% (3308088 kB) free!
[1581929419] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (385336 kB) free!
[1581929539] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26082340 kB) free!
[1581929909] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (24.7% change >= 20.0% threshold)
[1581929909] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1957200 kB) free!
[1581929928] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9164740 kB) free!
[1581930019] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.3% (21554816 kB) free.
[1581930048] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.7% (8744796 kB) free!
[1581930058] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 25.1% (12389908 kB) free.
[1581930095] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.1% (20860716 kB) free.
[1581930134] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.1% (21355764 kB) free.
[1581930139] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406632 kB) free!
[1581930168] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1581930168] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.0% (7883228 kB) free!
[1581930176] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 63.7% (14308328 kB) free.
[1581930215] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16463908 kB) free!
[1581930815] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1581930815] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 39.1% (51481240 kB) free.
[1581931674] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8557692 kB) free!
[1581931769] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 405 processes
[1581931794] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8296104 kB) free!
[1581931808] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(163) < 49.50.67.175 [17/Feb/2020:15:00:06 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1581931889] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1581931914] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.2% (29685768 kB) free.
[1581931929] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(19) < 92.42.104.166 [17/Feb/2020:15:00:22 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1581932032] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/02/17 14:55:23 [error] 19841#0: *55372940 open() "/usr/share/nginx/html/js/tmpl/404.ejs" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /js/tmpl/404.ejs?0.4552245870360392 HTTP/1.1", host: "qservercluster1.com:8091", referrer: "http://qservercluster1.com:8091/"
[1581932049] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581932152] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581933215] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23973180 kB) free!
[1581933219] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1581933335] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.7% (19411260 kB) free!
[1581933455] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13353452 kB) free!
[1581933735] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9668732 kB) free!
[1581933768] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3385440 kB) free!
[1581933855] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (9553556 kB) free!
[1581933975] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (9406700 kB) free!
[1581934655] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34620172 kB) free.
[1581935201] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1581935255] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9674368 kB) free!
[1581935266] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.3% (2808440 kB) free!
[1581935321] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1581935375] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.5% (49414652 kB) free.
[1581935386] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.7% (2880072 kB) free!
[1581935507] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (3029428 kB) free!
[1581935649] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [17/Feb/2020:16:02:29 +0530] TCP 502 0 0 60.002 "10.147.212.155:3001, 10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 0.000, 60.001, 0.000"
[1581935701] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 744 processes
[1581935752] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(156) < 2020/02/17 16:02:08 [error] 19841#0: *55399934 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1581935769] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581935821] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 326 processes
[1581935828] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 62.0% (30556168 kB) free.
[1581935872] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581935921] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1581935975] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15755524 kB) free!
[1581936025] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.1% (3468632 kB) free!
[1581936040] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 350 processes
[1581936042] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1581936095] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (11572704 kB) free!
[1581936106] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3672564 kB) free.
[1581936145] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (227520 kB) free!
[1581936215] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (7700004 kB) free!
[1581936265] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190256 kB) free!
[1581936641] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1581936761] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1581936795] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1581936815] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.4% (19013740 kB) free!
[1581936827] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.9% (4397744 kB) free!
[1581936881] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1581936947] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.2% (4213276 kB) free!
[1581937067] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.5% (4055236 kB) free!
[1581937415] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39831088 kB) free.
[1581937575] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4524260 kB) free!
[1581937895] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.9% (4146656 kB) free!
[1581938015] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13131800 kB) free!
[1581938015] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.4% (4008108 kB) free!
[1581938135] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.2% (45131200 kB) free.
[1581938135] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.8% (3881340 kB) free!
[1581938644] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.1% (21496940 kB) free.
[1581938681] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 31.6% (15591724 kB) free.
[1581938735] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17650640 kB) free!
[1581938855] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (13047928 kB) free!
[1581938867] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2181692 kB) free!
[1581938975] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (8813856 kB) free!
[1581939336] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1776328 kB) free!
[1581939576] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17235552 kB) free!
[1581939640] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1581939706] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.0% (2768340 kB) free!
[1581939786] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581939826] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (3208196 kB) free!
[1581939946] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.2% (3452608 kB) free.
[1581940176] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30361388 kB) free.
[1581940386] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4263444 kB) free!
[1581940506] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4090116 kB) free!
[1581940546] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (1533180 kB) free!
[1581940626] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.3% (3892620 kB) free!
[1581940637] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9535000 kB) free!
[1581940666] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.0% (4062004 kB) free.
[1581940757] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.9% (8840208 kB) free!
[1581940776] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5801432 kB) free!
[1581940877] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.0% (7891876 kB) free!
[1581940896] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37899564 kB) free.
[1581941047] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (578296 kB) free!
[1581941167] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (452872 kB) free!
[1581941266] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.1% (1965788 kB) free!
[1581941287] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (280184 kB) free!
[1581941386] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.7% (2552316 kB) free!
[1581941496] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9914836 kB) free!
[1581941506] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (3049792 kB) free!
[1581941616] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (5599848 kB) free!
[1581941736] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39085012 kB) free.
[1581941826] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1677196 kB) free!
[1581942107] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1581942107] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.8% (4189664 kB) free.
[1581942336] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (9994948 kB) free!
[1581942456] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (5753548 kB) free!
[1581942576] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36610568 kB) free.
[1581942746] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 69.2% (16953224 kB) free.
[1581943027] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.7% (15880668 kB) free.
[1581943088] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 65.1% (16065580 kB) free.
[1581943176] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6752676 kB) free!
[1581943296] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2350304 kB) free!
[1581943307] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (1042200 kB) free!
[1581943416] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35988284 kB) free.
[1581943427] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (1532264 kB) free!
[1581943547] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.9% (1927200 kB) free!
[1581944016] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7599880 kB) free!
[1581944136] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3314084 kB) free!
[1581944147] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.0% (3416588 kB) free.
[1581944257] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36378828 kB) free.
[1581944478] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3705284 kB) free!
[1581944747] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (1487412 kB) free!
[1581944769] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(478) < 103.255.146.164 [17/Feb/2020:18:30:51 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1581944856] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7340772 kB) free!
[1581944867] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.1% (1974740 kB) free!
[1581944889] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581944976] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (3062900 kB) free!
[1581944987] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (963452 kB) free!
[1581945096] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (375504 kB) free!
[1581945587] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (2774028 kB) free!
[1581946187] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.1% (4241928 kB) free.
[1581946296] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.4% (24284104 kB) free!
[1581946420] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 708 processes
[1581946540] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 706 processes
[1581946660] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 706 processes
[1581946688] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.8% (4136472 kB) free!
[1581946808] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.3% (3780908 kB) free!
[1581946896] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (429988 kB) free!
[1581946928] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1826180 kB) free!
[1581947387] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.9% (2258804 kB) free!
[1581947507] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.8% (2573916 kB) free!
[1581947600] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.2% (21281084 kB) free.
[1581947627] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (2815172 kB) free!
[1581947638] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23075612 kB) free!
[1581947642] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36224452 kB) free.
[1581947672] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 69.1% (34072652 kB) free.
[1581947701] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 73.5% (18010540 kB) free.
[1581948068] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9393444 kB) free!
[1581948157] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1581948188] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.6% (9160212 kB) free!
[1581948205] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.1% (2745764 kB) free!
[1581948227] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.8% (3864264 kB) free.
[1581948308] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.0% (8866312 kB) free!
[1581948326] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172160 kB) free!
[1581948446] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177180 kB) free!
[1581948737] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (4605020 kB) free!
[1581948858] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (4562976 kB) free!
[1581948978] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (4498800 kB) free!
[1581949428] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (1356684 kB) free!
[1581949548] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (1400500 kB) free!
[1581949668] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3319528 kB) free.
[1581949840] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 356 processes
[1581949960] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 394 processes
[1581950272] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(166) < 2020/02/17 20:07:35 [error] 12257#0: *55497196 access forbidden by rule while initializing session, client: 124.153.112.226, server: 0.0.0.0:4000
[1581950289] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(9) < 203.212.70.170 [17/Feb/2020:20:02:15 +0530] TCP 502 0 0 91.069 "10.147.212.154:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "31.069, 60.000, 0.000"
[1581950392] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(28) < 2020/02/17 20:09:35 [error] 12252#0: *55497881 access forbidden by rule while initializing session, client: 124.153.112.226, server: 0.0.0.0:4000
[1581950410] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581950513] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;HARD;3;(28) < 2020/02/17 20:11:37 [error] 12250#0: *55498589 access forbidden by rule while initializing session, client: 124.153.112.226, server: 0.0.0.0:4000
[1581951378] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165880 kB) free!
[1581951908] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4648468 kB) free!
[1581952445] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16396736 kB) free!
[1581952565] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.5% (49373400 kB) free.
[1581953106] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1581953165] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24219536 kB) free!
[1581953285] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.3% (20128232 kB) free!
[1581953406] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 39.0% (51393952 kB) free.
[1581954877] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1581955176] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 77.9% (19229624 kB) free.
[1581955205] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (19033312 kB) free!
[1581955227] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 35.6% (17561344 kB) free.
[1581955266] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 55.6% (13619560 kB) free.
[1581955325] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.2% (47727072 kB) free.
[1581955461] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1581955477] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9112664 kB) free!
[1581955582] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1581955597] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9113004 kB) free!
[1581955702] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1581955717] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.5% (9112524 kB) free!
[1581955925] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24345444 kB) free!
[1581956045] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.6% (20577604 kB) free!
[1581956165] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.5% (50763328 kB) free.
[1581956765] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (25105500 kB) free!
[1581956869] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (2795600 kB) free!
[1581956885] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.3% (21522488 kB) free!
[1581956989] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 38.0% (6173068 kB) free.
[1581957005] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18042416 kB) free!
[1581957605] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28993976 kB) free.
[1581958789] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1581958805] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17251224 kB) free!
[1581958925] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.8% (49867364 kB) free.
[1581959525] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23272780 kB) free!
[1581959646] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.8% (19488772 kB) free!
[1581959766] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.7% (49634652 kB) free.
[1581959989] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581960366] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24822112 kB) free!
[1581960461] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1581960486] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.9% (20961624 kB) free!
[1581960502] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1581960606] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17390136 kB) free!
[1581961806] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35638492 kB) free.
[1581962513] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581963006] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (24008768 kB) free!
[1581963114] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (233480 kB) free!
[1581963126] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.5% (20493308 kB) free!
[1581963234] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (234804 kB) free!
[1581963246] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.1% (47530036 kB) free.
[1581963355] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (212580 kB) free!
[1581963846] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24075948 kB) free!
[1581963897] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4161872 kB) free!
[1581963966] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.6% (20510760 kB) free!
[1581964017] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7459512 kB) free!
[1581964086] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.4% (46709900 kB) free.
[1581964137] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (7647236 kB) free!