[1546196863] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.30, 10.49, 10.20
[1546196983] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.91, 10.24, 10.14
[1546197103] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.59, 10.33, 10.18
[1546198395] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 129.97, 127.18, 53.88
[1546198515] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 17.60, 85.16, 47.37
[1546198635] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.64, 57.08, 41.66
[1546198903] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 6.56, 7.42, 9.16
[1546199835] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.33, 2.09, 12.16
[1546201035] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.08, 0.20, 3.46
[1546204904] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.51, 11.50, 10.26
[1546205024] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.70, 11.47, 10.39
[1546205144] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 13.09, 12.03, 10.73
[1546208744] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 6.60, 7.31, 8.69
[1546210107] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 355 processes
[1546213544] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.98, 10.63, 10.19
[1546213664] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.16, 10.49, 10.20
[1546213784] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.53, 10.37, 10.18
[1546216035] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 3.60, 55.27, 37.49
[1546216107] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 358 processes
[1546216155] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.82, 37.12, 32.99
[1546216184] SERVICE FLAPPING ALERT: qservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1546216184] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 7.36, 8.12, 9.70
[1546216276] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.13, 24.85, 28.99
[1546216876] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.13, 3.55, 15.32
[1546218076] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.32, 0.29, 4.40
[1546220984] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.58, 10.96, 10.01
[1546221104] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.72, 10.85, 10.08
[1546221224] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.42, 10.66, 10.10
[1546222424] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 6.93, 7.73, 8.98
[1546224507] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 360 processes
[1546228424] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.31, 10.33, 10.12
[1546228544] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.11, 10.50, 10.22
[1546228664] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.20, 10.23, 10.15
[1546229864] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 6.62, 8.44, 9.83
[1546229907] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 354 processes
[1546231556] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 17 users currently logged in
[1546231676] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 17 users currently logged in
[1546231796] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 17 users currently logged in
[1546233464] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.68, 11.25, 10.54
[1546233584] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.67, 10.73, 10.45
[1546233704] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.65, 10.39, 10.36
[1546238191] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546238249] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9662872 kB) free!
[1546238311] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546238369] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.5% (8152904 kB) free!
[1546238431] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546238490] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10278072 kB) free.
[1546239031] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546239873] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 338 processes
[1546242104] SERVICE FLAPPING ALERT: qservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1546244308] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 368 processes
[1546245676] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 0.72, 32.60, 37.95
[1546245796] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.73, 22.03, 33.43
[1546245873] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1546245916] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.61, 14.97, 29.46
[1546246516] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.56, 2.57, 15.77
[1546247431] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546247489] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9563200 kB) free!
[1546247551] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546247609] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8027364 kB) free!
[1546247671] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1546247716] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.16, 0.37, 4.67
[1546247729] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10276048 kB) free.
[1546253673] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1546253716] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 0.83, 39.93, 34.13
[1546253793] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 333 processes
[1546253836] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.41, 26.83, 30.04
[1546253914] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 333 processes
[1546253956] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1546253956] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.29, 18.04, 26.42
[1546254556] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.37, 2.72, 14.00
[1546254961] SERVICE ALERT: smscapp1;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1546255304] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.55, 4.75, 9.13
[1546255756] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.63, 0.64, 4.26
[1546256162] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1546256282] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1546256314] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 335 processes
[1546256402] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1546257002] SERVICE ALERT: smscapp1;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1546257602] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1546257722] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1546257842] SERVICE FLAPPING ALERT: smscapp1;CPU_Procs;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1546257842] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1546258442] SERVICE ALERT: smscapp1;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1546259042] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1546259163] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1546259283] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1546259476] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1546259479] HOST ALERT: sig2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.150)
[1546259490] SERVICE ALERT: sig2;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig2' was out of bounds)
[1546259543] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 201.67 ms
[1546259596] SERVICE ALERT: sig2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 316 processes
[1546259610] SERVICE ALERT: sig2;Sensors Status;OK;HARD;1;SENSORS OK
[1546259717] SERVICE ALERT: sig2;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1546261439] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 64.3% (10450476 kB) free.
[1546262071] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546262129] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9639892 kB) free!
[1546262191] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546262249] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.5% (8128180 kB) free!
[1546262311] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546262314] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 339 processes
[1546262369] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10260156 kB) free.
[1546262911] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1546264196] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 11 users currently logged in
[1546265911] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 289 processes
[1546265969] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7448632 kB) free!
[1546266031] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546266089] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10259088 kB) free.
[1546266556] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1546268356] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 634.16, 192.40, 67.07
[1546268476] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 86.11, 128.94, 59.01
[1546268596] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 12.30, 86.57, 51.96
[1546269796] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.73, 2.41, 14.88
[1546270683] SERVICE FLAPPING ALERT: smscapp1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1546270996] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.83, 0.83, 4.68
[1546276396] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 9.53, 75.12, 45.97
[1546276516] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.33, 50.62, 40.52
[1546276636] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1546276636] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.90, 34.09, 35.69
[1546277236] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.90, 5.78, 19.56
[1546278872] SERVICE ALERT: qservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 301 processes
[1546278992] SERVICE ALERT: qservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 302 processes
[1546279036] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.74, 0.71, 3.40
[1546279112] SERVICE ALERT: qservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 302 processes