[1544642152] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1544642272] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544642392] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544642992] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1544643630] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 560.49, 128.78, 43.12
[1544643750] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 476.91, 353.35, 143.79
[1544643870] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1544643870] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 65.52, 236.88, 126.52
[1544645670] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.93, 1.45, 18.94
[1544647470] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.04, 1.33, 3.88
[1544652244] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1544652870] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 3.63, 122.40, 97.83
[1544652990] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.33, 82.21, 86.08
[1544653110] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.96, 55.35, 75.78
[1544654910] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.74, 1.12, 11.81
[1544656110] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.00, 1.12, 4.10
[1544666910] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1544670510] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.76, 99.99, 96.14
[1544670630] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.49, 67.50, 84.71
[1544670750] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.95, 45.78, 74.66
[1544672550] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.60, 1.84, 12.19
[1544673750] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.54, 1.44, 4.41
[1544676244] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1544676365] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 333 processes
[1544676485] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 333 processes
[1544677192] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1544677312] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1544679485] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1544680127] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1544681327] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 314 processes
[1544681447] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 314 processes
[1544681567] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 314 processes
[1544682167] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1544682712] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544682832] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544682952] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1544683085] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 340 processes
[1544683205] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 339 processes
[1544683325] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 342 processes
[1544684305] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 319 processes
[1544688726] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1544689926] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1544690046] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 344 processes
[1544690166] SERVICE FLAPPING ALERT: smscapp2;CPU_Procs;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1544690166] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1544690715] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.0% (3083784 kB) free!
[1544690752] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1544690835] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.3% (2979472 kB) free!
[1544690872] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1544690955] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.7% (2876228 kB) free!
[1544690992] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1544691966] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 347 processes
[1544692155] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1192148 kB) free!
[1544693638] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1622380 kB) free!
[1544694592] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1544694712] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544694832] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1544696632] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544696753] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1544698203] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 41 users currently logged in
[1544698263] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 41 users currently logged in
[1544698323] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 41 users currently logged in
[1544698383] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 43 users currently logged in
[1544698949] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (285172 kB) free!
[1544699047] HOST ALERT: sservercluster1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.49)
[1544699047] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sservercluster1' was out of bounds)
[1544699069] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (1479028 kB) free!
[1544699134] HOST ALERT: sservercluster1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1544699167] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 302 processes
[1544699189] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175292 kB) free!
[1544699766] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 345 processes
[1544699767] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1544701552] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544701673] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544701793] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1544702355] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.8% (2725924 kB) free!
[1544704985] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1544705105] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1544705583] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 40 users currently logged in
[1544705767] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1544707193] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1544707313] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544707433] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1544708283] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 41 users currently logged in
[1544708343] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 41 users currently logged in
[1544708403] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 41 users currently logged in
[1544708463] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 41 users currently logged in
[1544708567] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 315 processes
[1544708906] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 2 warn out of 338 processes
[1544710263] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 40 users currently logged in
[1544711167] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 352 processes
[1544711287] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 350 processes
[1544711407] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 352 processes
[1544712833] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1544712953] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1544713073] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1544718007] SERVICE FLAPPING ALERT: smscapp2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)