[1590172531] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (446784 kB) free!
[1590172771] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 359 processes
[1590173131] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.9% (20960604 kB) free!
[1590173370] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 357 processes
[1590173416] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1590173416] SERVICE ALERT: sigqscluster1;Total Processes;OK;HARD;1;PROCS OK: 1425 processes
[1590173428] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (885560 kB) free!
[1590173731] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398972 kB) free!
[1590173970] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 363 processes
[1590174016] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1787 processes
[1590174028] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (3144372 kB) free!
[1590174136] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1878 processes
[1590174257] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1875 processes
[1590174332] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.7% (23382460 kB) free!
[1590174570] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1590174570] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 352 processes
[1590174838] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.46, 6.83, 5.77
[1590174852] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 636 processes
[1590174932] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382868 kB) free!
[1590174958] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.87, 6.48, 5.77
[1590175078] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.32, 6.67, 5.94
[1590175170] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 350 processes
[1590175228] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (702704 kB) free!
[1590175532] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30199804 kB) free.
[1590175770] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 350 processes
[1590175828] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.7% (2557392 kB) free!
[1590176132] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5088916 kB) free!
[1590176247] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41544 MB (18.48% inode=100%):
[1590176252] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (1036084 kB) free!
[1590176308] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 37326 MB (16.60% inode=100%):
[1590176368] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;3;DISK WARNING - free space: /home 31166 MB (13.86% inode=100%):
[1590176370] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 350 processes
[1590176372] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (412296 kB) free!
[1590176428] SERVICE ALERT: webserver2;HDD_Home status;WARNING;HARD;4;DISK WARNING - free space: /home 30286 MB (13.47% inode=100%):
[1590176728] SERVICE ALERT: webserver2;HDD_Home status;CRITICAL;HARD;4;DISK CRITICAL - free space: /home 14136 MB (6.29% inode=100%):
[1590176970] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 349 processes
[1590176972] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19962224 kB) free!
[1590177028] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1054240 kB) free!
[1590177105] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 83.5% (20461304 kB) free.
[1590177328] SERVICE ALERT: webserver2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 72110 MB (32.08% inode=100%):
[1590177570] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 354 processes
[1590177572] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403504 kB) free!
[1590177628] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.7% (2060800 kB) free!
[1590177753] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (237608 kB) free!
[1590177874] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (238012 kB) free!
[1590177946] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (414416 kB) free!
[1590177994] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237264 kB) free!
[1590178066] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (418008 kB) free!
[1590178170] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 350 processes
[1590178172] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.2% (23980020 kB) free!
[1590178186] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (450636 kB) free!
[1590178770] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 351 processes
[1590178772] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1468748 kB) free!
[1590179372] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.4% (25544884 kB) free!
[1590179794] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.2% (8000004 kB) free!
[1590179972] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1783104 kB) free!
[1590180029] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (938136 kB) free!
[1590180572] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29104608 kB) free.
[1590180629] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.1% (2284200 kB) free!
[1590181172] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4259520 kB) free!
[1590181229] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (1384512 kB) free!
[1590181293] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (402612 kB) free!
[1590181413] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384332 kB) free!
[1590181829] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.9% (3730148 kB) free.
[1590182013] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16902184 kB) free!
[1590182370] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 349 processes
[1590182429] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (1632588 kB) free!
[1590182549] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (1043232 kB) free!
[1590182613] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (433132 kB) free!
[1590182632] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.5% (46101848 kB) free.
[1590182669] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (516076 kB) free!
[1590182781] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 93.5% (46090948 kB) free.
[1590182970] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 352 processes
[1590183116] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(433) < 2020/05/23 03:01:35 [error] 22135#0: *90131604 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1590183199] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2643) < 203.212.70.170 [23/May/2020:03:02:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1590183213] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17670536 kB) free!
[1590183236] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590183269] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.2% (2635128 kB) free!
[1590183319] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590183395] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 89.7% (44202844 kB) free.
[1590183478] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1590183570] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1590183588] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(340) < 103.16.101.72 [23/May/2020:03:01:35 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1590183709] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590183813] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (427696 kB) free!
[1590183869] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3478768 kB) free.
[1590184170] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1590184413] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15672012 kB) free!
[1590185013] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (437252 kB) free!
[1590185371] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1590185491] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1590185613] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15211536 kB) free!
[1590185661] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590186091] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1590186211] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1590186213] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394916 kB) free!
[1590186586] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.4% (13740800 kB) free!
[1590186608] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 20.1% (4923132 kB) free.
[1590189813] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34683400 kB) free.
[1590190208] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 20.0% (4910776 kB) free!
[1590190328] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 20.0% (4909272 kB) free!
[1590190413] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4917144 kB) free!
[1590190448] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 20.0% (4907344 kB) free!
[1590190533] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1126324 kB) free!
[1590190653] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (477912 kB) free!
[1590191253] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14796972 kB) free!
[1590191853] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402936 kB) free!
[1590192114] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1590192211] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1590192234] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1590192241] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 39.7% (9794636 kB) free.
[1590192277] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 53.1% (13015320 kB) free.
[1590192332] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1590192452] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1590192453] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.7% (15368332 kB) free!
[1590193052] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1590193053] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34192680 kB) free.
[1590193652] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1590193653] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12906640 kB) free!
[1590193772] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1590193773] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8730884 kB) free!
[1590193892] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1590193893] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (4607696 kB) free!
[1590194386] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407900 kB) free!
[1590194493] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (18400088 kB) free!
[1590195094] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34194884 kB) free.
[1590195269] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590195693] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1590195694] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13697668 kB) free!
[1590195705] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (2406636 kB) free!
[1590195765] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 30.2% (7412176 kB) free.
[1590195814] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (9924396 kB) free!
[1590195934] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (6209760 kB) free!
[1590196065] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 16.0% (3916812 kB) free!
[1590196125] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 16.0% (3917056 kB) free!
[1590196185] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 16.0% (3914504 kB) free!
[1590196245] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 16.0% (3914520 kB) free!
[1590196293] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1590196413] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1590196477] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197040 kB) free!
[1590196534] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.4% (18968352 kB) free!
[1590196597] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193824 kB) free!
[1590196717] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177112 kB) free!
[1590197041] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196720 kB) free!
[1590197134] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31657984 kB) free.
[1590197161] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190128 kB) free!
[1590197281] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179424 kB) free!
[1590197734] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10749252 kB) free!
[1590197854] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6870712 kB) free!
[1590197974] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (3029684 kB) free!
[1590197986] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 74.6% (98388840 kB) free.
[1590198574] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16395292 kB) free!
[1590198768] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1590198813] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1590198834] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1590198888] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1590198933] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1590198954] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1590199054] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1590199074] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1590199174] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30737356 kB) free.
[1590199375] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.7% (1735156 kB) free!
[1590199436] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/23 07:28:21 [error] 22135#0: *90162832 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590199556] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590199654] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1590199774] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9825484 kB) free!
[1590199894] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (5938528 kB) free!
[1590200014] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2107992 kB) free!
[1590200254] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1590200374] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1590200494] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1590200614] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15275828 kB) free!
[1590201094] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1590201214] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29076376 kB) free.
[1590201815] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6612196 kB) free!
[1590201935] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (2464384 kB) free!
[1590202055] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399816 kB) free!
[1590202294] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1590202414] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1590202535] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1590203193] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.00, 6.70, 5.08
[1590203255] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29728920 kB) free.
[1590203313] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.95, 6.02, 5.02
[1590203434] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.70, 4.85, 4.72
[1590203735] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1590203855] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6334880 kB) free!
[1590203975] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (1959248 kB) free!
[1590204095] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (418292 kB) free!
[1590204695] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13598064 kB) free!
[1590205233] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 13.82, 12.11, 8.36
[1590205295] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33652632 kB) free.
[1590205353] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 29.94, 18.74, 11.27
[1590205473] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 17.98, 18.25, 12.02
[1590205535] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 350 processes
[1590205655] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 352 processes
[1590205775] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1590205879] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1590205895] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7548904 kB) free!
[1590206000] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1590206015] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2846612 kB) free!
[1590206120] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1590206136] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (445972 kB) free!
[1590206720] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(489) < 103.255.146.164 [23/May/2020:09:32:02 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590206720] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1590206736] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.9% (20955752 kB) free!
[1590206840] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590206975] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1590207336] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (438156 kB) free!
[1590207575] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1590207695] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1590207815] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 354 processes
[1590207936] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21200928 kB) free!
[1590208536] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (421716 kB) free!
[1590208640] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(626) < 103.255.146.164 [23/May/2020:10:01:31 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590208760] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590208804] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 45.5% (11163960 kB) free.
[1590209136] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23052064 kB) free!
[1590209736] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410096 kB) free!
[1590209960] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(675) < 103.255.146.164 [23/May/2020:10:24:10 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590210080] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590210175] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1631056 kB) free!
[1590210273] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 60.92, 35.42, 22.71
[1590210732] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.5% (3713864 kB) free!
[1590210775] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.5% (1701116 kB) free!
[1590210852] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.2% (3187272 kB) free!
[1590210972] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.9% (2665124 kB) free!
[1590211070] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.6920.6267> (10.147.212.180:33470 -> 10.147.212.160:5672, vhost: '/', user: 'guest'), channel 1:
[1590211190] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590211204] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 14.5% (3554932 kB) free!
[1590211324] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 12.9% (3170708 kB) free!
[1590211444] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 11.3% (2766544 kB) free!
[1590211536] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30312048 kB) free.
[1590211572] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (193792 kB) free!
[1590212032] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9514224 kB) free!
[1590212044] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (927992 kB) free!
[1590212136] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4571304 kB) free!
[1590212152] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.6% (9180848 kB) free!
[1590212256] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (382068 kB) free!
[1590212272] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.0% (8864120 kB) free!
[1590212376] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390360 kB) free!
[1590212880] SERVICE ALERT: qservercluster2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 804 processes
[1590212936] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.25058.581> (10.147.212.158:43584 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590212976] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16711968 kB) free!
[1590213000] SERVICE ALERT: qservercluster2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 804 processes
[1590213056] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590213120] SERVICE ALERT: qservercluster2;Total Processes;WARNING;HARD;3;PROCS WARNING: 804 processes
[1590213576] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405624 kB) free!
[1590213869] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (1371460 kB) free!
[1590213957] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 18.2% (6378460 kB) free!
[1590213989] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (1313448 kB) free!
[1590214072] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4037880 kB) free!
[1590214077] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 16.7% (5857972 kB) free!
[1590214109] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (1327936 kB) free!
[1590214177] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32804700 kB) free.
[1590214197] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 15.3% (5367784 kB) free!
[1590214415] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1590214674] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 327 processes
[1590214777] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3645952 kB) free!
[1590214794] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1590214797] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2915364 kB) free!
[1590214880] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(615) < 103.255.146.164 [23/May/2020:11:46:13 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590214897] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (395292 kB) free!
[1590214987] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 15.5% (3834176 kB) free!
[1590215000] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590215017] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399428 kB) free!
[1590215107] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 13.9% (3426864 kB) free!
[1590215227] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 12.3% (3032768 kB) free!
[1590215309] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.2% (3778868 kB) free.
[1590215617] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.0% (24992972 kB) free!
[1590215756] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/23 12:02:58 [error] 22126#0: *90207797 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590215774] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.9% (15929468 kB) free.
[1590215827] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (1016548 kB) free!
[1590215876] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590215998] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.8% (31798936 kB) free.
[1590216175] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1566068 kB) free!
[1590216217] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409664 kB) free!
[1590216509] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3171588 kB) free!
[1590216629] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.0% (2926984 kB) free!
[1590216749] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (2763840 kB) free!
[1590217027] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.0% (22952600 kB) free.
[1590217078] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 36.72, 28.47, 21.65
[1590217415] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1590217535] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 352 processes
[1590217655] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 353 processes
[1590217920] SERVICE ALERT: qservercluster2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 854 processes
[1590217949] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1590217949] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3370176 kB) free.
[1590218549] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (3097144 kB) free!
[1590218617] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.7% (15448124 kB) free!
[1590218669] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (3029416 kB) free!
[1590218790] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.3% (2483412 kB) free!
[1590218994] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1590219114] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1590219175] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1590219175] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.4% (1690704 kB) free!
[1590219217] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (377208 kB) free!
[1590219381] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9063896 kB) free!
[1590219390] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.3% (3793892 kB) free.
[1590219455] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 351 processes
[1590219501] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.7% (8737564 kB) free!
[1590219621] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (8416496 kB) free!
[1590219817] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32328880 kB) free.
[1590219990] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.4% (2007940 kB) free!
[1590220110] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.6% (2048144 kB) free!
[1590220230] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.0% (2117744 kB) free!
[1590220320] SERVICE ALERT: qservercluster2;Total Processes;OK;HARD;1;PROCS OK: 718 processes
[1590220417] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (421804 kB) free!
[1590220537] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (411764 kB) free!
[1590220657] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399012 kB) free!
[1590220856] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.11409.706> (10.147.212.158:51488 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590220921] SERVICE ALERT: qservercluster2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 752 processes
[1590220976] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590221041] SERVICE ALERT: qservercluster2;Total Processes;OK;HARD;1;PROCS OK: 747 processes
[1590221076] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 17.35, 16.51, 18.51
[1590221255] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 352 processes
[1590221258] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35325664 kB) free.
[1590221421] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4577188 kB) free!
[1590221575] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1566624 kB) free!
[1590221648] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.9% (2425276 kB) free!
[1590221774] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4200164 kB) free!
[1590221858] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (1246464 kB) free!
[1590221894] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.5% (3699396 kB) free!
[1590221978] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (405092 kB) free!
[1590222014] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.5% (3266116 kB) free!
[1590222098] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (447172 kB) free!
[1590222200] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(886) < 103.255.146.164 [23/May/2020:13:44:56 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590222276] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 33.59, 21.54, 18.39
[1590222320] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590222455] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 352 processes
[1590222614] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (24.8% change >= 20.0% threshold)
[1590222614] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (1014332 kB) free!
[1590222630] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1213604 kB) free!
[1590222775] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.4% (1689252 kB) free!
[1590223230] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3584148 kB) free.
[1590223655] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 350 processes
[1590223830] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (1397724 kB) free!
[1590223950] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.3% (1677928 kB) free!
[1590224070] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.2% (1979236 kB) free!
[1590224255] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 351 processes
[1590224498] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17756004 kB) free!
[1590224855] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 349 processes
[1590225098] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (442308 kB) free!
[1590225175] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (1465236 kB) free!
[1590225270] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1155752 kB) free!
[1590225455] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1590225598] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 17.4% (6082420 kB) free!
[1590225718] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 16.2% (5681648 kB) free!
[1590225838] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 15.2% (5317192 kB) free!
[1590225870] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.1% (2454532 kB) free!
[1590226076] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/23 14:53:45 [error] 22130#0: *90251844 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590226196] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590226214] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.1% (15960404 kB) free.
[1590226375] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.5% (1706992 kB) free!
[1590226439] SERVICE FLAPPING ALERT: sig5-vm2;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1590226439] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.9% (31841788 kB) free.
[1590226470] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3684240 kB) free.
[1590226520] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(687) < 185.255.8.58 [23/May/2020:15:00:04 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1590226641] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590226655] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 352 processes
[1590226914] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1590227034] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1590227070] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3113988 kB) free!
[1590227076] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.41, 8.53, 19.16
[1590227191] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (3048676 kB) free!
[1590227255] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1590227278] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 8.73, 11.48, 17.78
[1590227311] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (3020372 kB) free!
[1590227498] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17497484 kB) free!
[1590228098] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409888 kB) free!
[1590228440] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(121) < 103.255.146.164 [23/May/2020:15:31:28 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590228455] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 349 processes
[1590228560] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590228698] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21441728 kB) free!
[1590228776] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.28938.736> (10.147.212.158:55548 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590228860] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1662 processes
[1590228896] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590229055] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1590229100] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 659 processes
[1590229299] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (458528 kB) free!
[1590229460] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1839 processes
[1590229476] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1590229476] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 1.90, 2.38, 3.86
[1590229700] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 735 processes
[1590229820] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 740 processes
[1590229899] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30687116 kB) free.
[1590229940] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 736 processes
[1590230255] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 350 processes
[1590230499] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (387592 kB) free!
[1590230619] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (473356 kB) free!
[1590230676] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.65, 7.35, 5.13
[1590230739] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (431864 kB) free!
[1590230796] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.21, 7.00, 5.28
[1590230855] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 350 processes
[1590230916] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.13, 7.44, 5.67
[1590230960] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(447) < 103.255.146.164 [23/May/2020:16:12:04 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590231080] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590231339] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26115560 kB) free!
[1590231455] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 351 processes
[1590231939] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (411972 kB) free!
[1590232055] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 352 processes
[1590232111] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.5% (3649368 kB) free.
[1590232245] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.9% (209524 kB) free!
[1590232496] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.18276.737> (10.147.212.158:58436 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590232616] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590232711] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (469312 kB) free!
[1590232832] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.3% (3792892 kB) free.
[1590232975] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1605696 kB) free!
[1590233255] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 351 processes
[1590233431] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.0% (1942596 kB) free!
[1590233551] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.3% (4594680 kB) free.
[1590233575] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (1674752 kB) free!
[1590233878] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 34.25, 29.80, 20.59
[1590233916] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 25.82, 26.13, 16.70
[1590234014] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.0% (3152916 kB) free!
[1590234134] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.9% (2680552 kB) free!
[1590234255] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.1% (2276624 kB) free!
[1590234340] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17613224 kB) free!
[1590234478] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 13.00, 16.97, 18.53
[1590234516] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.02, 11.58, 14.02
[1590234751] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.8% (2405636 kB) free!
[1590234855] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (450448 kB) free!
[1590234871] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.8% (5013100 kB) free.
[1590234940] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387996 kB) free!
[1590235055] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 353 processes
[1590235375] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1413212 kB) free!
[1590235472] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (1545216 kB) free!
[1590235540] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33825044 kB) free.
[1590235592] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (2983704 kB) free!
[1590235712] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.2% (4423480 kB) free.
[1590235975] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.4% (1694160 kB) free!
[1590236052] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 455 processes
[1590236140] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (391460 kB) free!
[1590236260] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (443616 kB) free!
[1590236312] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (1587984 kB) free!
[1590236380] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388440 kB) free!
[1590236432] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.5% (2839452 kB) free!
[1590236552] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174512 kB) free!
[1590236816] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.10656.744> (10.147.212.158:33010 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590236878] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1590236878] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 1.12, 1.38, 4.85
[1590236937] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590236980] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29025684 kB) free.
[1590237152] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (3045412 kB) free!
[1590237255] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.9% (16144356 kB) free.
[1590237455] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 351 processes
[1590237580] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (396936 kB) free!
[1590237700] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (400308 kB) free!
[1590237752] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.8% (4025372 kB) free.
[1590237820] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (432364 kB) free!
[1590237834] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1590237839] SERVICE FLAPPING ALERT: sig5-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590237954] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1590238056] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 353 processes
[1590238196] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/23 18:11:22 [error] 22134#0: *90293152 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590238316] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590238352] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.9% (2095668 kB) free!
[1590238420] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31671464 kB) free.
[1590238472] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.7% (4020332 kB) free.
[1590239020] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (436392 kB) free!
[1590239072] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (564560 kB) free!
[1590239140] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (422972 kB) free!
[1590239193] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.9% (2415856 kB) free!
[1590239256] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 352 processes
[1590239260] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400200 kB) free!
[1590239313] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.3% (4119764 kB) free.
[1590239516] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/23 18:36:02 [error] 22129#0: *90297528 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590239636] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590239856] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 352 processes
[1590239860] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33618228 kB) free.
[1590239913] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (1037400 kB) free!
[1590240033] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (3074960 kB) free!
[1590240153] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.3% (4270788 kB) free.
[1590240456] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 351 processes
[1590240460] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (435900 kB) free!
[1590240581] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (385540 kB) free!
[1590240701] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (456824 kB) free!
[1590240753] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (1471144 kB) free!
[1590240873] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3750896 kB) free.
[1590241301] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36560616 kB) free.
[1590241375] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1601348 kB) free!
[1590241473] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (452792 kB) free!
[1590241593] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.5% (2188664 kB) free!
[1590241713] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.5% (2511412 kB) free!
[1590241901] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (1879952 kB) free!
[1590242021] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (407640 kB) free!
[1590242141] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (419056 kB) free!
[1590242314] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.8% (4038556 kB) free.
[1590242575] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (1681152 kB) free!
[1590242914] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (844544 kB) free!
[1590243034] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.5% (2523284 kB) free!
[1590243154] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.4% (4448844 kB) free.
[1590243754] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (1181352 kB) free!
[1590243874] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (2975624 kB) free!
[1590243994] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.9% (4370624 kB) free.
[1590244057] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 353 processes
[1590244541] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.4% (18955004 kB) free!
[1590244594] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (1566312 kB) free!
[1590244657] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 350 processes
[1590244714] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3575732 kB) free.
[1590245141] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (378576 kB) free!
[1590245154] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 336 processes
[1590245274] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1590245575] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1614144 kB) free!
[1590245655] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.7% (3984916 kB) free!
[1590245741] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35539948 kB) free.
[1590245775] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.7% (3748452 kB) free!
[1590245857] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 353 processes
[1590245895] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.6% (3507524 kB) free!
[1590245914] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.0% (2443104 kB) free!
[1590245916] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590246035] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.2% (4098368 kB) free.
[1590246175] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.2% (1664608 kB) free!
[1590246341] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (936968 kB) free!
[1590246457] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 352 processes
[1590246461] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (375036 kB) free!
[1590246478] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1590246582] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397056 kB) free!
[1590246634] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (512252 kB) free!
[1590246754] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.7% (2065692 kB) free!
[1590246775] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1614184 kB) free!
[1590246845] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 11.2% (2737596 kB) free!
[1590246874] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.7% (3857356 kB) free.
[1590247095] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (1138144 kB) free!
[1590247182] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32941076 kB) free.
[1590247445] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1773192 kB) free!
[1590247657] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1590247674] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1590247782] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (406904 kB) free!
[1590247794] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1590247902] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (386616 kB) free!
[1590247975] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.4% (1686820 kB) free!
[1590248022] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388592 kB) free!
[1590248074] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.0% (2274956 kB) free!
[1590248088] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 370 processes
[1590248194] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.0% (3738884 kB) free.
[1590248208] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 369 processes
[1590248257] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 349 processes
[1590248295] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.3% (16015800 kB) free.
[1590248316] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.11, 2.86, 4.11
[1590248575] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1622820 kB) free!
[1590248622] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29355492 kB) free.
[1590248794] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.9% (1934100 kB) free!
[1590248858] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1590248914] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3476580 kB) free.
[1590249222] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (950212 kB) free!
[1590249342] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (394964 kB) free!
[1590249462] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (423852 kB) free!
[1590249775] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (1680528 kB) free!
[1590250058] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 351 processes
[1590250062] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16261860 kB) free!
[1590250115] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (1362164 kB) free!
[1590250235] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.0% (2766732 kB) free!
[1590250355] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.5% (3502148 kB) free.
[1590250662] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (424728 kB) free!
[1590250955] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (179260 kB) free!
[1590251075] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (1555468 kB) free!
[1590251195] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.7% (2875876 kB) free!
[1590251262] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32196700 kB) free.
[1590251795] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3676648 kB) free.
[1590251858] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 351 processes
[1590251862] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4902776 kB) free!
[1590251982] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (551104 kB) free!
[1590252102] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406404 kB) free!
[1590252458] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 351 processes
[1590252702] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21275984 kB) free!
[1590253303] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403488 kB) free!
[1590253461] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1658 processes
[1590253658] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1590253903] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27882240 kB) free.
[1590254258] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 351 processes
[1590254503] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5523200 kB) free!
[1590254623] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1558932 kB) free!
[1590254743] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403396 kB) free!
[1590255343] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18021412 kB) free!
[1590255395] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.7% (2554876 kB) free!
[1590255459] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 350 processes
[1590255515] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3269620 kB) free.
[1590255943] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (404152 kB) free!
[1590256059] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 350 processes
[1590256115] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.5% (2028820 kB) free!
[1590256236] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.4% (2660040 kB) free!
[1590256356] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (3187748 kB) free!
[1590256543] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21547516 kB) free!
[1590257143] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (736448 kB) free!
[1590257259] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 350 processes
[1590257743] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.2% (22646500 kB) free!
[1590257859] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 352 processes
[1590258343] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1781468 kB) free!
[1590258459] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 354 processes