[1581013884] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.5% (29313576 kB) free.
[1581013959] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7950048 kB) free!
[1581014079] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31667452 kB) free.
[1581014230] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181920 kB) free!
[1581014350] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179004 kB) free!
[1581014398] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180568 kB) free!
[1581014470] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198332 kB) free!
[1581014502] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184436 kB) free!
[1581014518] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170988 kB) free!
[1581014623] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190580 kB) free!
[1581014638] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179256 kB) free!
[1581014679] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7903468 kB) free!
[1581014710] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 406 processes
[1581014743] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172612 kB) free!
[1581014799] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31563360 kB) free.
[1581014830] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1581015399] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8671288 kB) free!
[1581015519] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31517572 kB) free.
[1581015870] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176080 kB) free!
[1581015930] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 66.9% (16387264 kB) free.
[1581016119] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8858000 kB) free!
[1581016239] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31435880 kB) free.
[1581016301] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(147) < 2020/02/07 00:37:13 [error] 10221#0: *49991574 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1581016421] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581016580] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 103.16.101.72 [07/Feb/2020:00:36:52 +0530] TCP 502 0 0 60.002 "10.147.212.38:4000, 10.147.212.61:4000, 10.147.212.37:4000, 10.147.212.49:4000" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 0.000, 0.000, 60.001"
[1581016701] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581016839] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8999316 kB) free!
[1581016959] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31379312 kB) free.
[1581017559] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9100248 kB) free!
[1581017639] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 23.5% (5799460 kB) free.
[1581017679] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31267344 kB) free.
[1581017901] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 103.16.101.72 [07/Feb/2020:01:06:12 +0530] TCP 502 0 0 60.001 "10.147.212.38:4000, 10.147.212.61:4000, backend3" "0, 0, 0" "0, 0, 0" "0.000, 60.001, 0.000"
[1581018022] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581018221] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(80) < 2020/02/07 01:07:10 [error] 10219#0: *50002753 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:46/0, bytes from/to upstream:0/46
[1581018239] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184184 kB) free!
[1581018279] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8654032 kB) free!
[1581018341] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581018359] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186972 kB) free!
[1581018399] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26736660 kB) free.
[1581018479] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198212 kB) free!
[1581022271] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.0% (10054012 kB) free.
[1581022429] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (234320 kB) free!
[1581022549] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (246380 kB) free!
[1581022669] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (244852 kB) free!
[1581023230] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1581023269] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 81.8% (40299272 kB) free.
[1581023350] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1581024623] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(154) < 103.16.101.72 [07/Feb/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1581024708] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2675) < 172.105.59.173 [07/Feb/2020:03:01:41 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1581024744] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581024828] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(2) < 203.212.70.170 [07/Feb/2020:03:02:31 +0530] TCP 502 0 0 120.000 "10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1581024941] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(288) < 2020/02/07 03:01:32 [error] 10210#0: *50036779 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581024948] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581025061] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581025751] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1581025871] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1581028661] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(23) < 2020/02/07 04:00:48 [error] 10210#0: *50052669 connect() failed (111: Connection refused) while connecting to upstream, client: 124.153.112.226, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581028781] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581031872] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180460 kB) free!
[1581031992] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191820 kB) free!
[1581032113] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195820 kB) free!
[1581032879] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 48.7% (12014832 kB) free.
[1581033343] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.6% (20246852 kB) free.
[1581036261] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.8% (39799388 kB) free.
[1581036313] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.5% (21929424 kB) free.
[1581036369] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 36.7% (48434792 kB) free.
[1581036888] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2458728 kB) free!
[1581037679] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (1456116 kB) free!
[1581037800] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (1443088 kB) free!
[1581037920] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (1435312 kB) free!
[1581039285] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.3% (21394172 kB) free.
[1581039319] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.8% (21407464 kB) free.
[1581041667] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 977.42, 351.26, 127.35
[1581041750] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(14) < 203.212.70.170 [07/Feb/2020:07:45:45 +0530] TCP 502 0 0 120.001 "10.147.212.143:3001, 10.147.212.142:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 60.001, 60.000, 0.000"
[1581041787] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 132.45, 235.24, 111.98
[1581041870] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(17) < 203.212.70.170 [07/Feb/2020:07:46:45 +0530] TCP 502 0 0 180.000 "10.147.212.155:3001, 10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 0.000"
[1581041927] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1581041990] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581042047] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1581043181] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(56) < 2020/02/07 08:01:44 [error] 10216#0: *50129023 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581043301] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581044501] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(171) < 2020/02/07 08:25:39 [error] 10216#0: *50137973 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581044621] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581049871] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1581049991] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1581050688] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.3% (4736816 kB) free!
[1581050743] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (1738792 kB) free!
[1581050809] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187176 kB) free!
[1581050811] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1581050863] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192476 kB) free!
[1581050929] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170044 kB) free!
[1581050931] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1581050983] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199308 kB) free!
[1581051052] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1581051590] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(705) < 103.255.146.164 [07/Feb/2020:10:26:02 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1581051710] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581052129] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 82.2% (20154584 kB) free.
[1581052910] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(709) < 103.255.146.164 [07/Feb/2020:10:46:32 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1581053030] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581053127] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168620 kB) free!
[1581053247] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188476 kB) free!
[1581053367] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182000 kB) free!
[1581053801] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20947072 kB) free!
[1581053921] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.2% (16127720 kB) free!
[1581054041] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41572284 kB) free.
[1581054641] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16508896 kB) free!
[1581054761] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37777432 kB) free.
[1581054988] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.2% (3859964 kB) free!
[1581055109] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.7% (3526824 kB) free!
[1581055229] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.2% (3197432 kB) free!
[1581055829] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1426196 kB) free!
[1581056598] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1581056719] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 706 processes
[1581056839] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1581057029] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.0% (15718040 kB) free.
[1581057064] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 18.2% (4493200 kB) free!
[1581057171] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.8% (4859164 kB) free!
[1581057184] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 16.6% (4089684 kB) free!
[1581057231] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(9) < 203.212.70.170 [07/Feb/2020:12:02:17 +0530] TCP 502 0 0 60.001 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.001, 0.000"
[1581057291] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.7% (4826656 kB) free!
[1581057304] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 14.7% (3632652 kB) free!
[1581057351] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581057411] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.6% (4807528 kB) free!
[1581057439] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 691 processes
[1581057904] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.2% (22769232 kB) free.
[1581059068] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9794756 kB) free!
[1581059188] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (9792936 kB) free!
[1581059308] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (9784776 kB) free!
[1581060470] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9530152 kB) free!
[1581060590] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9089676 kB) free!
[1581060710] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.7% (8742564 kB) free!
[1581060747] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 41.35, 39.28, 29.12
[1581060761] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14820488 kB) free!
[1581060867] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 31.52, 36.42, 29.32
[1581060881] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39220388 kB) free.
[1581060885] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (273888 kB) free!
[1581060917] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (469500 kB) free!
[1581060987] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 30.48, 33.95, 29.25
[1581061005] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (274496 kB) free!
[1581061038] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165500 kB) free!
[1581061125] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (273028 kB) free!
[1581061158] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167904 kB) free!
[1581061481] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13907424 kB) free!
[1581061551] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(53) < 103.16.101.72 [07/Feb/2020:13:11:59 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1581061601] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40143168 kB) free.
[1581061671] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581062201] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13856972 kB) free!
[1581062321] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9751244 kB) free!
[1581062329] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.1% (4426000 kB) free!
[1581062441] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33032844 kB) free.
[1581062447] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1581062450] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.0% (4174476 kB) free!
[1581062567] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1581062570] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.4% (4014164 kB) free!
[1581062839] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1581062908] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3824060 kB) free!
[1581062959] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1581063041] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7951708 kB) free!
[1581063110] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3243480 kB) free!
[1581063162] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.4% (47976352 kB) free.
[1581063761] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (22243556 kB) free!
[1581063881] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.5% (48151748 kB) free.
[1581064011] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.4% (2314964 kB) free!
[1581064230] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.1% (3166032 kB) free!
[1581064350] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.1% (2720588 kB) free!
[1581064370] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1791016 kB) free!
[1581064470] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.1% (2271744 kB) free!
[1581064481] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22170384 kB) free!
[1581064601] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (18117268 kB) free!
[1581064721] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.9% (45998332 kB) free.
[1581065070] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (25.3% change >= 20.0% threshold)
[1581065070] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (474860 kB) free!
[1581065321] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20029412 kB) free!
[1581065441] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (15991416 kB) free!
[1581065561] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42162240 kB) free.
[1581066084] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9717956 kB) free!
[1581066161] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15761284 kB) free!
[1581066204] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.7% (9240064 kB) free!
[1581066281] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.1% (42317700 kB) free.
[1581066324] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.8% (8787220 kB) free!
[1581066881] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15980424 kB) free!
[1581066904] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4776548 kB) free!
[1581067001] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41529216 kB) free.
[1581067024] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4433396 kB) free!
[1581067144] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 16.6% (4108992 kB) free!
[1581067524] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4845252 kB) free!
[1581067601] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17591984 kB) free!
[1581067722] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (13134716 kB) free!
[1581067842] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40292980 kB) free.
[1581068070] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.9% (15692512 kB) free.
[1581068125] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.5% (29817552 kB) free.
[1581068227] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9569420 kB) free!
[1581068271] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(136) < 92.42.104.166 [07/Feb/2020:15:00:13 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1581068344] SERVICE FLAPPING ALERT: db1-vm2;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1581068344] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1271660 kB) free!
[1581068347] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (9416340 kB) free!
[1581068391] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581068442] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13584320 kB) free!
[1581068467] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (9254972 kB) free!
[1581068562] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (9616252 kB) free!
[1581068671] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 81.2% (19896824 kB) free.
[1581068682] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31491748 kB) free.
[1581068717] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 74.9% (18494952 kB) free.
[1581068747] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 67.9% (33481148 kB) free.
[1581068769] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 22.7% (11162568 kB) free.
[1581068800] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 64.5% (15813288 kB) free.
[1581068828] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 81.3% (19919736 kB) free.
[1581068867] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.5% (14224024 kB) free.
[1581068900] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 82.8% (20301204 kB) free.
[1581068944] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.4% (22807708 kB) free.
[1581068992] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2998376 kB) free!
[1581069112] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (2908944 kB) free!
[1581069232] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (2889664 kB) free!
[1581069282] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (1193168 kB) free!
[1581069402] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393652 kB) free!
[1581069522] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33422948 kB) free.
[1581069832] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.0% (3904864 kB) free.
[1581069919] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191680 kB) free!
[1581070008] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176836 kB) free!
[1581070037] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (2420460 kB) free!
[1581070039] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181448 kB) free!
[1581070122] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (2506364 kB) free!
[1581070128] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197396 kB) free!
[1581070157] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (2325436 kB) free!
[1581070159] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189224 kB) free!
[1581070242] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (391184 kB) free!
[1581070248] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191388 kB) free!
[1581070277] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2225452 kB) free!
[1581070362] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39403424 kB) free.
[1581070791] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(183) < 92.42.104.166 [07/Feb/2020:15:42:41 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1581070911] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581070962] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6277332 kB) free!
[1581071082] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (2442524 kB) free!
[1581071202] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38951972 kB) free.
[1581071330] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.5% (21943444 kB) free.
[1581071359] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.9% (21452616 kB) free.
[1581071392] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.9% (22290856 kB) free.
[1581071429] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 88.3% (43510392 kB) free.
[1581071511] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(247) < 103.255.146.164 [07/Feb/2020:16:01:16 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1581071621] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(68) < 2020/02/07 16:01:17 [error] 10218#0: *50304270 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581071632] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(60) < 182.18.144.246 [07/Feb/2020:16:03:22 +0530] TCP 502 0 0 120.002 "10.147.212.143:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.001, 0.000"
[1581071742] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581071752] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581071802] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6756772 kB) free!
[1581071923] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2833388 kB) free!
[1581072043] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403676 kB) free!
[1581072349] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9129592 kB) free!
[1581072469] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9129292 kB) free!
[1581072589] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.5% (9129180 kB) free!
[1581072832] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.0% (1951032 kB) free!
[1581072952] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.7% (2062708 kB) free!
[1581073072] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.8% (2240580 kB) free!
[1581073168] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.5% (3580672 kB) free!
[1581073243] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16814384 kB) free!
[1581073288] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (2250760 kB) free!
[1581073408] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1581073408] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1689260 kB) free!
[1581073843] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34768772 kB) free.
[1581074077] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.8% (21904176 kB) free.
[1581074098] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.8% (39801032 kB) free.
[1581074135] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1581074135] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 54.3% (8826640 kB) free.
[1581074443] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8620832 kB) free!
[1581074563] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36299596 kB) free.
[1581075163] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8423584 kB) free!
[1581075283] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4349096 kB) free!
[1581075403] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38658888 kB) free.
[1581075901] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172256 kB) free!
[1581075938] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (3013728 kB) free!
[1581076003] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12432488 kB) free!
[1581076021] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185380 kB) free!
[1581076058] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3403552 kB) free.
[1581076123] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41949768 kB) free.
[1581076141] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187868 kB) free!
[1581076723] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15818248 kB) free!
[1581076843] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (11255948 kB) free!
[1581076963] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40170556 kB) free.
[1581077070] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1581077258] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.3% (2483464 kB) free!
[1581077378] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.5% (2841948 kB) free!
[1581077498] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.9% (3242648 kB) free!
[1581077563] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14813084 kB) free!
[1581077684] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (10917320 kB) free!
[1581077804] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39123040 kB) free.
[1581078098] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.6% (4325472 kB) free.
[1581078261] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.4% (21815644 kB) free.
[1581078404] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13408060 kB) free!
[1581078524] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9255368 kB) free!
[1581078644] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37195244 kB) free.
[1581078698] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.1% (2131208 kB) free!
[1581078818] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.7% (2389356 kB) free!
[1581078909] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9314520 kB) free!
[1581078938] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (2782860 kB) free!
[1581079029] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (8916028 kB) free!
[1581079149] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.4% (8584024 kB) free!
[1581079198] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172312 kB) free!
[1581079244] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9450904 kB) free!
[1581079318] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179900 kB) free!
[1581079364] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.6% (23229072 kB) free!
[1581079438] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195244 kB) free!
[1581079484] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35727988 kB) free.
[1581079538] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.0% (3898304 kB) free.
[1581080038] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1581080038] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.9% (21062616 kB) free.
[1581080084] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (8557584 kB) free!
[1581080204] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4476932 kB) free!
[1581080324] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34426080 kB) free.
[1581080344] SERVICE FLAPPING ALERT: db1-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581080924] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8425624 kB) free!
[1581081044] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4429396 kB) free!
[1581081164] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34406736 kB) free.
[1581081764] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7189136 kB) free!
[1581081885] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (3073172 kB) free!
[1581081938] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.0% (2608416 kB) free!
[1581082005] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34002308 kB) free.
[1581082058] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.7% (2720908 kB) free!
[1581082178] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (2584692 kB) free!
[1581082605] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6592156 kB) free!
[1581082725] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2372236 kB) free!
[1581082845] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33893588 kB) free.
[1581083349] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4477444 kB) free!
[1581083445] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5615304 kB) free!
[1581083565] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1338864 kB) free!
[1581083668] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (4669496 kB) free!
[1581083685] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34343124 kB) free.
[1581083788] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.9% (3927116 kB) free!
[1581083908] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.5% (3579820 kB) free!
[1581083978] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.1% (4240456 kB) free.
[1581084285] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5912660 kB) free!
[1581084405] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1742636 kB) free!
[1581084508] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2208744 kB) free!
[1581084525] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34603388 kB) free.
[1581084767] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1581084887] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1581085125] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7467020 kB) free!
[1581085245] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3242420 kB) free!
[1581085365] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395580 kB) free!
[1581085487] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1581085607] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 331 processes
[1581085727] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 332 processes
[1581085759] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1581085879] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1581085965] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (24.7% change >= 20.0% threshold)
[1581085965] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13271720 kB) free!
[1581085999] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1581086152] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(26) < 203.212.70.170 [07/Feb/2020:20:03:46 +0530] TCP 502 0 0 183.008 "10.147.212.155:3001, 10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.142:3001" "0, 0, 0, 0" "0, 0, 0, 0" "3.008, 60.000, 60.000, 60.000"
[1581086272] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581086327] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1581087165] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33874572 kB) free.
[1581087765] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10259672 kB) free!
[1581087799] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1581087885] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6494788 kB) free!
[1581087919] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1581088005] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33727452 kB) free.
[1581088605] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10255440 kB) free!
[1581088725] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6423608 kB) free!
[1581088845] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33455468 kB) free.
[1581089120] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1581089240] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1581089360] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 705 processes
[1581089445] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6408116 kB) free!
[1581089565] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2162532 kB) free!
[1581089685] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395256 kB) free!
[1581089938] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.3% (4239644 kB) free!
[1581089960] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1581090015] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1581090051] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1581090059] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191508 kB) free!
[1581090136] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1581090172] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1581090179] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192240 kB) free!
[1581090256] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1581090277] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187620 kB) free!
[1581090292] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1581090398] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183288 kB) free!
[1581090518] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194996 kB) free!
[1581090575] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 82.6% (20246720 kB) free.
[1581090600] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.4% (19837608 kB) free.
[1581090635] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.0% (50074884 kB) free.
[1581090666] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 55.8% (27503752 kB) free.
[1581090691] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 67.1% (16449608 kB) free.
[1581090856] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1581090856] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1581090892] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1581090976] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1581091096] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1581091160] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 705 processes
[1581091280] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 708 processes
[1581091401] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes
[1581091438] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581091456] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1581091576] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 344 processes
[1581091696] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1581091696] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1581091810] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171748 kB) free!
[1581091898] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165388 kB) free!
[1581091930] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163516 kB) free!
[1581092001] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1581092001] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1581092018] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185408 kB) free!
[1581092050] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183024 kB) free!
[1581092138] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177976 kB) free!
[1581093496] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1581093496] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1581093801] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1581093921] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1581094042] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1581094178] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1581094641] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1581094761] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 705 processes
[1581094881] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 706 processes
[1581096081] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1581096682] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1581096802] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1581096922] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1581098725] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8250876 kB) free!
[1581098741] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/07 23:27:18 [error] 10210#0: *50486063 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1581098845] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (967504 kB) free!
[1581098861] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581098965] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (232292 kB) free!
[1581099922] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1581099969] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (391608 kB) free!
[1581100042] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1581100090] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (1472960 kB) free!
[1581100162] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes