[1554921807] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.20, 0.39, 14.67
[1554922121] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.52, 0.63, 3.10
[1554922584] SERVICE FLAPPING ALERT: db2;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1554923007] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.34, 0.27, 4.18
[1554923236] SERVICE ALERT: db1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923315] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923326] SERVICE ALERT: db1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923366] SERVICE ALERT: db1;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923384] SERVICE ALERT: db1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923416] SERVICE ALERT: db1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923445] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923456] SERVICE ALERT: db1;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923496] SERVICE ALERT: db1;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923509] SERVICE ALERT: db1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923514] SERVICE ALERT: db1;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923530] SERVICE ALERT: db1;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923546] SERVICE ALERT: db1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923566] SERVICE ALERT: db1;Mariadb status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923575] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923586] SERVICE ALERT: db1;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923639] SERVICE ALERT: db1;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923644] SERVICE ALERT: db1;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923672] SERVICE ALERT: db1;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923676] SERVICE ALERT: db1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923696] SERVICE ALERT: db1;Mariadb status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923769] SERVICE ALERT: db1;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923784] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44883 MB (19% inode=99%):
[1554923802] SERVICE ALERT: db1;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923826] SERVICE ALERT: db1;Mariadb status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923905] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56421 MB (25% inode=99%):
[1554923932] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1554923956] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1554924504] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44621 MB (19% inode=99%):
[1554924624] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44621 MB (19% inode=99%):
[1554924744] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56437 MB (25% inode=99%):
[1554924911] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1554925511] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1554925631] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1554925751] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1554926256] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2479600 kB) free!
[1554926544] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46508 MB (20% inode=99%):
[1554926644] SERVICE ALERT: db1;Total Processes;OK;HARD;1;PROCS OK: 345 processes
[1554926664] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46508 MB (20% inode=99%):
[1554926676] SERVICE ALERT: db1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1554926784] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46516 MB (20% inode=99%):
[1554926826] SERVICE ALERT: db1;Mariadb status;OK;HARD;1;PROCS OK: 1 process with command name 'mysqld'
[1554926932] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 60374 MB (28% inode=99%):
[1554926956] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.5% (5154276 kB) free!
[1554927096] SERVICE ALERT: db1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 46436 MB (90% inode=99%):
[1554927130] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1554927138] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.2% (11450656 kB) free.
[1554927176] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1554927186] SERVICE ALERT: db1;Sensors Status;OK;HARD;1;SENSORS OK
[1554927244] SERVICE ALERT: db1;Total Processes;CRITICAL;SOFT;1;(Return code of 255 for service 'Total Processes' on host 'db1' was out of bounds)
[1554927276] SERVICE ALERT: db1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'db1' was out of bounds)
[1554927367] SERVICE ALERT: db1;Total Processes;CRITICAL;SOFT;2;(Return code of 255 for service 'Total Processes' on host 'db1' was out of bounds)
[1554927384] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48239 MB (21% inode=99%):
[1554927399] SERVICE ALERT: db1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'db1' was out of bounds)
[1554927429] SERVICE ALERT: db1;Mariadb status;CRITICAL;SOFT;1;(Return code of 255 for service 'Mariadb status' on host 'db1' was out of bounds)
[1554927432] HOST ALERT: db1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.47)
[1554927490] SERVICE ALERT: db1;Total Processes;CRITICAL;HARD;2;(Return code of 255 for service 'Total Processes' on host 'db1' was out of bounds)
[1554927522] SERVICE ALERT: db1;LocalService-NTPD SERVICE;CRITICAL;HARD;2;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'db1' was out of bounds)
[1554927533] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'db1' was out of bounds)
[1554927552] SERVICE ALERT: db1;Mariadb status;CRITICAL;HARD;1;(Return code of 255 for service 'Mariadb status' on host 'db1' was out of bounds)
[1554927556] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (225220 kB) free!
[1554927697] SERVICE ALERT: db1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'db1' was out of bounds)
[1554927733] SERVICE FLAPPING ALERT: db1;Current Users;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1554927733] SERVICE ALERT: db1;Current Users;CRITICAL;HARD;3;(Return code of 255 for service 'Current Users' on host 'db1' was out of bounds)
[1554927740] SERVICE ALERT: db1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'db1' was out of bounds)
[1554927778] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'db1' was out of bounds)
[1554927789] SERVICE ALERT: db1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'db1' was out of bounds)
[1554927991] HOST ALERT: db1;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.47)
[1554928156] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.1% (4981604 kB) free!
[1554928757] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1554928757] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (220376 kB) free!
[1554929288] SERVICE FLAPPING ALERT: db1;Total Processes;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1554929288] SERVICE ALERT: db1;Total Processes;OK;HARD;1;PROCS OK: 293 processes
[1554929292] HOST ALERT: db1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1554929319] SERVICE FLAPPING ALERT: db1;LocalService-NTPD SERVICE;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1554929319] SERVICE ALERT: db1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1554929453] SERVICE FLAPPING ALERT: db1;HDD_Home status;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1554929453] SERVICE ALERT: db1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 60438 MB (28% inode=99%):
[1554929530] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1554929617] SERVICE FLAPPING ALERT: db1;HDD_Root status;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1554929617] SERVICE ALERT: db1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 46175 MB (90% inode=99%):
[1554929660] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.3% (6051396 kB) free!
[1554929698] SERVICE FLAPPING ALERT: db1;CPU_Procs;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1554929698] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 280 processes
[1554929709] SERVICE FLAPPING ALERT: db1;Sensors Status;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1554929709] SERVICE ALERT: db1;Sensors Status;OK;HARD;1;SENSORS OK
[1554929769] SERVICE ALERT: db1;CPU Load;OK;HARD;1;OK - load average: 1.00, 0.87, 0.48
[1554930260] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (219472 kB) free!
[1554930380] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;3;CRITICAL - 0.5% (223992 kB) free!
[1554930500] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;3;CRITICAL - 0.4% (220080 kB) free!
[1554930620] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;3;CRITICAL - 0.5% (245924 kB) free!
[1554930741] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 34.3% (16924416 kB) free.
[1554930984] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46510 MB (20% inode=99%):
[1554931055] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2455368 kB) free!
[1554931104] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42947 MB (19% inode=99%):
[1554931224] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56430 MB (25% inode=99%):
[1554931341] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (262972 kB) free!
[1554931461] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (966344 kB) free!
[1554931581] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1554931581] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3254944 kB) free!
[1554932351] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1554932358] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.5% (5187460 kB) free!
[1554932950] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 368.16, 132.31, 48.06
[1554932951] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1554932958] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (213588 kB) free!
[1554933070] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 435.44, 236.13, 96.32
[1554933071] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554933190] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 444.24, 305.39, 138.65
[1554933191] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1554933191] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554933230] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 260.31, 74.06, 26.40
[1554933274] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 12.0% (2932000 kB) free!
[1554933350] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 409.82, 192.82, 75.70
[1554933470] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 429.32, 272.13, 118.96
[1554933558] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.6% (5201164 kB) free!
[1554934158] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (227300 kB) free!
[1554934655] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.1% (2486960 kB) free!
[1554935255] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2462580 kB) free!
[1554935424] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44128 MB (19% inode=99%):
[1554935544] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56419 MB (25% inode=99%):
[1554935590] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.65, 0.87, 13.81
[1554935591] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1554936191] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554936311] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554936431] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554936470] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.02, 1.79, 13.36
[1554936790] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.64, 0.72, 4.27
[1554937671] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.00, 1.73, 4.91
[1554939053] SERVICE FLAPPING ALERT: db1;HDD_Home status;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1554939144] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44121 MB (19% inode=99%):
[1554939217] SERVICE FLAPPING ALERT: db1;HDD_Root status;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1554939265] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55390 MB (24% inode=99%):
[1554939298] SERVICE FLAPPING ALERT: db1;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1554939309] SERVICE FLAPPING ALERT: db1;Sensors Status;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1554939488] SERVICE FLAPPING ALERT: db1;Total Processes;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1554939519] SERVICE FLAPPING ALERT: db1;LocalService-NTPD SERVICE;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1554939730] SERVICE FLAPPING ALERT: db1;Current Users;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1554939864] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46167 MB (20% inode=99%):
[1554939984] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46168 MB (20% inode=99%):
[1554940031] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1554940104] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56152 MB (24% inode=99%):
[1554940631] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554940751] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554940871] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554941904] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46167 MB (20% inode=99%):
[1554942024] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56407 MB (25% inode=99%):
[1554942875] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2450268 kB) free!
[1554942982] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1554943758] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.6% (5213964 kB) free!
[1554944358] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224240 kB) free!
[1554945558] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.5% (5170992 kB) free!
[1554946158] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (228440 kB) free!
[1554947424] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45317 MB (20% inode=99%):
[1554947471] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1554947544] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55908 MB (24% inode=99%):
[1554948072] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554948192] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554948312] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554948531] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 9.6% (2358860 kB) free!
[1554948548] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 482 processes
[1554948608] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 481 processes
[1554948744] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44838 MB (19% inode=99%):
[1554948864] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45740 MB (20% inode=99%):
[1554948984] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45743 MB (20% inode=99%):
[1554949581] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 36.2% (17839240 kB) free.
[1554949584] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49100 MB (21% inode=99%):
[1554950184] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47012 MB (20% inode=99%):
[1554950304] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47012 MB (20% inode=99%):
[1554950321] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 460.90, 344.81, 165.24
[1554950424] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46117 MB (20% inode=99%):
[1554950441] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 170.14, 299.66, 172.36
[1554950561] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 23.95, 200.89, 151.61
[1554951024] SERVICE FLAPPING ALERT: db2;HDD_Home status;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1554951024] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48090 MB (21% inode=99%):
[1554951381] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.9% (8319028 kB) free!
[1554951501] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 14.2% (7001900 kB) free!
[1554951621] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.6% (5223024 kB) free!
[1554951624] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44273 MB (19% inode=99%):
[1554951744] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44273 MB (19% inode=99%):
[1554951864] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44279 MB (19% inode=99%):
[1554952209] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 491 processes
[1554952221] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (2633320 kB) free!
[1554952269] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 490 processes
[1554952431] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 15.7% (3848844 kB) free!
[1554952464] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48013 MB (21% inode=99%):
[1554952961] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.63, 0.83, 12.20
[1554953064] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45999 MB (20% inode=99%):
[1554953185] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56198 MB (25% inode=99%):
[1554953784] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45982 MB (20% inode=99%):
[1554953904] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45982 MB (20% inode=99%):
[1554954024] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45990 MB (20% inode=99%):
[1554954161] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.17, 0.90, 3.94
[1554954624] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47394 MB (21% inode=99%):
[1554955513] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1554956113] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554956233] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554956353] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554956424] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46172 MB (20% inode=99%):
[1554956544] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46172 MB (20% inode=99%):
[1554956664] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46172 MB (20% inode=99%):
[1554957264] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54041 MB (24% inode=99%):
[1554957560] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1554959065] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44182 MB (19% inode=99%):
[1554959185] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55955 MB (24% inode=99%):
[1554960190] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 468.18, 347.95, 166.05
[1554960310] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 468.48, 387.83, 202.72
[1554960430] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 64.02, 259.85, 178.27
[1554960875] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.8% (2649916 kB) free!
[1554960985] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46231 MB (20% inode=99%):
[1554961071] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 467.04, 311.53, 139.70
[1554961105] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55848 MB (24% inode=99%):
[1554961153] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1554961158] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.7% (5285744 kB) free!
[1554961191] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 468.54, 363.51, 179.58
[1554961235] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 47.0% (7633132 kB) free.
[1554961312] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 468.94, 398.37, 214.65
[1554961379] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 33609 MB (14% inode=99%):
[1554961475] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2248280 kB) free!
[1554961753] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1554961758] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (222696 kB) free!
[1554961835] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (170096 kB) free!
[1554961873] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554961955] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (166324 kB) free!
[1554961979] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 76725 MB (34% inode=99%):
[1554961993] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554962014] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24599 MB (10% inode=99%):
[1554962075] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 81.7% (13272436 kB) free.
[1554962305] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47035 MB (20% inode=99%):
[1554962356] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.1% (2947172 kB) free!
[1554962425] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47035 MB (20% inode=99%):
[1554962476] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.2% (2791072 kB) free!
[1554962545] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47035 MB (20% inode=99%):
[1554962597] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.3% (2652288 kB) free!
[1554962830] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.71, 0.75, 14.10
[1554963145] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47458 MB (21% inode=99%):
[1554963745] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43833 MB (19% inode=99%):
[1554963797] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (695836 kB) free!
[1554963865] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43833 MB (19% inode=99%):
[1554963985] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43833 MB (19% inode=99%):
[1554964030] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.93, 0.98, 4.52
[1554964312] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.83, 1.81, 11.69
[1554964585] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55507 MB (24% inode=99%):
[1554965185] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47058 MB (20% inode=99%):
[1554965305] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47058 MB (20% inode=99%):
[1554965425] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47058 MB (20% inode=99%):
[1554965512] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.49, 1.68, 4.44
[1554968450] SERVICE ALERT: sig1;Memory;WARNING;SOFT;1;WARNING - 12.9% (2092760 kB) free!
[1554968570] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 25.6% (4155812 kB) free.
[1554968675] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.4% (2663716 kB) free!
[1554968795] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.0% (2440964 kB) free!
[1554968915] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 13.7% (2227084 kB) free!
[1554969170] SERVICE ALERT: sig1;Memory;WARNING;SOFT;1;WARNING - 18.7% (3037924 kB) free!
[1554969290] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 32.7% (5315096 kB) free.
[1554969516] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1151268 kB) free!
[1554970225] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50594 MB (22% inode=99%):
[1554970825] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46498 MB (20% inode=99%):
[1554970945] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45474 MB (20% inode=99%):
[1554971014] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51414 MB (22% inode=99%):
[1554971065] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45474 MB (20% inode=99%):
[1554971255] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 30.1% (7386136 kB) free.
[1554971303] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 80.2% (19661384 kB) free.
[1554971561] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 466.82, 202.19, 77.42
[1554971681] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 506.83, 304.26, 129.91
[1554971801] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 159.58, 266.07, 139.05
[1554972059] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;1;WARNING - 17.1% (1258076 kB) free!
[1554972179] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;2;WARNING - 16.3% (1199440 kB) free!
[1554972291] SERVICE ALERT: sig1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (168496 kB) free!
[1554972299] SERVICE ALERT: sig4-vm1;Memory;WARNING;HARD;3;WARNING - 15.6% (1149196 kB) free!
[1554972411] SERVICE ALERT: sig1;Memory;WARNING;SOFT;2;WARNING - 16.8% (2724664 kB) free!
[1554972456] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 38.0% (9370528 kB) free.
[1554972531] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (176344 kB) free!
[1554972763] SERVICE ALERT: pservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 122450 MB (55% inode=99%):
[1554973056] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (254324 kB) free!
[1554973176] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (469080 kB) free!
[1554973296] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176084 kB) free!
[1554973393] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1554973732] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 18.6% (3015012 kB) free!
[1554974065] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54680 MB (24% inode=99%):
[1554974201] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.42, 1.27, 11.79
[1554974332] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 26.6% (4330388 kB) free.
[1554974464] SERVICE ALERT: sig6-vm1;Memory;WARNING;SOFT;1;WARNING - 19.8% (1456048 kB) free!
[1554974584] SERVICE ALERT: sig6-vm1;Memory;WARNING;SOFT;2;WARNING - 19.6% (1443764 kB) free!
[1554974665] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46485 MB (20% inode=99%):
[1554974704] SERVICE ALERT: sig6-vm1;Memory;WARNING;HARD;3;WARNING - 19.5% (1431876 kB) free!
[1554974785] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46485 MB (20% inode=99%):
[1554974905] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46485 MB (20% inode=99%):
[1554974933] SERVICE ALERT: sig1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (146704 kB) free!
[1554975053] SERVICE ALERT: sig1;Memory;WARNING;SOFT;2;WARNING - 19.3% (3142716 kB) free!
[1554975173] SERVICE FLAPPING ALERT: sig1;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1554975173] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (149548 kB) free!
[1554975187] SERVICE ALERT: sservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 79880 MB (36% inode=99%):
[1554975275] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 80.5% (19735236 kB) free.
[1554975401] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.16, 1.13, 3.99
[1554975773] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 10.5% (1700860 kB) free!
[1554975892] SERVICE ALERT: sig4-vm2;Memory;WARNING;SOFT;1;WARNING - 19.7% (1445340 kB) free!
[1554975968] SERVICE ALERT: sig5-vm1;Memory;WARNING;SOFT;1;WARNING - 19.7% (1450556 kB) free!
[1554976012] SERVICE ALERT: sig4-vm2;Memory;WARNING;SOFT;2;WARNING - 19.5% (1436304 kB) free!
[1554976088] SERVICE ALERT: sig5-vm1;Memory;WARNING;SOFT;2;WARNING - 19.6% (1440528 kB) free!
[1554976132] SERVICE ALERT: sig4-vm2;Memory;WARNING;HARD;3;WARNING - 19.4% (1426260 kB) free!
[1554976208] SERVICE ALERT: sig5-vm1;Memory;WARNING;HARD;3;WARNING - 19.5% (1431948 kB) free!
[1554976374] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 24.9% (4045116 kB) free.
[1554976499] SERVICE ALERT: sig4-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (699356 kB) free!
[1554976630] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 388.35, 127.35, 45.43
[1554976751] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 72.69, 103.13, 47.35
[1554976871] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 10.42, 69.26, 41.70
[1554976974] SERVICE ALERT: sig1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (164660 kB) free!
[1554977094] SERVICE ALERT: sig1;Memory;WARNING;SOFT;2;WARNING - 18.1% (2947256 kB) free!
[1554977214] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161092 kB) free!
[1554977272] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 19.0% (1397652 kB) free!
[1554977305] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47452 MB (21% inode=99%):
[1554977392] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 18.8% (1381516 kB) free!
[1554977512] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 18.6% (1365428 kB) free!
[1554977815] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 11.9% (1926320 kB) free!
[1554978071] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.63, 1.94, 11.94
[1554978415] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 22.4% (3637244 kB) free.
[1554978439] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;1;WARNING - 18.5% (1362884 kB) free!
[1554978503] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4754648 kB) free!
[1554978505] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46325 MB (20% inode=99%):
[1554978559] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;2;WARNING - 18.1% (1328892 kB) free!
[1554978624] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 65.4% (16029532 kB) free.
[1554978625] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46325 MB (20% inode=99%):
[1554978680] SERVICE ALERT: sig2-vm1;Memory;WARNING;HARD;3;WARNING - 17.6% (1294660 kB) free!
[1554978745] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46325 MB (20% inode=99%):
[1554979015] SERVICE ALERT: sig1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (1120540 kB) free!
[1554979135] SERVICE ALERT: sig1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (151696 kB) free!
[1554979255] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (1428368 kB) free!
[1554979271] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.85, 0.96, 3.93
[1554979689] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;1;WARNING - 18.6% (1364844 kB) free!
[1554979809] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;2;WARNING - 18.1% (1333256 kB) free!
[1554979897] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 37.4% (9233300 kB) free.
[1554979929] SERVICE ALERT: sig3-vm1;Memory;WARNING;HARD;3;WARNING - 17.6% (1296256 kB) free!
[1554979946] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54379 MB (24% inode=99%):
[1554980455] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 19.7% (3203176 kB) free!
[1554980546] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45927 MB (20% inode=99%):
[1554980666] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45927 MB (20% inode=99%):
[1554980704] SERVICE ALERT: sig6-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (648800 kB) free!
[1554980786] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45927 MB (20% inode=99%):
[1554980932] SERVICE ALERT: sig4-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (713120 kB) free!
[1554981008] SERVICE ALERT: sig5-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (677564 kB) free!
[1554981055] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (933164 kB) free!
[1554981080] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (607840 kB) free!
[1554981112] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (637648 kB) free!
[1554981986] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50009 MB (22% inode=99%):
[1554982329] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (646148 kB) free!
[1554982586] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45401 MB (20% inode=99%):
[1554982650] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.9% (1460172 kB) free!
[1554982706] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42328 MB (18% inode=99%):
[1554982770] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 19.7% (1448400 kB) free!
[1554982799] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554982826] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53962 MB (24% inode=99%):
[1554982855] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 19.9% (3230464 kB) free!
[1554982890] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 19.5% (1435876 kB) free!
[1554982919] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554983039] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554984055] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (1447240 kB) free!
[1554984626] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45889 MB (20% inode=99%):
[1554984746] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45889 MB (20% inode=99%):
[1554984866] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45889 MB (20% inode=99%):
[1554988439] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554988466] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49452 MB (22% inode=99%):
[1554988979] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47104 MB (20% inode=99%):
[1554989066] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44332 MB (19% inode=99%):
[1554989099] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47056 MB (20% inode=99%):
[1554989186] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42284 MB (18% inode=99%):
[1554989219] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47056 MB (20% inode=99%):
[1554989306] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53538 MB (23% inode=99%):
[1554989819] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47283 MB (21% inode=99%):
[1554989906] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45344 MB (20% inode=99%):
[1554990027] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45344 MB (20% inode=99%):
[1554990147] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45344 MB (20% inode=99%):
[1554990419] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46460 MB (20% inode=99%):
[1554990539] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46459 MB (20% inode=99%):
[1554990659] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46393 MB (20% inode=99%):
[1554990691] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (726232 kB) free!
[1554990697] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4768028 kB) free!
[1554990817] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4725660 kB) free!
[1554990839] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554990937] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.9% (4667012 kB) free!
[1554990959] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554991079] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554991855] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 20.2% (3290572 kB) free.
[1554991947] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 58263 MB (25% inode=99%):
[1554992455] SERVICE ALERT: sig1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (154720 kB) free!
[1554992575] SERVICE ALERT: sig1;Memory;WARNING;SOFT;2;WARNING - 13.9% (2254312 kB) free!
[1554992695] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (162436 kB) free!
[1554993023] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (2323640 kB) free!
[1554993143] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (2301148 kB) free!
[1554993193] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1554993265] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169608 kB) free!
[1554993296] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 10.2% (1651232 kB) free!
[1554993337] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (231240 kB) free!
[1554993455] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176380 kB) free!
[1554993575] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172212 kB) free!
[1554993695] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165996 kB) free!
[1554993794] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1554993875] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (1496268 kB) free!
[1554993896] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (158264 kB) free!
[1554993914] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554993995] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (563044 kB) free!
[1554994034] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1554994116] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (530280 kB) free!
[1554994465] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.6% (2597820 kB) free!
[1554994496] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 18.7% (3043648 kB) free!
[1554995065] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (215944 kB) free!
[1554995096] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (164688 kB) free!
[1554998879] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555000497] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 11.7% (1900380 kB) free!
[1555001234] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1555001697] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (159108 kB) free!
[1555001834] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1555001954] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1555002074] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1555002074] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1555002147] SERVICE FLAPPING ALERT: db2;HDD_Home status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1555004097] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 10.2% (1649832 kB) free!
[1555004471] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 503.92, 341.34, 153.78
[1555004591] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 333.28, 355.17, 183.03
[1555004607] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 278.35, 333.55, 173.52
[1555004697] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (154196 kB) free!
[1555004711] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 45.20, 237.79, 160.89
[1555004727] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 37.97, 223.39, 152.56
[1555004847] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 5.48, 149.61, 134.11
[1555004916] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.9% (2912232 kB) free!
[1555005096] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.6% (2855384 kB) free!
[1555005897] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 12.3% (1997220 kB) free!
[1555006274] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1555006497] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (424184 kB) free!
[1555006647] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.33, 0.66, 19.54
[1555006874] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1555006994] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1555007111] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.95, 0.94, 12.86
[1555007115] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes