[1563129057] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11117520 kB) free.
[1563129595] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563129657] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6734344 kB) free!
[1563129716] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563129777] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11124716 kB) free.
[1563130315] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563130377] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6450780 kB) free!
[1563130435] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563130497] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11125352 kB) free.
[1563131035] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563131097] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.5% (6153012 kB) free!
[1563131155] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563131218] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11123184 kB) free.
[1563131755] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1563131875] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563132417] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7832708 kB) free!
[1563132475] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563132537] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6216192 kB) free!
[1563132595] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563132615] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185908 kB) free!
[1563132623] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 489 processes
[1563132657] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11118608 kB) free.
[1563132683] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 485 processes
[1563132735] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169836 kB) free!
[1563132855] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183368 kB) free!
[1563133196] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563133257] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6436224 kB) free!
[1563133316] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563133377] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11123480 kB) free.
[1563133916] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563133977] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6292052 kB) free!
[1563134036] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563134097] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11123768 kB) free.
[1563134516] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (167084 kB) free!
[1563134636] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (154636 kB) free!
[1563134636] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563134697] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6454280 kB) free!
[1563134756] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (160036 kB) free!
[1563134756] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1563134817] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11122012 kB) free.
[1563135138] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (180560 kB) free!
[1563135258] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (185264 kB) free!
[1563135356] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563135378] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (178260 kB) free!
[1563135417] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6656228 kB) free!
[1563135476] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563135538] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11124684 kB) free.
[1563136077] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563136137] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6522368 kB) free!
[1563136197] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563136257] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11114048 kB) free.
[1563136578] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 75.6% (12280372 kB) free.
[1563136797] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563136857] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6410848 kB) free!
[1563136917] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563136977] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11120452 kB) free.
[1563137510] SERVICE ALERT: sigqscluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1563137517] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563137637] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563138177] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7575340 kB) free!
[1563138237] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563138297] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11122588 kB) free.
[1563138357] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563138897] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7437972 kB) free!
[1563138957] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563139017] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11121536 kB) free.
[1563139078] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563139617] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7371632 kB) free!
[1563139677] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563139737] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11122092 kB) free.
[1563139797] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563140338] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7120968 kB) free!
[1563140397] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563140458] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11111188 kB) free.
[1563140517] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563141058] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6803312 kB) free!
[1563141178] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11120980 kB) free.
[1563141717] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563141778] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6537748 kB) free!
[1563141837] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563141898] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11122228 kB) free.
[1563142438] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563142498] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6297044 kB) free!
[1563142558] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1563142618] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11121328 kB) free.
[1563143158] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563143218] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6261540 kB) free!
[1563143278] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563143338] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11123692 kB) free.
[1563143878] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563143998] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563144539] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7497152 kB) free!
[1563144598] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563144659] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11114808 kB) free.
[1563144719] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563145259] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7534276 kB) free!
[1563145318] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563145379] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11121660 kB) free.
[1563145438] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563145979] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7482468 kB) free!
[1563146038] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563146099] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11116424 kB) free.
[1563146158] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563146699] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7518020 kB) free!
[1563146758] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563146819] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11109784 kB) free.
[1563146878] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563147419] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7490896 kB) free!
[1563147478] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563147539] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11110184 kB) free.
[1563147599] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563147719] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563148739] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8627144 kB) free!
[1563148860] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7252904 kB) free!
[1563148980] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11080428 kB) free.
[1563149580] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9843688 kB) free!
[1563149700] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.2% (8486968 kB) free!
[1563149820] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.5% (7121436 kB) free!
[1563150420] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11053432 kB) free.
[1563150772] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 14325.45, 8529.24, 3600.80
[1563150892] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 14476.09, 10504.67, 4922.27
[1563150957] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 33.7% (5475476 kB) free.
[1563151012] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14496.29, 11826.39, 6083.47
[1563151020] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9464580 kB) free!
[1563151140] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.5% (8151176 kB) free!
[1563151260] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.8% (6802916 kB) free!
[1563151319] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563151860] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1563151860] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11028624 kB) free.
[1563152460] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7249140 kB) free!
[1563152519] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563152580] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11039088 kB) free.
[1563152639] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563153181] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7286384 kB) free!
[1563153239] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563153301] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11037204 kB) free.
[1563153359] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563153901] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7342244 kB) free!
[1563153959] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563154021] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11034684 kB) free.
[1563154079] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563154621] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7535264 kB) free!
[1563154679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563154741] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11036576 kB) free.
[1563154800] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563155341] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7546100 kB) free!
[1563155399] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563155461] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11028368 kB) free.
[1563155519] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563155639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563156336] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.3% (21797784 kB) free.
[1563156661] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9017320 kB) free!
[1563156781] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.7% (7746812 kB) free!
[1563156902] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.2% (6002212 kB) free!
[1563157012] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.28, 0.30, 16.69
[1563158212] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.33, 0.29, 4.78
[1563158639] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1563158639] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563158702] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10428496 kB) free.
[1563161215] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1563161335] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1563161639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563161702] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6904912 kB) free!
[1563161759] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563161822] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11022980 kB) free.
[1563161879] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563162479] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563162710] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1563162785] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (181892 kB) free!
[1563162830] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1563162905] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (169884 kB) free!
[1563162950] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1563162957] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (161552 kB) free!
[1563163025] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (176804 kB) free!
[1563163077] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (166324 kB) free!
[1563163197] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (152572 kB) free!
[1563163278] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (217188 kB) free!
[1563163398] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (238512 kB) free!
[1563163518] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (214204 kB) free!
[1563163623] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7312844 kB) free!
[1563163679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563163743] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11022180 kB) free.
[1563163799] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563164343] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7452752 kB) free!
[1563164399] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563164463] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11017560 kB) free.
[1563164519] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563165063] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7228552 kB) free!
[1563165119] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563165183] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11019944 kB) free.
[1563165239] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563165783] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7330400 kB) free!
[1563165839] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563165903] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (11009476 kB) free.
[1563165960] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1563166503] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7524408 kB) free!
[1563166546] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.3% (2800680 kB) free!
[1563166559] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563166623] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11017060 kB) free.
[1563166660] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193332 kB) free!
[1563166666] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.0% (2722000 kB) free!
[1563166679] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563166768] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186856 kB) free!
[1563166780] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183408 kB) free!
[1563166786] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.7% (2634632 kB) free!
[1563166889] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184252 kB) free!
[1563166900] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182304 kB) free!
[1563167009] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184392 kB) free!
[1563167042] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.81, 7.32, 5.05
[1563167162] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.06, 6.70, 5.09
[1563167223] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8440824 kB) free!
[1563167279] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563167282] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.42, 5.88, 4.98
[1563167343] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7545332 kB) free!
[1563167386] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (1212536 kB) free!
[1563167399] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563167463] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6184232 kB) free!
[1563167519] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563167882] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.28, 5.82, 5.41
[1563167934] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 57.53, 14.84, 5.45
[1563168002] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.06, 5.33, 5.28
[1563168054] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 119.96, 52.82, 20.50
[1563168119] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563168122] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.90, 5.60, 5.38
[1563168174] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 128.47, 78.29, 33.75
[1563168240] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563168360] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1563168378] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.3% (2644024 kB) free!
[1563168498] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.0% (2432764 kB) free!
[1563168618] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.2% (2637804 kB) free!
[1563169263] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (11002668 kB) free.
[1563169322] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.80, 3.88, 4.51
[1563169818] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (473956 kB) free!
[1563169922] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.26, 6.95, 5.70
[1563170042] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.46, 8.09, 6.29
[1563170125] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1563170162] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 8.96, 8.73, 6.76
[1563170245] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1563170259] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (13074144 kB) free!
[1563170379] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (19078456 kB) free.
[1563170397] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 62.2% (10103492 kB) free.
[1563170574] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.45, 1.74, 14.23
[1563170998] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (293732 kB) free!
[1563171063] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9200616 kB) free!
[1563171118] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (175896 kB) free!
[1563171183] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7896564 kB) free!
[1563171238] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (175432 kB) free!
[1563171303] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.6% (6203860 kB) free!
[1563171360] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563171434] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 13.8% (1012124 kB) free!
[1563171554] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (691376 kB) free!
[1563171674] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (417560 kB) free!
[1563171774] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.53, 0.62, 4.38
[1563171903] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10253128 kB) free.
[1563171960] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563172080] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563172201] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563172274] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.1% (6848700 kB) free.
[1563172503] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8879736 kB) free!
[1563172623] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.4% (7605596 kB) free!
[1563172735] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1563172743] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.0% (5936772 kB) free!
[1563172801] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563172855] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1563173344] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10338196 kB) free.
[1563173401] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563173521] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563173641] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563175012] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 83.96, 24.68, 9.40
[1563175088] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.11, 5.92, 5.24
[1563175132] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 131.39, 62.45, 25.09
[1563175144] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (6980092 kB) free!
[1563175208] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.37, 6.16, 5.41
[1563175252] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 137.42, 87.58, 38.83
[1563175264] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10977544 kB) free.
[1563175328] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.11, 6.48, 5.62
[1563175610] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1236 processes
[1563175612] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.4% (19839132 kB) free.
[1563175694] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 84.2% (20636736 kB) free.
[1563175730] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1237 processes
[1563175769] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.9% (39870764 kB) free.
[1563175819] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 82.9% (13475828 kB) free.
[1563175850] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 334 processes
[1563175891] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 81.6% (20001572 kB) free.
[1563175946] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.1% (14800752 kB) free.
[1563175974] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.61, 5.30, 5.48
[1563176016] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 68.8% (11176944 kB) free.
[1563176094] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.72, 4.18, 5.06
[1563176212] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183020 kB) free!
[1563176214] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.87, 3.06, 4.54
[1563176332] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164392 kB) free!
[1563176452] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1563176452] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174940 kB) free!
[1563176528] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 3.75, 4.15, 4.94
[1563176631] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (163892 kB) free!
[1563176751] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (182044 kB) free!
[1563176871] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1563176871] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (550376 kB) free!
[1563177052] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.57, 1.14, 13.33
[1563177674] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 14.5% (1068320 kB) free!
[1563177794] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (715300 kB) free!
[1563177914] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (347988 kB) free!
[1563178252] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.38, 0.47, 4.03
[1563178441] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563178825] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (3125872 kB) free!
[1563178864] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6242596 kB) free!
[1563178945] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.2% (2790740 kB) free!
[1563178960] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.1% (2942492 kB) free!
[1563178985] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10974064 kB) free.
[1563179065] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.2% (2465360 kB) free!
[1563179080] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.6% (2864644 kB) free!
[1563179200] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 16.8% (2726900 kB) free!
[1563179584] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6440196 kB) free!
[1563179665] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1563179665] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1235316 kB) free!
[1563179704] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10954076 kB) free.
[1563180242] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563180304] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.5% (6176696 kB) free!
[1563180362] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563180400] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (180440 kB) free!
[1563180424] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10963132 kB) free.
[1563180962] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563181024] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.0% (5920936 kB) free!
[1563181082] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563181144] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10961352 kB) free.
[1563181682] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563181802] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563181855] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1563181975] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1563182039] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 227.91, 184.26, 95.38
[1563182159] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 118.33, 170.48, 101.95
[1563182280] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 19.32, 115.26, 90.02
[1563182344] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7479696 kB) free!
[1563182402] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563182464] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10954948 kB) free.
[1563182522] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563182978] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (241900 kB) free!
[1563183098] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (222964 kB) free!
[1563183123] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563183218] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (259692 kB) free!
[1563183243] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563183314] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.6% (6812712 kB) free.
[1563183363] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563183664] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8518256 kB) free!
[1563183785] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.6% (7200416 kB) free!
[1563183905] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10935628 kB) free.
[1563184080] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.28, 5.14, 16.69
[1563184505] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9640668 kB) free!
[1563184625] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.0% (8356752 kB) free!
[1563184745] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.2% (7010636 kB) free!
[1563185763] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563185892] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 196 processes
[1563186012] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 194 processes
[1563186239] SERVICE ALERT: sig4-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (131456 kB) free!
[1563186359] SERVICE ALERT: sig4-vm2;Memory;OK;HARD;1;OK - 36.7% (2681412 kB) free.
[1563186480] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.29, 3.24, 4.28
[1563186545] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10754852 kB) free.
[1563187094] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (4611224 kB) free!
[1563187145] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9142008 kB) free!
[1563187214] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (4544456 kB) free!
[1563187265] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7802208 kB) free!
[1563187334] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (4477980 kB) free!
[1563187385] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6065088 kB) free!
[1563187671] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563187852] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563188763] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563188883] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563189185] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9923728 kB) free.
[1563189362] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.89, 4.32, 4.76
[1563189483] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563189603] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1563189736] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195976 kB) free!
[1563190203] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563190324] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563190844] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.2% (21520720 kB) free.
[1563190857] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 54.6% (13377008 kB) free.
[1563190878] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 64.2% (31661440 kB) free.
[1563190896] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 81.8% (13294928 kB) free.
[1563190923] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563190935] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 61.9% (10068076 kB) free.
[1563190971] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 85.8% (13952432 kB) free.
[1563191043] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563191448] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.5% (3073708 kB) free!
[1563191487] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (214992 kB) free!
[1563191496] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4899716 kB) free!
[1563191568] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182396 kB) free!
[1563191578] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (169536 kB) free!
[1563191585] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6917588 kB) free!
[1563191607] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (225448 kB) free!
[1563191616] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.8% (4854508 kB) free!
[1563191643] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563191688] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194124 kB) free!
[1563191699] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (161960 kB) free!
[1563191705] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10917836 kB) free.
[1563191714] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 14.4% (1062128 kB) free!
[1563191727] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (216928 kB) free!
[1563191736] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.7% (4820472 kB) free!
[1563191763] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563191819] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (162524 kB) free!
[1563191834] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 13.1% (964268 kB) free!
[1563191954] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 11.7% (860612 kB) free!
[1563192305] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6949208 kB) free!
[1563192363] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563192425] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10926808 kB) free.
[1563192483] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563192554] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (247904 kB) free!
[1563192775] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1563192895] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1563193025] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7117392 kB) free!
[1563193083] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563193145] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10931332 kB) free.
[1563193204] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563193233] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3162268 kB) free!
[1563193354] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.1% (3109384 kB) free!
[1563193365] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1563193368] HOST ALERT: sig4-vm2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.191)
[1563193415] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig4-vm2' was out of bounds)
[1563193474] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.9% (3067292 kB) free!
[1563193488] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig4-vm2' was out of bounds)
[1563193561] SERVICE ALERT: sig4-vm2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig4-vm2' was out of bounds)
[1563193630] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig4-vm2' was out of bounds)
[1563193745] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7102080 kB) free!
[1563193803] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563193813] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig4-vm2' was out of bounds)
[1563193813] SERVICE ALERT: sig4-vm2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig4-vm2' was out of bounds)
[1563193834] SERVICE ALERT: sig4-vm2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig4-vm2' was out of bounds)
[1563193866] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10921092 kB) free.
[1563193923] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563193935] HOST ALERT: sig4-vm2;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.191)
[1563194074] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 83.3% (13545316 kB) free.
[1563194355] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1563194355] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.2% (6853644 kB) free.
[1563194465] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6809628 kB) free!
[1563194523] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563194585] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10927044 kB) free.
[1563194643] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563194686] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 233 processes
[1563194690] HOST ALERT: sig4-vm2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1563194735] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1563194882] SERVICE ALERT: sig4-vm2;Memory;OK;HARD;1;OK - 86.6% (6327488 kB) free.
[1563194950] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 0.14, 0.15, 0.14
[1563195134] SERVICE ALERT: sig4-vm2;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1563195134] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 236 processes
[1563195154] SERVICE ALERT: sig4-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 119008 MB (92% inode=99%):
[1563195185] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6970312 kB) free!
[1563195243] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563195305] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10911808 kB) free.
[1563195338] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig4-vm2' was out of bounds)
[1563195341] HOST ALERT: sig4-vm2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.191)
[1563195363] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563195405] HOST ALERT: sig4-vm2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1563195459] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1563195705] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1563195905] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7142632 kB) free!
[1563195963] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563196025] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10923952 kB) free.
[1563196084] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563196412] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.1% (14646800 kB) free.
[1563196443] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 73.3% (18083880 kB) free.
[1563196457] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.7% (39768056 kB) free.
[1563196495] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1563196615] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1563196625] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7125564 kB) free!
[1563196683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563196745] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10923808 kB) free.
[1563196803] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1563196803] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563197043] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177620 kB) free!
[1563197163] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175040 kB) free!
[1563197283] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1563197283] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170636 kB) free!
[1563197403] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563197523] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563197643] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563197814] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 72.93, 33.87, 14.82
[1563197880] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 152.23, 40.55, 15.81
[1563197934] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 13.23, 24.05, 13.54
[1563198000] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 23.50, 28.23, 14.29
[1563198052] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 54.66, 12.15, 4.29
[1563198054] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.52, 17.14, 12.28
[1563198120] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.64, 19.81, 12.89
[1563198172] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 21.43, 20.46, 8.88
[1563198292] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.25, 13.85, 7.86
[1563198892] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.05, 2.60, 4.48
[1563199146] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6437192 kB) free!
[1563199254] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.07, 1.88, 4.90
[1563199266] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10903204 kB) free.
[1563199444] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563199905] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.4% (2991892 kB) free!
[1563199920] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.78, 3.04, 4.54
[1563200012] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (166100 kB) free!
[1563200025] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.7% (2880436 kB) free!
[1563200132] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (198628 kB) free!
[1563200145] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.0% (2756184 kB) free!
[1563200252] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1563200252] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (200116 kB) free!
[1563200454] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 22.54, 21.24, 11.46
[1563200466] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9386200 kB) free!
[1563200520] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.34, 14.93, 11.01
[1563200574] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.60, 15.64, 10.60
[1563200586] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1563200586] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10909152 kB) free.
[1563200640] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.23, 10.99, 10.04
[1563200692] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.53, 7.44, 6.42
[1563200694] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.87, 11.99, 9.88
[1563200760] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.87, 8.65, 9.30
[1563200812] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.56, 5.15, 5.70
[1563200815] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1563200932] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.53, 3.62, 5.08
[1563200935] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1563201055] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1563201301] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2336500 kB) free!
[1563201345] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1608628 kB) free!
[1563201532] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1563201532] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14469.07, 10160.13, 4665.13
[1563201786] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7290836 kB) free!
[1563201844] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563201906] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10895940 kB) free.
[1563201960] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1563201960] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 19.41, 27.31, 19.22
[1563201964] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1563202501] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 13.4% (3274304 kB) free!
[1563202506] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7052324 kB) free!
[1563202560] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.66, 6.18, 11.54
[1563202564] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563202626] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10908524 kB) free.
[1563202685] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563203094] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1563203094] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 63.12, 22.86, 13.54
[1563203226] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6836560 kB) free!
[1563203284] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563203347] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10906304 kB) free.
[1563203404] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563203694] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.98, 3.80, 7.48
[1563203760] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 108.93, 38.55, 21.15
[1563203946] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6723160 kB) free!
[1563204066] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10898024 kB) free.
[1563204360] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.31, 7.93, 12.56
[1563204555] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1563204604] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1563204666] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6706056 kB) free!
[1563204724] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563204735] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (555284 kB) free!
[1563204786] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10902648 kB) free.
[1563204855] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (164280 kB) free!
[1563204960] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14421.87, 9457.22, 4185.68
[1563204975] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (264824 kB) free!
[1563205324] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563205386] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6642160 kB) free!
[1563205444] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563205506] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10895884 kB) free.
[1563205645] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 2 warn out of 368 processes
[1563205765] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 1 warn out of 368 processes
[1563205885] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 1 warn out of 370 processes
[1563206045] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563206094] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 115.17, 29.54, 12.97
[1563206106] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6749408 kB) free!
[1563206165] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1563206226] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10898216 kB) free.
[1563206295] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.95, 7.78, 5.68
[1563206415] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.45, 8.05, 6.04
[1563206535] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.20, 8.00, 6.26
[1563206694] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.53, 6.36, 7.79
[1563206765] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563206826] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6517528 kB) free!
[1563206885] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563206946] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10902468 kB) free.
[1563207302] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2380364 kB) free!
[1563207485] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563207546] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6274596 kB) free!
[1563207605] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563207667] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10890192 kB) free.
[1563208205] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563208266] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.3% (6053512 kB) free!
[1563208325] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563208386] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10900900 kB) free.
[1563208652] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1563208683] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563208926] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563209046] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563209175] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 47.0% (7645336 kB) free.
[1563209332] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.24, 3.06, 10.74
[1563209586] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7411464 kB) free!
[1563209646] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563209706] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10888984 kB) free.
[1563209766] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1563209775] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (452920 kB) free!
[1563209895] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (167556 kB) free!
[1563209932] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 42.67, 24.44, 14.90
[1563210015] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165080 kB) free!
[1563210085] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1563210306] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7144936 kB) free!
[1563210366] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563210426] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10896232 kB) free.
[1563210486] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563211026] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7014532 kB) free!
[1563211086] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563211132] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.58, 4.28, 9.23
[1563211146] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10902636 kB) free.
[1563211206] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563211285] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 365 processes
[1563211336] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.99, 2.03, 3.92
[1563211405] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 360 processes
[1563211525] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1563211732] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.29, 0.81, 4.96
[1563211746] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7015916 kB) free!
[1563211807] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563211866] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10893044 kB) free.
[1563211927] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563212160] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.49, 8.25, 18.06
[1563212467] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6806964 kB) free!
[1563212587] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10899932 kB) free.
[1563212932] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.62, 4.02, 5.86
[1563213052] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.31, 2.78, 5.18
[1563213127] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563213172] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.41, 1.95, 4.59
[1563213187] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6577132 kB) free!
[1563213247] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563213307] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10889376 kB) free.
[1563213616] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1563213616] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (1682544 kB) free!
[1563213847] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563213894] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 94.13, 33.92, 18.76
[1563213907] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6302076 kB) free!
[1563213967] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1563214027] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10898080 kB) free.
[1563214216] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1555812 kB) free!
[1563214494] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.04, 5.26, 10.18
[1563214568] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563214627] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5967848 kB) free!
[1563214688] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563214747] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10898276 kB) free.
[1563214816] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.3% (1678104 kB) free!
[1563214972] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.27, 5.09, 6.39
[1563215092] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.56, 3.60, 5.69
[1563215212] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 43.51, 23.07, 12.70
[1563215288] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes