[1565721044] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11313768 kB) free.
[1565721066] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565721224] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.4% (16301124 kB) free!
[1565721644] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7977500 kB) free!
[1565721666] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1565721764] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.2% (6501288 kB) free!
[1565721786] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565721884] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11323240 kB) free.
[1565721906] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565722485] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6926724 kB) free!
[1565722507] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1565722605] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11319756 kB) free.
[1565722627] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1565723205] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7233252 kB) free!
[1565723227] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565723325] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11324744 kB) free.
[1565723347] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565723525] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 42.9% (10585504 kB) free.
[1565723925] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7471668 kB) free!
[1565723947] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565724044] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1565724045] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11314804 kB) free.
[1565724067] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1565724067] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565724647] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7551248 kB) free!
[1565724667] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565724767] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11322880 kB) free.
[1565724787] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565724824] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2062716 kB) free!
[1565725367] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7637784 kB) free!
[1565725388] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565725487] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11325208 kB) free.
[1565725508] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565726087] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7906176 kB) free!
[1565726108] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1565726207] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.5% (6164736 kB) free!
[1565726228] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1565726327] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11320896 kB) free.
[1565726828] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565726927] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6893280 kB) free!
[1565726948] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565727047] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11312448 kB) free.
[1565727068] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565727647] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7234820 kB) free!
[1565727668] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565727767] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11322084 kB) free.
[1565727788] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565728368] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7705820 kB) free!
[1565728389] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565728488] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11222520 kB) free.
[1565728509] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565729088] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7876356 kB) free!
[1565729109] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565729208] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.2% (6032564 kB) free!
[1565729229] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565729328] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11323320 kB) free.
[1565729525] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190076 kB) free!
[1565729645] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191496 kB) free!
[1565729765] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (201784 kB) free!
[1565729829] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565729901] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1565729928] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6679872 kB) free!
[1565729949] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565730021] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 397 processes
[1565730048] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11312296 kB) free.
[1565730069] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1565730332] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.9% (42825696 kB) free.
[1565730648] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7028044 kB) free!
[1565730669] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565730768] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11319028 kB) free.
[1565730790] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565731368] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7218756 kB) free!
[1565731389] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565731488] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11323800 kB) free.
[1565731509] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565732088] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7431992 kB) free!
[1565732109] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565732208] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11320244 kB) free.
[1565732229] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565732433] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 46.4% (11370172 kB) free.
[1565732454] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 43.2% (10579240 kB) free.
[1565732809] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7568624 kB) free!
[1565732829] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565732929] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11321396 kB) free.
[1565732949] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565733529] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7654892 kB) free!
[1565733549] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565733649] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11314572 kB) free.
[1565733670] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565734249] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7966052 kB) free!
[1565734269] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565734369] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.7% (6259328 kB) free!
[1565734389] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565734489] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11320916 kB) free.
[1565734509] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565735089] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6425040 kB) free!
[1565735109] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1565735209] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11320048 kB) free.
[1565735229] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1565735809] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6528720 kB) free!
[1565735829] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565735929] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11321928 kB) free.
[1565735949] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565736529] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6664164 kB) free!
[1565736550] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565736650] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11313496 kB) free.
[1565736670] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565737249] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7014540 kB) free!
[1565737270] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565737369] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11318980 kB) free.
[1565737390] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565737969] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7235396 kB) free!
[1565737990] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565738089] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11322448 kB) free.
[1565738110] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565738689] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7422416 kB) free!
[1565738710] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565738809] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11319652 kB) free.
[1565738830] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565739409] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7618684 kB) free!
[1565739430] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565739529] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11321812 kB) free.
[1565739551] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565740129] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7825652 kB) free!
[1565740150] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565740249] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.3% (6044792 kB) free!
[1565740270] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565740369] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11319888 kB) free.
[1565740870] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1565740969] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6442084 kB) free!
[1565740990] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565741090] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11320204 kB) free.
[1565741110] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565741689] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7069984 kB) free!
[1565741710] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1565741809] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11318316 kB) free.
[1565741830] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565742409] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7291288 kB) free!
[1565742431] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1565742529] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11319456 kB) free.
[1565742551] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565743129] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7440508 kB) free!
[1565743151] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565743249] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11310132 kB) free.
[1565743271] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565743424] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.8% (26119116 kB) free!
[1565743849] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7893708 kB) free!
[1565743871] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565743969] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6387632 kB) free!
[1565743991] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565744089] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11317668 kB) free.
[1565744111] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1565744689] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7163324 kB) free!
[1565744711] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565744809] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11318572 kB) free.
[1565744832] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565745409] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7754572 kB) free!
[1565745431] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565745530] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.3% (6056164 kB) free!
[1565745551] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565745650] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11313168 kB) free.
[1565746151] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565746250] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6685044 kB) free!
[1565746271] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565746370] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11318128 kB) free.
[1565746391] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565746847] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1565746967] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1565746970] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7707452 kB) free!
[1565746991] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565747087] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1565747090] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6117588 kB) free!
[1565747111] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565747210] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11317652 kB) free.
[1565747232] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565747435] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186668 kB) free!
[1565747455] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190340 kB) free!
[1565747558] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191808 kB) free!
[1565747575] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181528 kB) free!
[1565747678] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187552 kB) free!
[1565747696] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184848 kB) free!
[1565747810] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7627416 kB) free!
[1565747831] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565747930] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6232564 kB) free!
[1565747951] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565748050] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11312104 kB) free.
[1565748071] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565748650] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7355588 kB) free!
[1565748671] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565748770] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11213896 kB) free.
[1565748791] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565749370] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8507712 kB) free!
[1565749391] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565749490] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.2% (7489196 kB) free!
[1565749511] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1565749611] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6076484 kB) free!
[1565749631] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565750232] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565750352] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565750472] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565750811] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (9982876 kB) free.
[1565751072] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565751192] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565751792] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565751912] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565752032] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565752611] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7224900 kB) free!
[1565752633] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565752731] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11317780 kB) free.
[1565752753] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565753331] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7984076 kB) free!
[1565753353] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1565753451] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6221328 kB) free!
[1565753473] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565753571] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11313620 kB) free.
[1565753593] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565754171] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7257408 kB) free!
[1565754193] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565754291] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11318072 kB) free.
[1565754313] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565754892] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7863912 kB) free!
[1565754913] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565755012] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.5% (6147992 kB) free!
[1565755034] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565755132] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11314460 kB) free.
[1565755154] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565755732] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7490452 kB) free!
[1565755754] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565755852] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.5% (6148360 kB) free!
[1565755874] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565755972] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11315064 kB) free.
[1565755994] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565756134] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.4% (14690064 kB) free.
[1565756209] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.0% (22296716 kB) free.
[1565756268] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 83.0% (40909588 kB) free.
[1565756314] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.3% (23114612 kB) free.
[1565756377] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.8% (22642408 kB) free.
[1565756428] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.2% (21629948 kB) free.
[1565756572] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7481396 kB) free!
[1565756594] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565756692] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6091004 kB) free!
[1565756714] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1565756812] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11314872 kB) free.
[1565756834] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1565757412] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6852780 kB) free!
[1565757434] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565757532] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11315520 kB) free.
[1565757555] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565758132] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7369896 kB) free!
[1565758154] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565758253] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11306292 kB) free.
[1565758274] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1565758544] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (176984 kB) free!
[1565758664] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (177180 kB) free!
[1565758675] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (231944 kB) free!
[1565758784] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161660 kB) free!
[1565758795] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (244972 kB) free!
[1565758852] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7846916 kB) free!
[1565758874] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565758915] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224776 kB) free!
[1565758972] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.1% (6474976 kB) free!
[1565758994] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565759092] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11312316 kB) free.
[1565759114] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565759692] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7732948 kB) free!
[1565759714] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565759812] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.1% (6438660 kB) free!
[1565759834] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565759932] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11312608 kB) free.
[1565759955] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565760532] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7709116 kB) free!
[1565760554] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565760652] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.0% (6420472 kB) free!
[1565760674] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1565760772] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11309188 kB) free.
[1565760794] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1565761372] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7379808 kB) free!
[1565761394] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565761492] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.0% (5900192 kB) free!
[1565761514] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565761613] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11307312 kB) free.
[1565761784] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189616 kB) free!
[1565761904] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193136 kB) free!
[1565762014] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 65 users currently logged in
[1565762024] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186784 kB) free!
[1565762074] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 65 users currently logged in
[1565762114] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565762134] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 65 users currently logged in
[1565762194] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 65 users currently logged in
[1565762212] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6863332 kB) free!
[1565762234] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565762332] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11312532 kB) free.
[1565762355] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565762799] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 365 processes
[1565762919] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 361 processes
[1565762932] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7663780 kB) free!
[1565762954] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565763039] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 363 processes
[1565763052] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6227140 kB) free!
[1565763074] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565763172] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11306468 kB) free.
[1565763194] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565763269] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.25, 7.19, 5.71
[1565763389] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.25, 6.22, 5.55
[1565763509] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.05, 5.19, 5.25
[1565763640] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1565763772] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6730100 kB) free!
[1565763794] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565763892] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11307280 kB) free.
[1565763914] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1565764109] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.01, 2.54, 3.83
[1565764492] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7445016 kB) free!
[1565764514] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565764612] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.2% (6028824 kB) free!
[1565764634] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565764732] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11304800 kB) free.
[1565765155] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1197.86, 817.83, 370.50
[1565765183] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 644 processes
[1565765194] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 68 users currently logged in
[1565765235] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565765246] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1167.20, 687.82, 289.51
[1565765275] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1200.81, 944.71, 471.23
[1565765303] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 639 processes
[1565765332] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6747676 kB) free!
[1565765355] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565765366] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1180.37, 851.46, 397.77
[1565765395] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1201.67, 1029.79, 559.82
[1565765423] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 401 processes
[1565765452] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11294072 kB) free.
[1565765475] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565765486] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1182.35, 961.06, 492.94
[1565766053] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7536832 kB) free!
[1565766075] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1565766173] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.6% (6226308 kB) free!
[1565766195] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565766224] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (11691840 kB) free!
[1565766293] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11129588 kB) free.
[1565766315] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565766640] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 367 processes
[1565766760] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 364 processes
[1565766880] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 369 processes
[1565767109] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.58, 7.33, 5.72
[1565767114] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.23, 5.50, 5.26
[1565767229] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.68, 7.42, 5.95
[1565767234] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.60, 5.87, 5.42
[1565767349] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.13, 6.44, 5.79
[1565767355] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.64, 6.61, 5.75
[1565767480] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1565767480] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 368 processes
[1565767515] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1565767949] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1565767949] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.40, 2.88, 4.30
[1565768080] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 3 warn out of 374 processes
[1565768093] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.7% (5788752 kB) free!
[1565768201] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 3 warn out of 376 processes
[1565768213] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10629672 kB) free.
[1565768321] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 370 processes
[1565768555] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.26, 3.90, 4.82
[1565768715] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565768813] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9333040 kB) free!
[1565768835] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565768921] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1565768933] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (8007368 kB) free!
[1565768955] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565768995] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.43, 2.68, 14.66
[1565769053] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.8% (6283656 kB) free!
[1565769086] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.64, 0.59, 12.04
[1565769653] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11124372 kB) free.
[1565769755] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.34, 5.95, 5.23
[1565769875] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.80, 5.24, 5.07
[1565769995] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.73, 4.49, 4.83
[1565770155] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1565770155] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565770286] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.85, 0.70, 3.76
[1565770795] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.40, 2.97, 4.61
[1565771024] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 38.4% (9466444 kB) free.
[1565771956] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565772076] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565772224] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188868 kB) free!
[1565772344] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173256 kB) free!
[1565772464] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180524 kB) free!
[1565772621] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 397 processes
[1565772653] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7097076 kB) free!
[1565772676] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565772741] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 397 processes
[1565772773] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11270840 kB) free.
[1565772796] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565773996] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565774116] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565774132] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9059904 kB) free!
[1565774236] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565774252] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8869780 kB) free!
[1565774372] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.6% (8649488 kB) free!
[1565774573] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9238328 kB) free!
[1565774693] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7897332 kB) free!
[1565774813] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.4% (6122092 kB) free!
[1565775413] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1565775413] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11048960 kB) free.
[1565775442] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.4% (14537020 kB) free.
[1565775491] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 75.3% (37108852 kB) free.
[1565775528] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.3% (22784448 kB) free.
[1565775577] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 53.1% (69969124 kB) free.
[1565775621] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 87.9% (21532760 kB) free.
[1565776013] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9453716 kB) free!
[1565776133] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173208 kB) free!
[1565776134] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8019600 kB) free!
[1565776199] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1565776253] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182804 kB) free!
[1565776254] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.9% (6357788 kB) free!
[1565776373] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1565776373] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182960 kB) free!
[1565776655] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (173268 kB) free!
[1565776721] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 365 processes
[1565776772] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4423452 kB) free!
[1565776775] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (163916 kB) free!
[1565776841] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 372 processes
[1565776854] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11231112 kB) free.
[1565776895] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (150808 kB) free!
[1565777324] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4864092 kB) free!
[1565777441] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 1 warn out of 372 processes
[1565777444] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (4827544 kB) free!
[1565777517] SERVICE ALERT: sig4-vm2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 279 processes
[1565777562] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 2 warn out of 374 processes
[1565777564] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (4792632 kB) free!
[1565777637] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 272 processes
[1565777682] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 374 processes
[1565777775] SERVICE ALERT: webservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.66, 5.52, 5.17
[1565777836] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565777895] SERVICE ALERT: webservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.10, 5.71, 5.28
[1565778015] SERVICE ALERT: webservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.16, 5.81, 5.37
[1565778149] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.22, 6.82, 5.61
[1565778269] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.62, 6.49, 5.64
[1565778282] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 370 processes
[1565778390] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.19, 6.35, 5.69
[1565778436] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565778491] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (2278428 kB) free!
[1565778556] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1565778611] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (213236 kB) free!
[1565778731] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (214964 kB) free!
[1565778882] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 3 warn out of 371 processes
[1565779002] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 3 warn out of 372 processes
[1565779122] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 368 processes
[1565779156] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565779215] SERVICE ALERT: webservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.00, 4.21, 4.76
[1565779254] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6087592 kB) free!
[1565779277] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565779374] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11231760 kB) free.
[1565779876] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1565779974] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.3% (6071272 kB) free!
[1565779996] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565780087] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1565780094] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11231656 kB) free.
[1565780394] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1565780596] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565780694] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6818320 kB) free!
[1565780716] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565780814] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11229096 kB) free.
[1565780836] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565780922] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 366 processes
[1565781190] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.2% (14171476 kB) free.
[1565781194] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.38, 5.33, 5.10
[1565781219] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 43.5% (21461764 kB) free.
[1565781256] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 59.1% (29131668 kB) free.
[1565781314] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.54, 5.40, 5.16
[1565781334] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.4% (23296060 kB) free.
[1565781415] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7362788 kB) free!
[1565781434] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.37, 5.01, 5.04
[1565781436] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565781523] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 370 processes
[1565781535] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11233832 kB) free.
[1565781557] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565781599] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1565781643] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 3 warn out of 376 processes
[1565781763] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 372 processes
[1565781791] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (167200 kB) free!
[1565781911] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (173412 kB) free!
[1565781942] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187500 kB) free!
[1565782031] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1565782031] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (176228 kB) free!
[1565782034] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.50, 4.78, 4.93
[1565782062] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168976 kB) free!
[1565782135] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7785088 kB) free!
[1565782156] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565782183] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164860 kB) free!
[1565782255] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6097096 kB) free!
[1565782276] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1565782363] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 372 processes
[1565782375] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10803248 kB) free.
[1565782396] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565782794] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1565782964] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2025600 kB) free!
[1565782975] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9716072 kB) free!
[1565783095] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.9% (8337472 kB) free!
[1565783106] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.2% (7014700 kB) free!
[1565783194] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 64 users currently logged in
[1565783215] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.0% (6899684 kB) free!
[1565783226] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (4760120 kB) free!
[1565783346] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1565783346] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3611544 kB) free!
[1565783563] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 1 warn out of 371 processes
[1565783683] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 1 warn out of 369 processes
[1565783803] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 374 processes
[1565783815] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11220780 kB) free.
[1565784390] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1565784694] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1565785003] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 372 processes
[1565785015] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5939976 kB) free!
[1565785135] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11221340 kB) free.
[1565785590] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.20, 3.24, 4.31
[1565785736] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7623368 kB) free!
[1565785856] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.8% (6326724 kB) free!
[1565785976] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11075140 kB) free.
[1565786190] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.43, 6.82, 5.37
[1565786195] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.79, 5.70, 5.25
[1565786203] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 3 warn out of 378 processes
[1565786310] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.25, 8.08, 6.02
[1565786315] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.57, 5.76, 5.33
[1565786323] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 3 warn out of 372 processes
[1565786430] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.09, 8.79, 6.53
[1565786435] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.40, 5.61, 5.33
[1565786443] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 4 warn out of 378 processes
[1565787043] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 6 warn out of 378 processes
[1565787415] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176384 kB) free!
[1565787535] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163292 kB) free!
[1565787655] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164112 kB) free!
[1565788243] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 376 processes
[1565788376] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9169432 kB) free!
[1565788496] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7846996 kB) free!
[1565788616] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.8% (6291032 kB) free!
[1565788830] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1565788830] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.43, 3.47, 4.98
[1565788894] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 65 users currently logged in
[1565789025] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.7% (8723392 kB) free!
[1565789041] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 660 processes
[1565789087] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1565789101] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 661 processes
[1565789145] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.1% (8440404 kB) free!
[1565789208] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1565789216] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11193036 kB) free.
[1565789264] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 13.2% (3241700 kB) free!
[1565789265] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.5% (8133100 kB) free!
[1565789325] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 12.4% (3045464 kB) free!
[1565789328] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1565789385] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 12.5% (3060292 kB) free!
[1565789445] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 12.5% (3061036 kB) free!
[1565790196] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1565790417] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8776628 kB) free!
[1565790537] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.9% (7355116 kB) free!
[1565790657] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11178368 kB) free.
[1565790797] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565790917] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565790945] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 26.4% (6475516 kB) free.
[1565791037] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565791065] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3381172 kB) free!
[1565791294] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 55 users currently logged in
[1565791358] SERVICE ALERT: sigqscluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1565791478] SERVICE ALERT: sigqscluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1565791598] SERVICE ALERT: sigqscluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1565791857] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9466132 kB) free!
[1565791941] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 397 processes
[1565791977] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8202484 kB) free!
[1565792061] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 398 processes
[1565792097] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.0% (6900852 kB) free!
[1565792198] SERVICE ALERT: sigqscluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1565792383] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1565792697] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11171500 kB) free.
[1565792946] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1565793431] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1565793897] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9243184 kB) free!
[1565794017] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8008264 kB) free!
[1565794137] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.4% (6625564 kB) free!
[1565794737] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11150908 kB) free.
[1565795237] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1565795837] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1565795937] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9223964 kB) free!
[1565795957] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1565796057] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7925156 kB) free!
[1565796077] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1565796177] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.9% (6378632 kB) free!
[1565796777] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11134672 kB) free.
[1565796861] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 398 processes
[1565796981] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 398 processes
[1565797377] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9807968 kB) free!
[1565797497] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.4% (8597712 kB) free!
[1565797617] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.9% (7318728 kB) free!
[1565798217] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11129712 kB) free.
[1565798431] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1565798435] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.63, 3.67, 4.55
[1565798478] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1565798486] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 3736.43, 1399.76, 510.85
[1565798606] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4291.79, 2385.30, 979.77
[1565798726] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4366.95, 3044.79, 1391.86
[1565798817] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7644428 kB) free!
[1565798937] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.7% (5780168 kB) free!
[1565799043] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1565799055] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.0% (2931236 kB) free!
[1565799057] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11120960 kB) free.
[1565799658] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.0% (5913620 kB) free!
[1565799778] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11128000 kB) free.
[1565800278] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1565800378] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6108808 kB) free!
[1565800398] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1565800498] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.7% (10670008 kB) free.
[1565800518] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565800855] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2333452 kB) free!
[1565801098] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9530076 kB) free!
[1565801218] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8156636 kB) free!
[1565801338] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.6% (6706392 kB) free!
[1565801938] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11119700 kB) free.
[1565803526] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.86, 0.73, 18.48
[1565804338] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9734540 kB) free!
[1565804459] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.2% (8491680 kB) free!
[1565804579] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.7% (7231900 kB) free!
[1565804718] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1565805179] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11100756 kB) free.
[1565805326] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.77, 0.63, 3.20
[1565805595] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 700.20, 286.16, 108.28
[1565805715] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 771.15, 450.24, 189.97
[1565805779] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7778044 kB) free!
[1565805835] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 780.23, 559.89, 261.71
[1565805899] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6088072 kB) free!
[1565805918] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1565806019] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11100448 kB) free.
[1565806038] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1565806619] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6545536 kB) free!
[1565806638] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1565806739] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11103360 kB) free.
[1565806759] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1565807339] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7009888 kB) free!
[1565807358] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes