[1573151787] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184504 kB) free!
[1573151870] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199944 kB) free!
[1573151889] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23457996 kB) free!
[1573151990] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172512 kB) free!
[1573152002] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (202356 kB) free!
[1573152009] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.4% (44079056 kB) free.
[1573152026] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1573152026] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.6% (13998988 kB) free!
[1573152110] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190196 kB) free!
[1573152122] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (201932 kB) free!
[1573152245] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193752 kB) free!
[1573153227] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (5704092 kB) free!
[1573153570] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167816 kB) free!
[1573153630] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 58.3% (14284176 kB) free.
[1573153809] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23222004 kB) free!
[1573153827] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.8% (14218676 kB) free!
[1573153930] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.3% (43945012 kB) free.
[1573155729] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22968280 kB) free!
[1573155736] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.6% (8159800 kB) free!
[1573155794] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573155849] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.0% (43490016 kB) free.
[1573155856] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.8% (7805516 kB) free!
[1573155976] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.2% (7481792 kB) free!
[1573156576] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (3109492 kB) free!
[1573157649] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23277412 kB) free!
[1573157769] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43404460 kB) free.
[1573160159] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (231204 kB) free!
[1573160437] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (221188 kB) free!
[1573160557] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (229376 kB) free!
[1573160677] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (239040 kB) free!
[1573160759] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 77.3% (38109460 kB) free.
[1573161873] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1573161877] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.3% (45977008 kB) free.
[1573161969] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23943224 kB) free!
[1573161993] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1573162090] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43352196 kB) free.
[1573163029] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 54.7% (13491724 kB) free.
[1573164489] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25472340 kB) free!
[1573164609] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43646220 kB) free.
[1573164848] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 70.81, 474.64, 330.38
[1573164968] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 10.18, 317.82, 290.40
[1573164975] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(33) < 2019/11/08 03:44:06 [error] 24130#0: *3796751 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1573165095] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573165227] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573167009] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25357436 kB) free!
[1573167127] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1573167129] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.1% (43627132 kB) free.
[1573167247] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1573167367] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1573167495] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(32) < 2019/11/08 04:28:12 [error] 24132#0: *3810619 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:85/0, bytes from/to upstream:0/85
[1573167615] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(16) < 2019/11/08 04:29:41 [error] 24130#0: *3811208 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.182.41, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:37/0, bytes from/to upstream:0/37
[1573167736] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(2) < 2019/11/08 04:31:15 [error] 24130#0: *3811840 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:39/0, bytes from/to upstream:0/39
[1573167967] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1573167987] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.5% (2825060 kB) free!
[1573169613] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(85) < 2019/11/08 04:56:30 [error] 23648#0: *1778782 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:56/17, bytes from/to upstream:17/56
[1573169734] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573169793] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1573169913] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1573170387] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2427612 kB) free!
[1573170933] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(49) < 2019/11/08 05:17:05 [error] 23641#0: *1780452 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:88/49, bytes from/to upstream:49/88
[1573170987] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (2516332 kB) free!
[1573171053] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573171310] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.2% (2741136 kB) free!
[1573171587] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1573171587] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2325144 kB) free!
[1573171929] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22470608 kB) free!
[1573172045] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.4% (2789348 kB) free!
[1573172049] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.1% (42368320 kB) free.
[1573172313] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1573172433] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1573172537] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573173034] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 410 processes
[1573173110] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180436 kB) free!
[1573173154] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 409 processes
[1573173230] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (158244 kB) free!
[1573173245] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186572 kB) free!
[1573173351] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182200 kB) free!
[1573173471] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168408 kB) free!
[1573173627] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (9125224 kB) free!
[1573173849] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22084752 kB) free!
[1573173969] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42545568 kB) free.
[1573174227] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.3% (33361448 kB) free.
[1573174310] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.1% (20849492 kB) free.
[1573174587] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.2% (9845808 kB) free.
[1573175045] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 83.9% (20553392 kB) free.
[1573175290] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 248868 MB (23% inode=99%):
[1573177910] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9759700 kB) free!
[1573178030] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.7% (9732088 kB) free!
[1573178150] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.7% (9695180 kB) free!
[1573178169] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (25196372 kB) free!
[1573178289] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.7% (21966900 kB) free!
[1573178410] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42231672 kB) free.
[1573179009] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25797196 kB) free!
[1573179129] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.2% (22643144 kB) free!
[1573179249] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (42100112 kB) free.
[1573180550] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 20.1% (9924472 kB) free.
[1573181150] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9759564 kB) free!
[1573181270] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.7% (9730508 kB) free!
[1573181390] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.4% (9548324 kB) free!
[1573182216] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1906 processes
[1573182336] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1906 processes
[1573182457] SERVICE ALERT: smscapp2;Total Processes;WARNING;HARD;3;PROCS WARNING: 1904 processes
[1573183190] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1573183190] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 63.9% (31473708 kB) free.
[1573183449] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24353352 kB) free!
[1573183569] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.9% (21010372 kB) free!
[1573183588] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1573183689] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41357968 kB) free.
[1573184289] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21137812 kB) free!
[1573184409] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.0% (17134168 kB) free!
[1573184530] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40694096 kB) free.
[1573185129] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (23010612 kB) free!
[1573185249] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.8% (19460664 kB) free!
[1573185369] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15811316 kB) free!
[1573186501] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.4% (22319108 kB) free.
[1573186534] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.8% (40320484 kB) free.
[1573186569] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30213880 kB) free.
[1573187187] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4893272 kB) free!
[1573187307] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (4788400 kB) free!
[1573187427] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (4620244 kB) free!
[1573187769] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16401908 kB) free!
[1573187889] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11871852 kB) free!
[1573188009] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (7659116 kB) free!
[1573188317] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 357 processes
[1573188437] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1573188609] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1573188609] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17138192 kB) free!
[1573188909] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.7% (4365112 kB) free!
[1573189029] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.8% (3405968 kB) free!
[1573189149] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.4% (3308580 kB) free!
[1573189209] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27712796 kB) free.
[1573189227] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2276688 kB) free!
[1573189749] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180516 kB) free!
[1573189809] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24172764 kB) free!
[1573189929] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32884544 kB) free.
[1573190530] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9584264 kB) free!
[1573190650] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (5698324 kB) free!
[1573190770] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34166328 kB) free.
[1573190832] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 373 processes
[1573190952] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 448 processes
[1573191137] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(43) < 2019/11/08 10:59:50 [error] 24131#0: *4001170 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573191257] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573191370] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9908740 kB) free!
[1573191490] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (5582056 kB) free!
[1573191610] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1437120 kB) free!
[1573192810] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.1% (22596124 kB) free!
[1573192819] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.5% (2800556 kB) free!
[1573192939] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 10.6% (2369724 kB) free!
[1573193059] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1931744 kB) free!
[1573193390] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1573193410] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393832 kB) free!
[1573193645] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (4365056 kB) free!
[1573193765] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.1% (4181624 kB) free!
[1573193885] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.3% (3983584 kB) free!
[1573193987] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 18.2% (6324568 kB) free!
[1573194107] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 16.6% (5766360 kB) free!
[1573194227] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 15.0% (5209892 kB) free!
[1573194257] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(33) < 2019/11/08 11:53:56 [error] 24134#0: *4028346 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:83/0, bytes from/to upstream:0/83
[1573194367] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 384 processes
[1573194377] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2) < 2019/11/08 11:55:59 [error] 24128#0: *4028774 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:83/0, bytes from/to upstream:0/83
[1573194487] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 428 processes
[1573194497] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(11) < 2019/11/08 11:57:16 [error] 24132#0: *4029020 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:87/0, bytes from/to upstream:0/87
[1573194610] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20530960 kB) free!
[1573194627] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (2516784 kB) free!
[1573194828] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 92.9% (32353072 kB) free.
[1573194859] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 83.2% (18692108 kB) free.
[1573195085] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1810088 kB) free!
[1573195097] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573195210] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (442744 kB) free!
[1573195227] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1573195227] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (1037568 kB) free!
[1573195752] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1573195872] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 362 processes
[1573195910] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.7% (4096284 kB) free!
[1573195992] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1573196030] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.6% (3581080 kB) free!
[1573196150] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 24.1% (5908796 kB) free.
[1573196298] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(39) < 2019/11/08 12:25:00 [error] 24129#0: *4034402 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:83/0, bytes from/to upstream:0/83
[1573196410] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29622296 kB) free.
[1573196418] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573196592] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1573196677] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9529736 kB) free!
[1573196737] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9699660 kB) free!
[1573196797] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9101132 kB) free!
[1573196857] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (9690492 kB) free!
[1573196857] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 2045 processes
[1573196918] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.6% (8676392 kB) free!
[1573196977] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.7% (9691148 kB) free!
[1573197011] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (4710876 kB) free!
[1573197131] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (724556 kB) free!
[1573197251] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (471544 kB) free!
[1573197950] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4797264 kB) free!
[1573198071] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4762880 kB) free!
[1573198118] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4228228 kB) free!
[1573198177] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (239232 kB) free!
[1573198191] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (4730724 kB) free!
[1573198451] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26588524 kB) free.
[1573198818] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(8) < 2019/11/08 13:09:09 [error] 24130#0: *4041510 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.182.41, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:37/0, bytes from/to upstream:0/37
[1573198938] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(4) < 2019/11/08 13:10:53 [error] 24130#0: *4041785 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.182.41, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:37/0, bytes from/to upstream:0/37
[1573199051] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (373968 kB) free!
[1573199059] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(5) < 2019/11/08 13:14:15 [error] 24132#0: *4042315 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.182.41, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:37/0, bytes from/to upstream:0/37
[1573199171] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (403656 kB) free!
[1573199291] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (465388 kB) free!
[1573199795] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < operation none caused a connection exception frame_error: "Malformed UTF-8 in shortstr"
[1573199916] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573200393] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 64.4% (15898204 kB) free.
[1573200419] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.1% (21827540 kB) free.
[1573200454] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.4% (40132576 kB) free.
[1573200484] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 55.6% (27398212 kB) free.
[1573200491] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29406660 kB) free.
[1573200515] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.2998.14> (10.147.212.42:55642 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1573200593] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 49.2% (12052552 kB) free.
[1573200635] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573200999] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171880 kB) free!
[1573201091] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (2508792 kB) free!
[1573201119] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191160 kB) free!
[1573201211] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (458840 kB) free!
[1573201239] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182748 kB) free!
[1573201331] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (441924 kB) free!
[1573201459] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (1734968 kB) free!
[1573201565] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1105 processes
[1573201579] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (1234984 kB) free!
[1573201685] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 487 processes
[1573201699] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (741984 kB) free!
[1573201720] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 623 processes
[1573201840] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 637 processes
[1573201960] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 396 processes
[1573202391] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2195636 kB) free!
[1573202531] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29183788 kB) free.
[1573202759] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9550436 kB) free!
[1573202853] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(162) < 2019/11/08 14:17:06 [error] 23634#0: *1880108 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573202856] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (224592 kB) free!
[1573202879] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (9327876 kB) free!
[1573202973] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(73) < 2019/11/08 14:17:42 [error] 23633#0: *1880437 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:74/0, bytes from/to upstream:0/74
[1573202977] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (259156 kB) free!
[1573202999] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (9135236 kB) free!
[1573203027] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.0% (23732576 kB) free!
[1573203093] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573203097] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1573203097] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (218668 kB) free!
[1573203132] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (389260 kB) free!
[1573203147] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.3% (17473688 kB) free!
[1573203252] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (442496 kB) free!
[1573203260] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1573203260] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573203267] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.2% (31889624 kB) free.
[1573203372] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (435992 kB) free!
[1573203596] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4648008 kB) free!
[1573203716] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (4424180 kB) free!
[1573203836] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.3% (4237416 kB) free!
[1573203867] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 10.9% (14391984 kB) free!
[1573203987] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 16.0% (21133608 kB) free!
[1573204026] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.207.436> (10.147.212.158:36812 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1573204108] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.3% (20130924 kB) free!
[1573204147] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573204428] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 15.5% (5387928 kB) free!
[1573204548] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 14.2% (4932944 kB) free!
[1573204572] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30051360 kB) free.
[1573204668] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 12.8% (4468424 kB) free!
[1573204708] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.7% (31192796 kB) free.
[1573204799] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4751784 kB) free!
[1573205025] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 17.3% (6046736 kB) free!
[1573205036] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1735560 kB) free!
[1573205145] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 16.1% (5626660 kB) free!
[1573205172] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2878608 kB) free!
[1573205175] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 791 processes
[1573205265] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 14.9% (5202520 kB) free!
[1573205268] SERVICE FLAPPING ALERT: sig2-vm2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1573205268] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (1932416 kB) free!
[1573205292] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (448488 kB) free!
[1573205296] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 793 processes
[1573205308] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.1% (15949148 kB) free!
[1573205408] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 53.61, 43.22, 22.22
[1573205412] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16628516 kB) free!
[1573205416] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 791 processes
[1573205428] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.4% (30841796 kB) free.
[1573205528] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 45.92, 44.00, 25.06
[1573205648] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 30.69, 39.42, 25.71
[1573205865] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.9% (31825120 kB) free.
[1573205869] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 93.3% (32514924 kB) free.
[1573205899] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 83.5% (18742348 kB) free.
[1573206012] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (11399756 kB) free!
[1573206028] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.2% (18738164 kB) free!
[1573206148] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.9% (30190084 kB) free.
[1573206248] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 10.49, 13.37, 17.87
[1573206260] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(85) < 2019/11/08 15:12:45 [error] 24127#0: *4083494 recv() failed (104: Connection reset by peer) while proxying connection, client: 202.65.140.29, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:12998/199313, bytes from/to upstream:199313/12998
[1573206380] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573206612] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13800200 kB) free!
[1573207212] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (457144 kB) free!
[1573207948] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.8% (20778892 kB) free!
[1573208068] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.8% (30053824 kB) free.
[1573208180] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(190) < 2019/11/08 15:46:03 [error] 24130#0: *4098862 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.182.41, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:37/0, bytes from/to upstream:0/37
[1573208288] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.2% (8479780 kB) free!
[1573208300] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(6) < 2019/11/08 15:47:41 [error] 24125#0: *4099854 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:195/17, bytes from/to upstream:17/195
[1573208409] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.5% (8122852 kB) free!
[1573208412] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27542056 kB) free.
[1573208420] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573208529] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.8% (7778488 kB) free!
[1573208969] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 55.7% (13749572 kB) free.
[1573208989] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 75.1% (37026584 kB) free.
[1573209012] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3340692 kB) free!
[1573209097] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.4% (38656832 kB) free.
[1573209129] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1573209129] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 53.0% (26127616 kB) free.
[1573209133] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (392248 kB) free!
[1573209196] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 44.5% (10895724 kB) free.
[1573209211] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.1% (22824260 kB) free.
[1573209253] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35555360 kB) free.
[1573209575] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186204 kB) free!
[1573209695] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191556 kB) free!
[1573209697] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (214340 kB) free!
[1573209815] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1573209815] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182436 kB) free!
[1573209817] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (227932 kB) free!
[1573209937] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (242016 kB) free!
[1573210028] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4701692 kB) free!
[1573210148] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (4486596 kB) free!
[1573210268] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.4% (4273000 kB) free!
[1573210820] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(18) < 2019/11/08 16:23:28 [error] 24133#0: *4116466 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1573210940] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573211065] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11559532 kB) free!
[1573211068] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.5% (23093088 kB) free!
[1573211185] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7343980 kB) free!
[1573211188] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.9% (34205864 kB) free.
[1573211305] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (40956536 kB) free.
[1573211540] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(12) < 2019/11/08 16:34:07 [error] 24133#0: *4117929 connect() failed (111: Connection refused) while connecting to upstream, client: 185.255.8.58, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573211596] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4857696 kB) free!
[1573211660] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1573211660] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573211716] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (4738576 kB) free!
[1573211836] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (4597808 kB) free!
[1573211905] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16527652 kB) free!
[1573212025] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12305604 kB) free!
[1573212069] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2182712 kB) free!
[1573212129] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9606480 kB) free!
[1573212145] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (8195772 kB) free!
[1573212250] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9067756 kB) free!
[1573212260] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(14) < 2019/11/08 16:51:46 [error] 24130#0: *4124618 connect() failed (111: Connection refused) while connecting to upstream, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573212370] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.3% (8503724 kB) free!
[1573212380] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573212388] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.9% (20942684 kB) free!
[1573212436] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 20.1% (4921588 kB) free.
[1573212508] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.6% (32486892 kB) free.
[1573212745] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.0% (21025708 kB) free!
[1573213036] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.6% (3822808 kB) free!
[1573213156] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.0% (3684456 kB) free!
[1573213276] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.5% (3559116 kB) free!
[1573213345] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35826420 kB) free.
[1573213570] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2938444 kB) free!
[1573213708] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.5% (17732828 kB) free!
[1573213823] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 103.255.146.164 [08/Nov/2019:17:17:13 +0530] TCP 502 0 0 0.001 "10.147.212.37:5001, 10.147.212.49:5001, 10.147.212.38:5001" "0, 0, 0" "0, 0, 0" "0.000, 0.001, 0.000"
[1573213828] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.0% (30383520 kB) free.
[1573213893] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/11/08 17:17:13 [error] 23647#0: *1932840 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:14001, upstream: "10.147.212.38:5001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573213943] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573213945] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11545884 kB) free!
[1573214013] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573214066] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7240164 kB) free!
[1573214186] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39649164 kB) free.
[1573214476] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2325448 kB) free!
[1573214544] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 103.255.146.164 [08/Nov/2019:17:23:52 +0530] TCP 502 0 0 0.003 "10.147.212.37:5001, 10.147.212.49:5001, 10.147.212.38:5001" "0, 0, 0" "0, 0, 0" "0.002, 0.001, 0.000"
[1573214613] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/11/08 17:23:52 [error] 23644#0: *1934767 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:14001, upstream: "10.147.212.38:5001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573214664] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573214733] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573214786] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16211856 kB) free!
[1573214906] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (11399668 kB) free!
[1573215026] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (6857688 kB) free!
[1573215626] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19860628 kB) free!
[1573215628] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.0% (22358204 kB) free!
[1573215748] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 27.1% (35775280 kB) free.
[1573216168] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 1768 processes
[1573216226] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35894336 kB) free.
[1573216826] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11530004 kB) free!
[1573216833] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 516 processes
[1573216946] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7196156 kB) free!
[1573216948] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.5% (20375412 kB) free!
[1573217066] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (3114224 kB) free!
[1573217068] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.9% (32834228 kB) free.
[1573217270] SERVICE FLAPPING ALERT: sig2-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573217648] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.36, 3.58, 4.88
[1573217666] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17864032 kB) free!
[1573218266] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30839280 kB) free.
[1573218268] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.8% (18202604 kB) free!
[1573218388] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.1% (30482296 kB) free.
[1573218601] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8814096 kB) free!
[1573218721] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.1% (8428764 kB) free!
[1573218758] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1573218782] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1573218836] SERVICE ALERT: smscapp1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1573218841] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 16.3% (8054844 kB) free!
[1573218848] SERVICE ALERT: smscapp1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1573218860] SERVICE ALERT: smscapp1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1573218866] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8480288 kB) free!
[1573218888] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1573218912] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1573218966] SERVICE ALERT: smscapp1;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1573218967] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1573218978] SERVICE ALERT: smscapp1;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1573218986] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4548972 kB) free!
[1573218990] SERVICE ALERT: smscapp1;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1573219018] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1573219042] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1573219075] SERVICE ALERT: smscapp1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1573219085] SERVICE ALERT: smscapp1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1573219096] SERVICE ALERT: smscapp1;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1573219098] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1573219106] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35404256 kB) free.
[1573219108] SERVICE ALERT: smscapp1;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1573219120] SERVICE ALERT: smscapp1;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1573219205] SERVICE ALERT: smscapp1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1573219215] SERVICE ALERT: smscapp1;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1573219228] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1573219335] SERVICE ALERT: smscapp1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1573219345] SERVICE ALERT: smscapp1;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1573219707] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12646516 kB) free!
[1573219827] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8514896 kB) free!
[1573219947] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (4412308 kB) free!
[1573220041] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3895748 kB) free!
[1573220236] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.3% (21056644 kB) free.
[1573220244] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.3% (22380656 kB) free.
[1573220259] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 29.5% (14543908 kB) free.
[1573220276] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 43.2% (56922948 kB) free.
[1573220286] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 65.5% (32280696 kB) free.
[1573220380] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 49.9% (12223280 kB) free.
[1573220836] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185128 kB) free!
[1573220842] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 667 processes
[1573220887] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236988 kB) free!
[1573220956] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192892 kB) free!
[1573221007] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (214592 kB) free!
[1573221077] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188300 kB) free!
[1573221127] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (215348 kB) free!
[1573221452] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1573221582] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1573221583] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1573221705] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573221712] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1573222068] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7371456 kB) free!
[1573222188] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.7% (6735348 kB) free!
[1573222308] SERVICE FLAPPING ALERT: qservercluster2;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1573222308] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 12.4% (6127800 kB) free!
[1573222676] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18072856 kB) free!
[1573222796] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.9% (14313880 kB) free!
[1573222835] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.7620.142> (10.147.212.53:35570 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1573222908] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (3034276 kB) free!
[1573222913] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1573222913] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 329 processes
[1573222916] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (10630088 kB) free!
[1573222945] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 0.56, 0.38, 0.15
[1573222955] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573223170] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 97.0% (47796868 kB) free.
[1573223219] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 116873 MB (52% inode=99%):
[1573223296] SERVICE ALERT: smscapp1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1573223308] SERVICE ALERT: smscapp1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 39127 MB (76% inode=98%):
[1573223320] SERVICE ALERT: smscapp1;Sensors Status;OK;HARD;1;SENSORS OK
[1573223428] SERVICE ALERT: smscapp1;CPU_Procs;OK;HARD;1;CPU OK: 607 processes
[1573223516] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.2% (23940140 kB) free!
[1573224077] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 74.8% (18461324 kB) free.
[1573224104] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < {writer,send_failed,{error,timeout}}
[1573224107] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 14.0% (6897460 kB) free!
[1573224116] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 46.4% (61178252 kB) free.
[1573224127] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 72.3% (35631296 kB) free.
[1573224226] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573224678] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191280 kB) free!
[1573224728] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (1053424 kB) free!
[1573224798] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194860 kB) free!
[1573224848] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (247004 kB) free!
[1573224918] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188140 kB) free!
[1573224968] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (230960 kB) free!
[1573225189] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1573225956] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4650304 kB) free!
[1573226516] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25560064 kB) free!
[1573226579] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (468684 kB) free!
[1573226636] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.6% (21896196 kB) free!
[1573226699] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (467584 kB) free!
[1573226757] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18312432 kB) free!
[1573226819] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (466332 kB) free!
[1573226975] SERVICE ALERT: smscapp1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1573227333] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(36) < 2019/11/08 21:00:10 [error] 23648#0: *1981523 connect() failed (111: Connection refused) while connecting to upstream, client: 49.50.67.175, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573227357] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28898412 kB) free.
[1573227453] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573227807] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.7% (14576312 kB) free.
[1573227826] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 64.5% (15925212 kB) free.
[1573227896] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.0% (38945564 kB) free.
[1573228426] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193140 kB) free!
[1573228497] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (239008 kB) free!
[1573228540] SERVICE FLAPPING ALERT: qservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1573228546] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188848 kB) free!
[1573228557] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (23019576 kB) free!
[1573228580] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/08 21:26:03 [error] 24130#0: *4254220 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1573228617] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (263212 kB) free!
[1573228666] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (542440 kB) free!
[1573228677] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.8% (19567348 kB) free!
[1573228700] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(34) < 2019/11/08 21:28:05 [error] 24130#0: *4254262 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573228737] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (261248 kB) free!
[1573228797] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16316568 kB) free!
[1573228821] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573229170] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1573229389] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.6% (4556636 kB) free!
[1573229397] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27823984 kB) free.
[1573229509] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (4475796 kB) free!
[1573229629] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.9% (4391816 kB) free!
[1573230020] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(20) < 2019/11/08 21:49:20 [error] 24129#0: *4258311 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.182.41, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:37/0, bytes from/to upstream:0/37
[1573230140] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(5) < 2019/11/08 21:51:56 [error] 24130#0: *4258824 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.182.41, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:37/0, bytes from/to upstream:0/37
[1573230260] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(5) < 2019/11/08 21:53:41 [error] 24130#0: *4259200 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.182.41, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:37/0, bytes from/to upstream:0/37
[1573230597] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (23039700 kB) free!
[1573230717] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.0% (19794316 kB) free!
[1573230837] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16506236 kB) free!
[1573231437] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28538716 kB) free.
[1573232637] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (22704508 kB) free!
[1573232757] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.7% (19443056 kB) free!
[1573232877] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16220192 kB) free!
[1573233478] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26900384 kB) free.
[1573233714] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1573233829] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2312836 kB) free!
[1573234054] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(120) < 2019/11/08 22:57:03 [error] 23644#0: *1990385 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:56/17, bytes from/to upstream:17/56
[1573234174] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573234588] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7722188 kB) free!
[1573234678] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20659488 kB) free!
[1573234690] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 199914 MB (18% inode=99%):
[1573234708] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (746916 kB) free!
[1573234798] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.2% (17356520 kB) free!
[1573234810] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 189685 MB (17% inode=99%):
[1573234828] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (3850700 kB) free!
[1573234918] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41982824 kB) free.
[1573234930] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 179455 MB (16% inode=99%):
[1573235029] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 48.7% (11948216 kB) free.
[1573235060] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573235088] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 78.9% (104052908 kB) free.
[1573235235] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.1% (38974536 kB) free.
[1573235266] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 58.3% (14389512 kB) free.
[1573235518] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25598624 kB) free!
[1573235638] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.0% (22390144 kB) free!
[1573235660] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(16) < 2019/11/08 23:16:06 [error] 24129#0: *4283429 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573235758] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19265492 kB) free!
[1573235781] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573235844] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (252480 kB) free!
[1573235866] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190856 kB) free!
[1573235964] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (224332 kB) free!
[1573235986] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183016 kB) free!
[1573236085] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (220208 kB) free!
[1573236106] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194108 kB) free!
[1573236358] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30572012 kB) free.
[1573237740] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.9% (39356388 kB) free.
[1573237797] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.4% (22073076 kB) free.