[1549305879] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1549305999] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1549306538] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4926796 kB) free!
[1549307138] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.0% (5443444 kB) free!
[1549307199] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1549307320] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1549307919] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1549308039] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1549308341] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1620336 kB) free!
[1549309011] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 327 processes
[1549309239] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1549309359] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1549309611] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 326 processes
[1549311117] SERVICE FLAPPING ALERT: webservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1549313559] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1549313679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1549313738] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4900440 kB) free!
[1549313799] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1549314338] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.3% (7532588 kB) free!
[1549316402] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 519 processes
[1549316462] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 515 processes
[1549316799] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1549316919] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1549317039] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1549319138] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4683112 kB) free!
[1549319738] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.2% (7508572 kB) free!
[1549320039] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1549320159] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1549320279] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1549320411] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 327 processes
[1549321011] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 327 processes
[1549322811] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 328 processes
[1549323279] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1549323399] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1549323411] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 326 processes
[1549323519] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1549325211] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 328 processes
[1549325811] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 327 processes
[1549326519] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549326639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549326759] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1549327612] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 327 processes
[1549328212] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 326 processes
[1549330012] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 328 processes
[1549330359] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1549330479] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1549330538] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4820560 kB) free!
[1549330599] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1549330612] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 327 processes
[1549331138] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.2% (7488232 kB) free!
[1549331199] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1549331319] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1549331439] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1549332338] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4256016 kB) free!
[1549332412] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 328 processes
[1549332938] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.7% (6768240 kB) free!
[1549333013] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 325 processes
[1549333239] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1549334213] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 328 processes
[1549334738] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4422028 kB) free!
[1549335338] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.7% (7250552 kB) free!
[1549335413] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 326 processes
[1549335639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1549335759] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1549335880] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1549336013] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 327 processes
[1549336613] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 327 processes
[1549337679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1549337799] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1549337919] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1549338413] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 328 processes
[1549338519] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1549338519] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1549338938] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4458244 kB) free!
[1549339013] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 327 processes
[1549339538] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.5% (7168552 kB) free!
[1549339613] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 328 processes
[1549339720] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1549339840] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1549339960] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1549339988] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.4% (6604660 kB) free!
[1549340213] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 326 processes
[1549340813] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 329 processes
[1549341160] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549341280] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1549341400] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1549341413] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 327 processes
[1549342000] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1549342013] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 327 processes
[1549342121] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1549342255] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 94.5% (15358944 kB) free.
[1549342614] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 327 processes
[1549342720] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1549342840] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1549343440] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549343560] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1549343714] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1549343834] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1549343954] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1549344160] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549344280] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549344401] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549344414] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 325 processes
[1549344911] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24290 MB (10% inode=99%):
[1549344938] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (301580 kB) free!
[1549345014] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 325 processes
[1549345538] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.9% (15705092 kB) free.
[1549346201] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1549346214] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 325 processes
[1549346801] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549346921] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549347041] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1549347415] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 324 processes
[1549348511] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 166011 MB (73% inode=99%):
[1549348842] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1549349442] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549349562] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549349682] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1549349787] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.91, 7.23, 5.68
[1549349907] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.99, 7.48, 5.96
[1549350027] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.00, 7.66, 6.21
[1549350415] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 328 processes
[1549351154] SERVICE ALERT: sig4;Memory;WARNING;SOFT;1;WARNING - 20.0% (3251952 kB) free!
[1549351274] SERVICE ALERT: sig4;Memory;WARNING;SOFT;2;WARNING - 20.0% (3251572 kB) free!
[1549351394] SERVICE ALERT: sig4;Memory;WARNING;HARD;3;WARNING - 20.0% (3250888 kB) free!
[1549351410] SERVICE ALERT: webservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1549351482] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1549351530] SERVICE ALERT: webservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1549351827] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.22, 4.02, 4.65
[1549352427] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.91, 8.42, 6.57
[1549352547] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.99, 8.95, 6.98
[1549352667] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.00, 9.29, 7.35
[1549353882] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549354002] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549354122] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549354722] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1549355215] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 327 processes
[1549355322] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549355442] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549355563] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549356383] SERVICE ALERT: sig2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3238344 kB) free!
[1549356415] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 329 processes
[1549356503] SERVICE ALERT: sig2;Memory;WARNING;SOFT;2;WARNING - 19.9% (3231780 kB) free!
[1549356585] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47177 MB (20% inode=99%):
[1549356623] SERVICE ALERT: sig2;Memory;WARNING;HARD;3;WARNING - 19.8% (3220980 kB) free!
[1549356705] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47048 MB (20% inode=99%):
[1549356825] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47040 MB (20% inode=99%):
[1549356938] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1549357289] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.9% (3653176 kB) free!
[1549357615] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 328 processes
[1549357963] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1549357988] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4561060 kB) free!
[1549358216] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 329 processes
[1549358563] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549358588] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.2% (6004372 kB) free!
[1549358683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549358803] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549358816] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 328 processes
[1549359403] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1549359416] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 327 processes
[1549359454] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1549359574] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1549359694] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1549360003] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549360016] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 329 processes
[1549360124] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549360244] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549360616] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 329 processes
[1549360844] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1549360988] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3966172 kB) free!
[1549361216] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 328 processes
[1549361444] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549361489] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2198468 kB) free!
[1549361564] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549361631] SERVICE ALERT: sig3;Memory;WARNING;SOFT;1;WARNING - 20.0% (3254516 kB) free!
[1549361684] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549361742] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 82.7% (13434408 kB) free.
[1549361751] SERVICE ALERT: sig3;Memory;WARNING;SOFT;2;WARNING - 20.0% (3242612 kB) free!
[1549361816] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 329 processes
[1549361872] SERVICE ALERT: sig3;Memory;WARNING;HARD;3;WARNING - 19.9% (3232200 kB) free!
[1549362788] SERVICE FLAPPING ALERT: db1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1549362788] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.6% (5231336 kB) free!
[1549362868] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.97, 3.43, 4.27
[1549362884] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1549363016] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 328 processes
[1549363294] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1549363389] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3607564 kB) free!
[1549363393] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 24.28, 9.99, 4.63
[1549363485] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549363513] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.66, 7.25, 4.28
[1549363605] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549363617] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 328 processes
[1549363725] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1549363894] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 18 users currently logged in
[1549364014] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 18 users currently logged in
[1549364134] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 18 users currently logged in
[1549364817] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 327 processes
[1549364954] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1549365417] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 328 processes
[1549366017] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 327 processes
[1549367217] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 329 processes
[1549367325] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1549367926] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549368046] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1549368417] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 329 processes
[1549369018] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 329 processes
[1549369246] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1549369366] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1549369389] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.5% (5167424 kB) free!
[1549369486] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1549369618] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 328 processes
[1549369989] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3423176 kB) free!
[1549370218] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 329 processes
[1549370354] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1549370448] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 41 users currently logged in
[1549370474] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1549370508] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 40 users currently logged in
[1549370589] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.7% (5292804 kB) free!
[1549371190] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4424668 kB) free!
[1549371418] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 327 processes
[1549371934] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1549372308] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 42 users currently logged in
[1549372368] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 42 users currently logged in
[1549372390] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.0% (5413472 kB) free!
[1549372428] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 42 users currently logged in
[1549372488] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 42 users currently logged in
[1549372534] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1549372618] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 329 processes
[1549372654] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1549372774] SERVICE FLAPPING ALERT: pservercluster2;Current Users;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1549372774] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1549373388] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 39 users currently logged in
[1549373686] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1549373806] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1549373819] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 327 processes
[1549373926] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1549373974] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1549374190] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3867896 kB) free!
[1549374419] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 328 processes
[1549374440] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 93.4% (15172668 kB) free.
[1549374574] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1549374656] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.5% (3168820 kB) free!
[1549374694] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1549374776] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.4% (3159112 kB) free!
[1549374790] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.3% (5587316 kB) free!
[1549374814] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1549374896] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.4% (3149664 kB) free!
[1549375019] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 327 processes
[1549375040] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (257248 kB) free!
[1549375126] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1549375161] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (150172 kB) free!
[1549375246] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1549375281] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156820 kB) free!
[1549375366] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1549375390] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4422816 kB) free!
[1549375990] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.9% (5348180 kB) free!
[1549376120] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1549376566] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1549376590] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4367840 kB) free!
[1549376819] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 328 processes
[1549377190] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.2% (5532244 kB) free!
[1549378019] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 327 processes
[1549378390] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4083304 kB) free!
[1549378991] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.5% (5675660 kB) free!
[1549379219] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 328 processes
[1549379474] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1549379594] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1549379714] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1549380166] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549380286] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549380406] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1549380419] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 327 processes
[1549380914] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1549381391] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4165204 kB) free!
[1549381882] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 93.4% (15186040 kB) free.
[1549381991] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.6% (5713000 kB) free!
[1549383420] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 325 processes
[1549384020] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 325 processes
[1549384391] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4238316 kB) free!
[1549384606] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1549384620] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 325 processes
[1549384726] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1549384846] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1549384991] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.6% (5737132 kB) free!
[1549385447] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1549385820] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 324 processes
[1549386214] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1549386814] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1549386934] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1549387020] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 326 processes
[1549387054] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1549387514] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1549387620] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 324 processes
[1549387634] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1549387754] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1549388354] SERVICE FLAPPING ALERT: db2;Current Users;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1549388354] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1549388854] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1549389647] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1549389768] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1549389888] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1549390020] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 324 processes
[1549390391] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4609964 kB) free!
[1549390620] SERVICE ALERT: db1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 323 processes
[1549390991] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.6% (5696528 kB) free!