[1587235196] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11144320 kB) free!
[1587235316] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (7108160 kB) free!
[1587235436] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35235144 kB) free.
[1587235771] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1587236036] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10772376 kB) free!
[1587236156] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (7142464 kB) free!
[1587236276] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35280812 kB) free.
[1587236641] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166052 kB) free!
[1587236702] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192420 kB) free!
[1587236762] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (171360 kB) free!
[1587236822] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (174708 kB) free!
[1587236877] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11430368 kB) free!
[1587236997] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7854040 kB) free!
[1587237117] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34498352 kB) free.
[1587237422] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 76.2% (18677248 kB) free.
[1587237572] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (779936 kB) free!
[1587237692] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (1094584 kB) free!
[1587237717] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11130040 kB) free!
[1587237812] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1087016 kB) free!
[1587237837] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7609352 kB) free!
[1587237957] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33364064 kB) free.
[1587238557] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11213512 kB) free!
[1587238677] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34151880 kB) free.
[1587239012] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.8% (2573916 kB) free!
[1587239277] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13812672 kB) free!
[1587239397] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (10075180 kB) free!
[1587239517] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35595076 kB) free.
[1587239688] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.9% (26286708 kB) free!
[1587239808] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 12.5% (16537980 kB) free!
[1587239928] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (6133380 kB) free!
[1587240117] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12897668 kB) free!
[1587240123] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (233488 kB) free!
[1587240237] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (9390712 kB) free!
[1587240243] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (230620 kB) free!
[1587240357] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34868092 kB) free.
[1587240363] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (214132 kB) free!
[1587240957] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12679304 kB) free!
[1587241078] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.9% (23620096 kB) free!
[1587241198] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32691312 kB) free.
[1587241798] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11373504 kB) free!
[1587241918] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35409784 kB) free.
[1587242012] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3676152 kB) free.
[1587242163] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.3% (9520380 kB) free!
[1587242518] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6360580 kB) free!
[1587242638] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (789696 kB) free!
[1587242758] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32317692 kB) free.
[1587243114] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3479468 kB) free!
[1587243234] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (252844 kB) free!
[1587243354] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (276416 kB) free!
[1587243358] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5731672 kB) free!
[1587243478] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34982548 kB) free.
[1587243954] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 76.7% (37775204 kB) free.
[1587244078] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10117604 kB) free!
[1587244198] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6132016 kB) free!
[1587244257] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.8% (43752540 kB) free.
[1587244318] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34991332 kB) free.
[1587244798] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 38.8% (9517072 kB) free.
[1587244918] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9766740 kB) free!
[1587245038] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3592428 kB) free!
[1587245158] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32272544 kB) free.
[1587245184] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1587245758] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8413612 kB) free!
[1587245764] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.5% (44594072 kB) free.
[1587245879] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34961880 kB) free.
[1587245947] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(347) < 172.105.57.57 [19/Apr/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1587245954] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(470) < 2020/04/19 03:01:42 [error] 22133#0: *84792982 no live upstreams while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1587245972] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2040) < 203.212.70.170 [19/Apr/2020:03:01:41 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1587245998] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (223072 kB) free!
[1587246067] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587246075] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587246092] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587246118] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (202060 kB) free!
[1587246238] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (240048 kB) free!
[1587246478] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10994048 kB) free!
[1587246598] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34898448 kB) free.
[1587247198] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14701424 kB) free!
[1587247318] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (11222656 kB) free!
[1587247438] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34931460 kB) free.
[1587248038] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15501428 kB) free!
[1587248158] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35796320 kB) free.
[1587248638] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.9% (10271020 kB) free.
[1587248758] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16786008 kB) free!
[1587248878] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1587248878] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36688672 kB) free.
[1587249238] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.5% (2830236 kB) free!
[1587249358] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192404 kB) free!
[1587249478] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17503396 kB) free!
[1587249478] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1587249478] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178600 kB) free!
[1587249598] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39085412 kB) free.
[1587250198] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20885352 kB) free!
[1587250318] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39081600 kB) free.
[1587250678] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 38.9% (9537096 kB) free.
[1587250918] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (20439372 kB) free!
[1587251039] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39065524 kB) free.
[1587251638] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (21004860 kB) free!
[1587251758] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39055936 kB) free.
[1587252358] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21659600 kB) free!
[1587252478] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39036684 kB) free.
[1587253078] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169132 kB) free!
[1587253078] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22104656 kB) free!
[1587253198] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (200872 kB) free!
[1587253198] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39009096 kB) free.
[1587253318] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188284 kB) free!
[1587253798] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21619196 kB) free!
[1587253816] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 78.0% (19120156 kB) free.
[1587253912] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 73.4% (17999504 kB) free.
[1587253918] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (38982252 kB) free.
[1587254022] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 57.6% (14207400 kB) free.
[1587254518] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20777588 kB) free!
[1587254638] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38933168 kB) free.
[1587255238] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19385060 kB) free!
[1587255358] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38077608 kB) free.
[1587255958] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18312560 kB) free!
[1587256078] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37356536 kB) free.
[1587257278] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19547108 kB) free!
[1587257398] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38298092 kB) free.
[1587257998] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19627872 kB) free!
[1587258118] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37708720 kB) free.
[1587258528] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 52.7% (69474940 kB) free.
[1587258823] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187260 kB) free!
[1587258943] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (207996 kB) free!
[1587259063] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195056 kB) free!
[1587259277] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 326 processes
[1587259318] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20701192 kB) free!
[1587259397] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 329 processes
[1587259438] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (38019196 kB) free.
[1587259518] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1587260038] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18403132 kB) free!
[1587260158] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36413444 kB) free.
[1587261358] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19729056 kB) free!
[1587261478] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37568152 kB) free.
[1587262078] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18652708 kB) free!
[1587262198] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37083328 kB) free.
[1587262798] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17805620 kB) free!
[1587262918] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1587262919] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (36039244 kB) free.
[1587264118] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18346456 kB) free!
[1587264238] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36827128 kB) free.
[1587264838] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (17003456 kB) free!
[1587264958] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (36014784 kB) free.
[1587265558] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14908068 kB) free!
[1587265678] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34453956 kB) free.
[1587266878] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16204800 kB) free!
[1587266998] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35569504 kB) free.
[1587267598] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13389256 kB) free!
[1587267718] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34110424 kB) free.
[1587267810] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.78, 6.78, 5.32
[1587267930] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.16, 6.50, 5.40
[1587268050] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.31, 5.89, 5.32
[1587268318] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12707024 kB) free!
[1587268438] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32680220 kB) free.
[1587269250] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 5.30, 4.94, 4.90
[1587269638] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14092096 kB) free!
[1587269758] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33441412 kB) free.
[1587269850] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.72, 5.22, 5.11
[1587269970] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.88, 5.17, 5.08
[1587270090] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.83, 5.23, 5.13
[1587270358] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12771284 kB) free!
[1587270478] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32883592 kB) free.
[1587271078] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10700612 kB) free!
[1587271198] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32148080 kB) free.
[1587271584] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1587271704] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1587271717] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1587271798] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10310748 kB) free!
[1587271816] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192516 kB) free!
[1587271824] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1587271837] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1587271918] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31235924 kB) free.
[1587271936] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185252 kB) free!
[1587271957] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1587272057] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194112 kB) free!
[1587272424] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1587272518] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7851608 kB) free!
[1587272558] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1587272638] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35146360 kB) free.
[1587273112] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198400 kB) free!
[1587273232] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (4710032 kB) free!
[1587273238] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (24625616 kB) free!
[1587273352] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.1% (3458976 kB) free!
[1587273358] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34360216 kB) free.
[1587273952] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (1310584 kB) free!
[1587273958] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26329932 kB) free!
[1587274078] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33815292 kB) free.
[1587274290] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1587274290] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 7.20, 5.16, 4.96
[1587274678] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10349036 kB) free!
[1587274799] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33264552 kB) free.
[1587274824] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1587274890] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.26, 6.40, 5.63
[1587274944] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1587275003] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 67.1% (16447448 kB) free.
[1587275010] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.79, 7.15, 6.02
[1587275012] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3240972 kB) free!
[1587275064] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1587275130] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.01, 7.10, 6.15
[1587275132] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (3115120 kB) free!
[1587275253] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.2% (3444828 kB) free.
[1587275492] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(220) < 103.255.146.164 [19/Apr/2020:11:18:21 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587275612] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587275852] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3254888 kB) free!
[1587275972] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (3216224 kB) free!
[1587275998] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11921928 kB) free!
[1587276092] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3167844 kB) free!
[1587276118] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7482156 kB) free!
[1587276238] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31287732 kB) free.
[1587276692] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3293676 kB) free.
[1587276838] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6913676 kB) free!
[1587276860] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.9% (3571244 kB) free!
[1587276958] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31432412 kB) free.
[1587276980] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.7% (3306888 kB) free!
[1587277100] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.6% (3052688 kB) free!
[1587277292] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2992512 kB) free!
[1587277412] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(456) < 103.255.146.164 [19/Apr/2020:11:48:27 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587277412] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (3051372 kB) free!
[1587277532] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587277532] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (3015896 kB) free!
[1587277558] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5975640 kB) free!
[1587277678] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29954488 kB) free.
[1587277700] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1778924 kB) free!
[1587278278] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5284792 kB) free!
[1587278300] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.7% (16106716 kB) free.
[1587278398] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30936480 kB) free.
[1587278664] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1587278664] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1587278998] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6167336 kB) free!
[1587279118] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30091200 kB) free.
[1587279332] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1587279332] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (635548 kB) free!
[1587279332] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(170) < 103.255.146.164 [19/Apr/2020:12:24:12 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587279452] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587279667] SERVICE ALERT: webservercluster1;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: /home 42337 MB (18% inode=99%):
[1587279718] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6560268 kB) free!
[1587279787] SERVICE ALERT: webservercluster1;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: /home 42337 MB (18% inode=99%):
[1587279839] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2110788 kB) free!
[1587279907] SERVICE ALERT: webservercluster1;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: /home 42337 MB (18% inode=99%):
[1587279932] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (2914648 kB) free!
[1587279959] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27440448 kB) free.
[1587280559] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (3118992 kB) free!
[1587280679] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28375496 kB) free.
[1587281060] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 38.2% (6214104 kB) free.
[1587281107] SERVICE ALERT: webservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: /home 88751 MB (39% inode=99%):
[1587281117] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1587281237] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1587281279] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (4680848 kB) free!
[1587281399] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26467716 kB) free.
[1587281852] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(330) < 103.255.146.164 [19/Apr/2020:12:59:20 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587281972] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587281999] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4850292 kB) free!
[1587282119] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (506528 kB) free!
[1587282239] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27464956 kB) free.
[1587282258] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 37.1% (9089032 kB) free.
[1587282260] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (183332 kB) free!
[1587282381] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 69.8% (11347268 kB) free.
[1587282839] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (2778732 kB) free!
[1587282959] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28239192 kB) free.
[1587283173] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(110) < 103.255.146.164 [19/Apr/2020:13:20:25 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587283293] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587283559] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (4695520 kB) free!
[1587283679] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (387680 kB) free!
[1587283799] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27546004 kB) free.
[1587284399] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2924612 kB) free!
[1587284520] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28194532 kB) free.
[1587284603] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.7% (4335116 kB) free!
[1587284723] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.0% (4168260 kB) free!
[1587284843] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.3% (3999708 kB) free!
[1587285093] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(167) < 103.255.146.164 [19/Apr/2020:14:01:12 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587285119] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (398344 kB) free!
[1587285213] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587285239] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393732 kB) free!
[1587285359] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34453244 kB) free.
[1587285932] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (1436816 kB) free!
[1587285959] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (3133064 kB) free!
[1587286044] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2356208 kB) free!
[1587286079] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (398968 kB) free!
[1587286199] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35081172 kB) free.
[1587286414] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(195) < 103.255.146.164 [19/Apr/2020:14:19:55 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587286530] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1587286532] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.1% (2124364 kB) free!
[1587286534] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587286799] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4227276 kB) free!
[1587286919] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (397172 kB) free!
[1587287039] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35575548 kB) free.
[1587287065] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1587287639] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5932400 kB) free!
[1587287734] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(122) < 103.255.146.164 [19/Apr/2020:14:43:49 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587287759] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (730004 kB) free!
[1587287854] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587287879] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33885148 kB) free.
[1587288330] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.81, 3.61, 4.81
[1587288332] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1528084 kB) free!
[1587288479] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5731964 kB) free!
[1587288500] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.7% (3746472 kB) free!
[1587288599] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (18017736 kB) free!
[1587288620] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.6% (16071464 kB) free.
[1587288719] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32510308 kB) free.
[1587288932] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.9% (2257500 kB) free!
[1587289055] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(99) < 185.255.8.58 [19/Apr/2020:15:00:03 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1587289175] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587289320] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6644480 kB) free!
[1587289440] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32571428 kB) free.
[1587290040] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6425112 kB) free!
[1587290160] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4450944 kB) free!
[1587290280] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 20.0% (26414208 kB) free!
[1587290375] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(123) < 103.255.146.164 [19/Apr/2020:15:20:48 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587290495] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587290730] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.02, 5.27, 5.20
[1587290850] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.66, 4.84, 5.06
[1587290880] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (3196688 kB) free!
[1587290970] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.99, 4.52, 4.91
[1587291095] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(134) < 103.255.146.164 [19/Apr/2020:15:41:33 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587291216] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587291570] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.85, 7.64, 6.35
[1587291690] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.47, 7.02, 6.29
[1587291810] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.37, 5.80, 5.93
[1587292080] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16555808 kB) free!
[1587292410] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.49, 3.84, 4.85
[1587293132] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1180700 kB) free!
[1587293280] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1587293280] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389876 kB) free!
[1587293615] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(101) < 103.255.146.164 [19/Apr/2020:16:18:13 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587293735] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587294332] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.0% (1794636 kB) free!
[1587294480] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15838880 kB) free!
[1587295037] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1587295157] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1587295257] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9628912 kB) free!
[1587295377] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.3% (9501152 kB) free!
[1587295497] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.0% (9387248 kB) free!
[1587295532] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1530800 kB) free!
[1587295535] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(652) < 103.255.146.164 [19/Apr/2020:16:46:28 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587295655] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587295680] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409168 kB) free!
[1587296732] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.7% (2220980 kB) free!
[1587296754] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9520416 kB) free!
[1587296874] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (9447580 kB) free!
[1587296881] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (15920532 kB) free!
[1587296995] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (9373716 kB) free!
[1587297210] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.60, 5.42, 5.27
[1587297258] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185836 kB) free!
[1587297264] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 356 processes
[1587297330] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.09, 4.89, 5.09
[1587297332] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1522052 kB) free!
[1587297379] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182616 kB) free!
[1587297384] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 356 processes
[1587297450] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1587297450] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.59, 4.93, 5.09
[1587297499] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173404 kB) free!
[1587297504] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 357 processes
[1587297932] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.4% (2015576 kB) free!
[1587298056] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(460) < 103.255.146.164 [19/Apr/2020:17:31:20 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587298176] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587298532] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (878292 kB) free!
[1587298681] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (6759844 kB) free!
[1587298776] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(169) < 103.255.146.164 [19/Apr/2020:17:42:28 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587298896] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587299250] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.92, 3.29, 4.05
[1587299281] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15806736 kB) free!
[1587299496] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(219) < 103.255.146.164 [19/Apr/2020:18:01:08 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587299617] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587299732] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (3144000 kB) free!
[1587300481] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30875284 kB) free.
[1587301081] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5492360 kB) free!
[1587301104] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1587301201] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (918576 kB) free!
[1587301321] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29174156 kB) free.
[1587301497] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4786196 kB) free!
[1587301921] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (3048996 kB) free!
[1587302041] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31324360 kB) free.
[1587302250] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.37, 5.74, 5.29
[1587302370] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.02, 5.18, 5.14
[1587302490] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.51, 4.24, 4.81
[1587302641] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6232116 kB) free!
[1587302761] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1572024 kB) free!
[1587302881] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31887628 kB) free.
[1587303332] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3688204 kB) free.
[1587303482] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5153804 kB) free!
[1587303499] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 39.6% (9715048 kB) free.
[1587303602] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (951972 kB) free!
[1587303722] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29565124 kB) free.
[1587303932] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (3191176 kB) free!
[1587304052] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (2862168 kB) free!
[1587304172] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.4% (2510892 kB) free!
[1587304322] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4168172 kB) free!
[1587304442] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31326200 kB) free.
[1587304772] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3516604 kB) free.
[1587305042] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6360632 kB) free!
[1587305162] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2054848 kB) free!
[1587305282] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30140820 kB) free.
[1587305781] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (155464 kB) free!
[1587305882] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (5037164 kB) free!
[1587305899] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.9% (3648956 kB) free!
[1587305901] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (168516 kB) free!
[1587306002] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30914432 kB) free.
[1587306020] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.8% (3633556 kB) free!
[1587306021] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (175836 kB) free!
[1587306140] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 14.8% (3633692 kB) free!
[1587306595] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4872344 kB) free!
[1587306602] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7139508 kB) free!
[1587306722] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2659608 kB) free!
[1587306842] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29870936 kB) free.
[1587307172] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.6% (2043736 kB) free!
[1587307292] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3277316 kB) free.
[1587307442] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4072008 kB) free!
[1587307562] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30555628 kB) free.
[1587307892] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.5% (2516416 kB) free!
[1587308012] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.3% (1996208 kB) free!
[1587308132] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (3206248 kB) free!
[1587308162] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6626312 kB) free!
[1587308283] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2101564 kB) free!
[1587308403] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28509744 kB) free.
[1587309003] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (1520692 kB) free!
[1587309123] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31171588 kB) free.
[1587309690] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1587309723] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (7013312 kB) free!
[1587309843] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (3025288 kB) free!
[1587309963] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28256152 kB) free.
[1587310563] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (4004812 kB) free!
[1587310683] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30072308 kB) free.
[1587311283] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5751408 kB) free!
[1587311403] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30705468 kB) free.
[1587312003] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6618696 kB) free!
[1587312123] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30571072 kB) free.
[1587312723] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6901216 kB) free!
[1587312843] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30407564 kB) free.
[1587312933] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.2% (4091292 kB) free.
[1587313443] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7190240 kB) free!
[1587313526] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1696 processes
[1587313564] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30327720 kB) free.
[1587314127] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1841 processes
[1587314163] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7670756 kB) free!
[1587314283] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3736928 kB) free!
[1587314403] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1587314403] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27217744 kB) free.
[1587315003] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4421112 kB) free!
[1587315123] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27765384 kB) free.
[1587316323] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7523068 kB) free!
[1587316443] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29879244 kB) free.
[1587316533] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (3095504 kB) free!
[1587316653] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.9% (3883396 kB) free.
[1587317043] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7634012 kB) free!
[1587317163] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29749436 kB) free.
[1587317763] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7609784 kB) free!
[1587317883] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29710840 kB) free.
[1587318483] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7282776 kB) free!
[1587318603] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29216024 kB) free.
[1587319203] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7549288 kB) free!
[1587319324] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29571228 kB) free.
[1587319923] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7223304 kB) free!
[1587320043] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27752516 kB) free.
[1587320254] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)