[1554489010] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1554489045] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 339 processes
[1554489130] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554489203] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (1948932 kB) free!
[1554489251] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554489328] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.8% (9778272 kB) free!
[1554489358] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 321 processes
[1554489803] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.8% (7294756 kB) free!
[1554490403] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9885312 kB) free.
[1554490451] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1554490528] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4685208 kB) free!
[1554490927] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 488 processes
[1554490987] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 475 processes
[1554491003] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4322680 kB) free!
[1554491051] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554491124] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (3164280 kB) free!
[1554491171] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554491244] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (2077636 kB) free!
[1554491291] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554491322] SERVICE ALERT: pservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 109421 MB (49% inode=99%):
[1554491667] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 39.5% (9747976 kB) free.
[1554491728] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.4% (10054180 kB) free.
[1554491844] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.7% (7222720 kB) free!
[1554492328] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.9% (8306180 kB) free!
[1554492444] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9900364 kB) free.
[1554492448] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.5% (5644024 kB) free!
[1554492491] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1554492568] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2966332 kB) free!
[1554493044] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4342420 kB) free!
[1554493091] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554493164] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (3167380 kB) free!
[1554493168] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.5% (10082484 kB) free.
[1554493212] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554493284] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (1881216 kB) free!
[1554493332] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554493555] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1554493768] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9123452 kB) free!
[1554493884] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.8% (7775868 kB) free!
[1554493889] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.3% (6575008 kB) free!
[1554494009] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3917680 kB) free!
[1554494484] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (431356 kB) free!
[1554494532] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1554495085] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.9% (5887444 kB) free!
[1554495132] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554495209] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.5% (10108116 kB) free.
[1554495252] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554495267] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165812 kB) free!
[1554495372] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554495387] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179636 kB) free!
[1554495477] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45372 MB (20% inode=99%):
[1554495508] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197248 kB) free!
[1554495597] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57665 MB (25% inode=99%):
[1554495685] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9912552 kB) free.
[1554496285] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3571188 kB) free!
[1554496405] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (2295284 kB) free!
[1554496525] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (363532 kB) free!
[1554496572] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1554497009] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8826648 kB) free!
[1554497125] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.1% (5470360 kB) free!
[1554497129] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.6% (6221864 kB) free!
[1554497172] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554497249] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (3517148 kB) free!
[1554497292] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554497412] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554497725] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9914912 kB) free.
[1554498325] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (3257620 kB) free!
[1554498445] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (2011608 kB) free!
[1554498449] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.1% (9896148 kB) free.
[1554498566] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9931488 kB) free.
[1554498612] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1554499049] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.3% (7537684 kB) free!
[1554499165] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.2% (5002508 kB) free!
[1554499170] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (4860960 kB) free!
[1554499213] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554499285] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (3733756 kB) free!
[1554499290] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (2041972 kB) free!
[1554499333] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554499405] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (2453600 kB) free!
[1554499453] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554499890] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.1% (9902168 kB) free.
[1554500005] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.6% (7697268 kB) free!
[1554500490] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.0% (7405384 kB) free!
[1554500605] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9928476 kB) free.
[1554500610] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (4797108 kB) free!
[1554500653] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1554500730] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (2024184 kB) free!
[1554500866] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.9% (2671976 kB) free!
[1554500917] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.9% (2917948 kB) free!
[1554501205] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (4425928 kB) free!
[1554501253] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554501325] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (3124560 kB) free!
[1554501330] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 20.0% (9851760 kB) free!
[1554501373] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554501445] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (1844924 kB) free!
[1554501493] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554502045] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.2% (6977596 kB) free!
[1554502693] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1554503130] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (3078916 kB) free!
[1554503245] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4162452 kB) free!
[1554503293] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554503413] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554503533] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554503730] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.6% (10169588 kB) free.
[1554503845] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 20.0% (9867288 kB) free!
[1554504330] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.4% (8566252 kB) free!
[1554504445] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (2062664 kB) free!
[1554504450] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.3% (6085660 kB) free!
[1554504571] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (3350692 kB) free!
[1554505066] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2440504 kB) free!
[1554505645] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.2% (6017348 kB) free!
[1554505771] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.8% (9259232 kB) free!
[1554505953] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1593460 kB) free!
[1554506677] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 18.90, 7.92, 3.26
[1554506797] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 21.01, 12.37, 5.46
[1554506846] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (234352 kB) free!
[1554506917] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 21.67, 15.41, 7.40
[1554507508] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.2% (3263844 kB) free!
[1554507571] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (3519328 kB) free!
[1554507598] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45832 MB (20% inode=99%):
[1554507718] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56522 MB (25% inode=99%):
[1554508117] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.54, 0.89, 3.06
[1554508318] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45252 MB (20% inode=99%):
[1554508438] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57542 MB (25% inode=99%):
[1554508646] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.1% (6462096 kB) free!
[1554508771] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.7% (9701564 kB) free!
[1554509846] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (730876 kB) free!
[1554510572] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4411828 kB) free!
[1554510733] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1554511046] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.3% (5088192 kB) free!
[1554511333] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554511438] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47004 MB (20% inode=99%):
[1554511453] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1554511558] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57404 MB (25% inode=99%):
[1554511574] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1554511772] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.1% (9395944 kB) free!
[1554512173] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554512246] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (2900800 kB) free!
[1554512293] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554512413] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554512846] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.1% (8450244 kB) free!
[1554512972] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4825228 kB) free!
[1554513446] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (222488 kB) free!
[1554513613] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1554513645] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 375 processes
[1554514046] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.2% (5046104 kB) free!
[1554514067] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 12.1% (2959800 kB) free!
[1554514172] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.3% (9497704 kB) free!
[1554514213] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1554514245] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 420 processes
[1554514333] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554514454] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554515246] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2977036 kB) free!
[1554515654] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1554515846] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.8% (8783244 kB) free!
[1554515972] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (2786064 kB) free!
[1554516255] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1554516375] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554516446] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (398116 kB) free!
[1554516495] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1554516508] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2479624 kB) free!
[1554516572] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.3% (9501544 kB) free!
[1554517047] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.5% (5675448 kB) free!
[1554517095] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554517215] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554517336] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554518247] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3664192 kB) free!
[1554518372] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3249180 kB) free!
[1554518847] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.5% (8123280 kB) free!
[1554519572] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.3% (8522592 kB) free!
[1554520047] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (2656872 kB) free!
[1554520555] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554520557] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4912456 kB) free!
[1554520617] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4910712 kB) free!
[1554520675] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554520678] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 20.0% (4913352 kB) free!
[1554520738] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 20.0% (4912268 kB) free!
[1554520772] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (3504236 kB) free!
[1554520795] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554521248] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.8% (9242508 kB) free!
[1554521268] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185768 kB) free!
[1554521317] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (463584 kB) free!
[1554521848] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (459500 kB) free!
[1554521972] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.5% (8621120 kB) free!
[1554522736] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1554523172] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (3454104 kB) free!
[1554523336] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1554523457] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1554523577] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 319 processes
[1554523648] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.9% (7846600 kB) free!
[1554524372] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.9% (8807836 kB) free!
[1554524758] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46990 MB (20% inode=99%):
[1554524849] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (1957428 kB) free!
[1554524878] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47231 MB (21% inode=99%):
[1554524995] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554525572] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (3462572 kB) free!
[1554525645] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 338 processes
[1554526049] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.3% (7053580 kB) free!
[1554526649] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10133880 kB) free.
[1554526772] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.1% (9927712 kB) free.
[1554527249] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (1517748 kB) free!
[1554527370] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (252428 kB) free!
[1554527372] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.1% (8904516 kB) free!
[1554527490] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.1% (9416608 kB) free!
[1554527492] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.0% (5905804 kB) free!
[1554527612] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2954608 kB) free!
[1554527868] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 15.5% (3794540 kB) free!
[1554528090] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (212268 kB) free!
[1554528812] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.6% (10175664 kB) free.
[1554528977] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1554529412] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.7% (8703424 kB) free!
[1554529532] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.8% (5792392 kB) free!
[1554529578] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1554529652] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (2807408 kB) free!
[1554529678] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46170 MB (20% inode=99%):
[1554529698] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554529798] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55896 MB (24% inode=99%):
[1554529818] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554529890] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.1% (7438256 kB) free!
[1554530268] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2333508 kB) free!
[1554530398] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44110 MB (19% inode=99%):
[1554530518] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56401 MB (25% inode=99%):
[1554531090] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (1196920 kB) free!
[1554531452] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.5% (10617584 kB) free.
[1554531916] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.7% (2884380 kB) free!
[1554532036] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.0% (2769248 kB) free!
[1554532052] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.7% (8250088 kB) free!
[1554532156] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.4% (2669480 kB) free!
[1554532172] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.8% (5315876 kB) free!
[1554532290] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.7% (6236192 kB) free!
[1554532292] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (2352716 kB) free!
[1554532556] SERVICE ALERT: sig4-vm1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 280 processes
[1554532676] SERVICE ALERT: sig4-vm1;CPU_Procs;OK;HARD;1;CPU OK: 275 processes
[1554532688] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (538736 kB) free!
[1554532809] SERVICE ALERT: sig3-vm1;Memory;OK;HARD;1;OK - 96.2% (7073060 kB) free.
[1554533356] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (1540180 kB) free!
[1554533490] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (862708 kB) free!
[1554534556] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 84.2% (13688288 kB) free.
[1554534690] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6146488 kB) free!
[1554535218] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1554535818] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 319 processes
[1554535890] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (548792 kB) free!
[1554535918] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45841 MB (20% inode=99%):
[1554535939] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1554536038] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56077 MB (24% inode=99%):
[1554536059] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1554537092] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.9% (5843820 kB) free!
[1554537691] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.7% (7230928 kB) free!
[1554537692] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (1903540 kB) free!
[1554538195] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554538316] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554538436] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554538891] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (1198504 kB) free!
[1554539059] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1554539638] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47039 MB (20% inode=99%):
[1554539659] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554539758] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47039 MB (20% inode=99%):
[1554539780] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1554539878] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55229 MB (24% inode=99%):
[1554539900] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1554540091] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.8% (6778808 kB) free!
[1554540478] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47037 MB (20% inode=99%):
[1554540500] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554540598] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44989 MB (20% inode=99%):
[1554540621] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1554540688] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171068 kB) free!
[1554540692] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.8% (5834288 kB) free!
[1554540718] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55193 MB (24% inode=99%):
[1554540741] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554540808] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196668 kB) free!
[1554540928] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189116 kB) free!
[1554541291] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (1436824 kB) free!
[1554541292] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (2530356 kB) free!
[1554541318] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47033 MB (20% inode=99%):
[1554541438] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42863 MB (19% inode=99%):
[1554541558] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54648 MB (24% inode=99%):
[1554542158] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46968 MB (20% inode=99%):
[1554542278] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55158 MB (24% inode=99%):
[1554542308] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 35.7% (8815228 kB) free.
[1554542373] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46575 MB (20% inode=99%):
[1554542492] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.3% (6577300 kB) free!
[1554542493] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46573 MB (20% inode=99%):
[1554542613] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46549 MB (20% inode=99%):
[1554543478] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46867 MB (20% inode=99%):
[1554543598] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46867 MB (20% inode=99%):
[1554543687] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 16.27, 9.88, 4.43
[1554543692] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (1311064 kB) free!
[1554543718] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46872 MB (20% inode=99%):
[1554543807] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 16.31, 12.04, 5.89
[1554543927] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.97, 9.11, 5.60
[1554544109] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (361292 kB) free!
[1554544229] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190080 kB) free!
[1554544319] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47648 MB (21% inode=99%):
[1554544349] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181980 kB) free!
[1554544527] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.49, 1.63, 3.16
[1554544892] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6062332 kB) free!
[1554546092] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (486984 kB) free!
[1554546692] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.1% (8439984 kB) free!
[1554546719] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46529 MB (20% inode=99%):
[1554546839] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46561 MB (20% inode=99%):
[1554546959] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46566 MB (20% inode=99%):
[1554547559] SERVICE FLAPPING ALERT: db2;HDD_Home status;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1554547559] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48451 MB (21% inode=99%):
[1554547892] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.0% (7885248 kB) free!
[1554547892] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (1084496 kB) free!
[1554548492] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240136 kB) free!
[1554548541] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1554549141] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554549262] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554549382] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554550292] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.8% (6818344 kB) free!
[1554551492] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.5% (9124140 kB) free!
[1554551492] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (1506060 kB) free!
[1554552692] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (2086352 kB) free!
[1554552692] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.4% (7105544 kB) free!
[1554553892] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (2693976 kB) free!
[1554554036] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554555092] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.9% (7830356 kB) free!
[1554555692] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (247100 kB) free!
[1554556756] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (1213360 kB) free!
[1554556877] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (1201576 kB) free!
[1554556997] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1190436 kB) free!
[1554557036] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554557156] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554557276] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554557492] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.8% (6820476 kB) free!
[1554557550] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.5% (2838844 kB) free!
[1554558150] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198312 kB) free!
[1554558692] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (1603476 kB) free!
[1554558959] SERVICE FLAPPING ALERT: db2;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1554559292] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.7% (10177372 kB) free.
[1554559559] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45069 MB (20% inode=99%):
[1554559679] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56247 MB (25% inode=99%):
[1554559892] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6037220 kB) free!
[1554559892] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.5% (7165516 kB) free!
[1554560012] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.4% (6616188 kB) free!
[1554560132] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.2% (6031256 kB) free!
[1554560492] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10452820 kB) free.
[1554560732] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (482340 kB) free!
[1554560782] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1554561092] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (2632428 kB) free!
[1554561212] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (900872 kB) free!
[1554561332] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.8% (9258540 kB) free!
[1554561709] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 17.02, 13.70, 7.54
[1554561829] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.36, 12.77, 8.00
[1554561832] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1216 processes
[1554561932] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (1016364 kB) free!
[1554561949] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.97, 9.22, 7.28
[1554561952] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 714 processes
[1554562072] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 636 processes
[1554562532] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.5% (5185812 kB) free!
[1554562533] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.8% (7311040 kB) free!
[1554562549] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.47, 2.66, 4.62
[1554562673] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 381 processes
[1554563133] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4551332 kB) free!
[1554563152] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1554563272] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1554563392] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1554563782] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1554563917] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.13, 10.56, 5.18
[1554564037] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.50, 7.30, 4.64
[1554564328] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.9% (2675740 kB) free!
[1554564333] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (2097808 kB) free!
[1554564383] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1554564479] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46643 MB (20% inode=99%):
[1554564503] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1554564599] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55601 MB (24% inode=99%):
[1554565103] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554565223] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1554565343] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554565533] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.6% (5223996 kB) free!
[1554566733] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.1% (9925160 kB) free.
[1554566733] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (225084 kB) free!
[1554567333] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.0% (6887096 kB) free!
[1554567453] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.7% (6254556 kB) free!
[1554567573] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.1% (5493544 kB) free!
[1554567599] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46826 MB (20% inode=99%):
[1554567719] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56362 MB (25% inode=99%):
[1554567928] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2438332 kB) free!
[1554567933] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.4% (5625168 kB) free!
[1554568076] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1554569133] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (1140292 kB) free!
[1554569373] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4418468 kB) free!
[1554570143] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1554570334] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.8% (6817716 kB) free!
[1554570743] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1554570864] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1554570928] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (2553816 kB) free!
[1554570984] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1554571534] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (2815464 kB) free!
[1554571584] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1554571704] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554571824] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1554572134] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.0% (8864120 kB) free!
[1554572425] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554572545] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554572665] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1554572665] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554572734] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235212 kB) free!
[1554573335] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.4% (5119624 kB) free!
[1554573865] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1554574465] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1554574535] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (2499296 kB) free!
[1554574585] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554574706] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1554575135] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.5% (8137748 kB) free!