[1559154687] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1559154807] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559154985] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (2737132 kB) free!
[1559155060] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1559155077] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559155105] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (3710636 kB) free.
[1559155677] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559155705] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (484312 kB) free!
[1559155797] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559155825] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3149896 kB) free!
[1559155946] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (3695872 kB) free.
[1559156397] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559156517] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1559156545] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (2942152 kB) free!
[1559156665] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (3690176 kB) free.
[1559157265] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (604812 kB) free!
[1559157345] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.4% (3479568 kB) free.
[1559157385] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (3689464 kB) free.
[1559157511] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9300672 kB) free!
[1559157631] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.0% (14792188 kB) free.
[1559157717] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559157807] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1559157837] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1559157985] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (664976 kB) free!
[1559158105] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (3682128 kB) free.
[1559158437] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559158557] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559158705] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (2291968 kB) free!
[1559158825] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (3680724 kB) free.
[1559159007] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559159127] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1559159145] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.9% (2897028 kB) free!
[1559159265] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (185124 kB) free!
[1559159386] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.2% (2948572 kB) free!
[1559159727] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1559159847] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559159967] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559160357] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559160477] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559160568] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559160631] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8562360 kB) free!
[1559160688] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1559160751] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.0% (14785940 kB) free.
[1559161288] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1559161352] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9489240 kB) free!
[1559161408] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559161472] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.0% (14785384 kB) free.
[1559161528] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1559161786] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.8% (3529840 kB) free.
[1559161825] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (1986588 kB) free!
[1559161945] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (3662012 kB) free.
[1559162128] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1559162248] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1559162386] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (167340 kB) free!
[1559162506] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (1013080 kB) free!
[1559162545] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (2453704 kB) free!
[1559162626] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (1001672 kB) free!
[1559162665] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559162665] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (3664716 kB) free.
[1559162849] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559162861] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 306.23, 163.94, 66.55
[1559162969] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559162981] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 313.31, 213.72, 96.60
[1559163056] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1559163089] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1559163101] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 136.51, 198.62, 106.54
[1559163689] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1559163809] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559164409] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559164465] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (712508 kB) free!
[1559164472] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9688884 kB) free!
[1559164529] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1559164585] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (3656800 kB) free.
[1559164592] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.0% (14780976 kB) free.
[1559164649] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1559164901] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.30, 0.76, 15.54
[1559165185] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (544300 kB) free!
[1559165250] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559165305] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (3651128 kB) free.
[1559165370] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559165626] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.3% (1671408 kB) free!
[1559165877] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559165906] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (2004876 kB) free!
[1559165970] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559165997] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559165997] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1559166026] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (3649008 kB) free.
[1559166090] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1559166101] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.29, 0.19, 4.41
[1559166210] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559166226] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1276232 kB) free!
[1559166810] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559166930] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1559167256] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 12.36, 112.56, 81.70
[1559167376] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.93, 75.41, 71.83
[1559167497] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.69, 50.58, 63.16
[1559167531] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1559167593] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8739836 kB) free!
[1559167651] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1559167713] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 30.0% (14780012 kB) free.
[1559168026] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.4% (1855780 kB) free!
[1559168251] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559168371] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559168491] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1559168697] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.58, 1.31, 17.59
[1559169091] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1559169211] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1559169211] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559169226] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (193160 kB) free!
[1559169811] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559169931] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1559170052] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559170497] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.79, 0.58, 2.98
[1559170651] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559170771] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1559171026] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.4% (2341012 kB) free!
[1559171371] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559171491] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1559171626] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.1% (3754732 kB) free.
[1559172091] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1559172211] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559172331] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1559172597] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559172626] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (491432 kB) free!
[1559172717] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1559172746] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (3602856 kB) free.
[1559172932] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1559173052] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559173652] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1559173713] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8743992 kB) free!
[1559173772] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1559173834] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.9% (14754852 kB) free.
[1559173917] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559173946] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (2571208 kB) free!
[1559174026] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.1% (2778568 kB) free!
[1559174037] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1559174066] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (3640080 kB) free.
[1559174146] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.2% (2784876 kB) free!
[1559174266] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.1% (2771068 kB) free!
[1559174637] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559174666] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (619408 kB) free!
[1559174757] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559174786] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.3% (2642908 kB) free!
[1559174877] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559174906] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (3671104 kB) free.
[1559174972] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559175092] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1559175212] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559175477] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559175506] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (1198572 kB) free!
[1559175597] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559175626] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.2% (2793512 kB) free!
[1559175746] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (3663344 kB) free.
[1559176197] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559176317] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559176347] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (1146692 kB) free!
[1559176437] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559176467] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.2% (2786328 kB) free!
[1559176587] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (3654728 kB) free.
[1559176666] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174200 kB) free!
[1559177037] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559177157] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559177187] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (1162244 kB) free!
[1559177266] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 30.8% (4998180 kB) free.
[1559177277] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559177307] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.5% (2839412 kB) free!
[1559177427] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (3644972 kB) free.
[1559177878] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559177998] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559178027] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (1322936 kB) free!
[1559178118] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1559178147] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.0% (2913184 kB) free!
[1559178212] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559178267] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (3615120 kB) free.
[1559178633] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9345828 kB) free!
[1559178718] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559178753] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.4% (14508164 kB) free.
[1559178812] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559178838] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559178867] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (1255704 kB) free!
[1559178932] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1559178958] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559178987] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.6% (3021416 kB) free!
[1559179053] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559179107] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (3662388 kB) free.
[1559179558] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559179678] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559179707] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (1838136 kB) free!
[1559179798] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559179827] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (3435852 kB) free.
[1559179946] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 290.84, 158.56, 65.79
[1559180066] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 297.05, 204.68, 93.94
[1559180187] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 297.53, 235.43, 118.64
[1559180398] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559180427] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (491432 kB) free!
[1559180518] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559180547] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.0% (2272628 kB) free!
[1559180638] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1559180667] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (3656352 kB) free.
[1559180853] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1559180853] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559181238] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559181267] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (1457440 kB) free!
[1559181358] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1559181388] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (3287760 kB) free.
[1559181453] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559181573] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559181693] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1559181958] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559182078] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559182199] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559182338] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 3 warn out of 575 processes
[1559182357] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8697248 kB) free!
[1559182398] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 569 processes
[1559182450] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 17.5% (4279764 kB) free!
[1559182477] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.1% (14350876 kB) free.
[1559182510] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 26.2% (6418624 kB) free.
[1559182587] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (2172068 kB) free!
[1559182587] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.01, 1.83, 14.18
[1559182707] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (3657852 kB) free.
[1559182798] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559182918] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559183038] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1559183298] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 575 processes
[1559183307] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (1223892 kB) free!
[1559183358] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 570 processes
[1559183427] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.9% (3070192 kB) free!
[1559183493] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1559183547] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (3696900 kB) free.
[1559183638] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559183658] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 567 processes
[1559183718] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 564 processes
[1559183758] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559183878] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1559184093] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1559184147] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (2123604 kB) free!
[1559184213] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1559184267] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (3643084 kB) free.
[1559184333] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1559184387] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.46, 1.59, 3.45
[1559184478] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559184598] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559184867] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (387736 kB) free!
[1559184987] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.4% (2005552 kB) free!
[1559185107] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (3545152 kB) free.
[1559185198] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559185318] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559185438] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1559185477] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9850384 kB) free!
[1559185597] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7641044 kB) free!
[1559185707] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (1649088 kB) free!
[1559185717] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.8% (13691788 kB) free.
[1559185827] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 20.0% (3250304 kB) free!
[1559185947] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (3662636 kB) free.
[1559186038] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559186097] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 67.51, 15.20, 5.26
[1559186158] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1559186217] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 271.48, 110.54, 41.41
[1559186278] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559186337] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 299.09, 174.34, 73.16
[1559186547] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (1923988 kB) free!
[1559186667] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (3555984 kB) free.
[1559186878] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559186999] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1559187267] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (1675940 kB) free!
[1559187387] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (3312764 kB) free.
[1559187598] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559187718] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559187838] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559187988] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (298348 kB) free!
[1559188108] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (2163228 kB) free!
[1559188117] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9479680 kB) free!
[1559188137] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.56, 1.48, 19.18
[1559188228] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (3524292 kB) free.
[1559188237] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7263968 kB) free!
[1559188357] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.2% (13920004 kB) free.
[1559188438] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559188558] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1559188666] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559188828] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (1006084 kB) free!
[1559188948] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.3% (2798892 kB) free!
[1559189068] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (3552100 kB) free.
[1559189158] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559189278] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1559189398] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1559189668] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (1763616 kB) free!
[1559189788] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (3460132 kB) free.
[1559189937] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.24, 0.32, 3.06
[1559189998] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1559190118] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1559190388] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (453500 kB) free!
[1559190508] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.0% (2439972 kB) free!
[1559190628] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (3466344 kB) free.
[1559190718] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559190838] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559190958] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1559191066] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.1% (2930448 kB) free!
[1559191186] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.4% (3796592 kB) free.
[1559191228] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (1311860 kB) free!
[1559191348] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.8% (3055196 kB) free!
[1559191357] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9072256 kB) free!
[1559191468] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (3398476 kB) free.
[1559191477] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.9% (14252076 kB) free.
[1559191558] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559191679] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559192068] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (2075260 kB) free!
[1559192154] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 42.3% (6867244 kB) free.
[1559192278] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559192398] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559192518] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1559192760] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (2103500 kB) free!
[1559192880] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (3476000 kB) free.
[1559193118] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559193238] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1559193480] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (1362760 kB) free!
[1559193586] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.9% (2097500 kB) free!
[1559193600] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (3508028 kB) free.
[1559193706] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.5% (3978792 kB) free.
[1559193838] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559193958] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559194078] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1559194200] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (737120 kB) free!
[1559194306] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.0% (2113352 kB) free!
[1559194320] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.5% (2678388 kB) free!
[1559194426] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.5% (3487620 kB) free.
[1559194440] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (3389448 kB) free.
[1559194678] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559194798] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1559195026] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.4% (2982612 kB) free!
[1559195041] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (2401952 kB) free!
[1559195077] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8890068 kB) free!
[1559195146] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.4% (2505908 kB) free!
[1559195161] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (3274116 kB) free.
[1559195197] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.5% (14057048 kB) free.
[1559195266] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.9% (2098492 kB) free!
[1559195393] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 346 processes
[1559195398] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559195513] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 346 processes
[1559195518] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559195633] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1559195733] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559195761] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (1999212 kB) free!
[1559195801] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (162148 kB) free!
[1559195819] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (227008 kB) free!
[1559195881] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (3255872 kB) free.
[1559195921] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (163264 kB) free!
[1559195937] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.23, 13.54, 7.61
[1559195939] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199512 kB) free!
[1559196041] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (162860 kB) free!
[1559196057] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.99, 9.20, 6.74
[1559196059] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195348 kB) free!
[1559196118] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559196177] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.46, 6.26, 5.96
[1559196238] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559196333] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1559196466] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.6% (3667012 kB) free.
[1559196481] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (1671980 kB) free!
[1559196601] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.8% (3204836 kB) free!
[1559196721] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165740 kB) free!
[1559196777] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (25.6% change >= 20.0% threshold)
[1559196777] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.60, 1.21, 3.31
[1559196786] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 93.2% (15141080 kB) free.
[1559196818] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 39.3% (6368404 kB) free.
[1559196839] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559196859] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.2% (22609748 kB) free.
[1559196934] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1559196959] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1559197054] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559197174] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559197422] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (166236 kB) free!
[1559197542] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.7% (2055184 kB) free!
[1559197559] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559197662] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.7% (3198004 kB) free!
[1559197666] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.3% (2814596 kB) free!
[1559197679] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1559197786] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 16.8% (2726792 kB) free!
[1559197906] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.1% (2615456 kB) free!
[1559198262] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1559198262] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (171024 kB) free!
[1559198279] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559198399] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 370 processes
[1559198519] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 385 processes
[1559198797] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9201200 kB) free!
[1559198917] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7142332 kB) free!
[1559198974] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1559199038] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.5% (13566304 kB) free.
[1559199424] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.8% (3062964 kB) free!
[1559199462] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.0% (2593672 kB) free!
[1559199544] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.1% (2945276 kB) free!
[1559199574] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559199619] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 565 processes
[1559199665] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.4% (2821260 kB) free!
[1559199679] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 572 processes
[1559199694] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1559199706] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1559199706] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (176296 kB) free!
[1559199719] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 368 processes
[1559199739] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 571 processes
[1559199815] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1559199815] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559199839] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1559200306] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.4% (2504424 kB) free!
[1559200865] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (1477480 kB) free!
[1559201039] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559201159] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 368 processes
[1559201262] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1052576 kB) free!
[1559201279] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 363 processes
[1559201438] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9521288 kB) free!
[1559201558] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7672900 kB) free!
[1559201678] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.5% (13573632 kB) free.
[1559202106] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1018576 kB) free!
[1559202215] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1559203062] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (1717020 kB) free!
[1559203415] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559203535] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1559203656] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559203679] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 392 processes
[1559203906] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.5% (2350680 kB) free!
[1559204262] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155228 kB) free!
[1559204279] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 400 processes
[1559204399] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 390 processes
[1559204519] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 391 processes
[1559204678] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9833972 kB) free!
[1559204799] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8188616 kB) free!
[1559204919] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12241680 kB) free.
[1559205456] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1559205462] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (1982348 kB) free!
[1559205706] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161204 kB) free!
[1559205719] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1559205719] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1559206062] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (172728 kB) free!
[1559206306] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.0% (1780828 kB) free!
[1559206319] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 395 processes
[1559206439] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 371 processes
[1559206559] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 387 processes
[1559206656] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1559206662] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (1933224 kB) free!
[1559206675] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 16.1% (1180624 kB) free!
[1559206776] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1559206795] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 14.7% (1078704 kB) free!
[1559206896] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1559206915] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 13.2% (970296 kB) free!
[1559207262] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (167044 kB) free!
[1559207506] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174792 kB) free!
[1559207515] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (361008 kB) free!
[1559207759] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 386 processes
[1559207919] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9664844 kB) free!
[1559208039] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (7988616 kB) free!
[1559208106] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.7% (2066896 kB) free!
[1559208159] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.4% (5626784 kB) free!
[1559208177] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559208359] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 406 processes
[1559208479] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 387 processes
[1559208599] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 399 processes
[1559208706] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1125448 kB) free!
[1559208715] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.9% (6830072 kB) free.
[1559208759] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.7% (13158060 kB) free.
[1559209296] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559209306] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 20.0% (3250924 kB) free!
[1559209662] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (1847220 kB) free!
[1559209896] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559210016] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1559210136] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559210262] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (173348 kB) free!
[1559210399] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 416 processes
[1559210506] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.4% (3312400 kB) free.
[1559210862] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (1870840 kB) free!
[1559210999] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 397 processes
[1559211106] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.1% (2940484 kB) free!
[1559211119] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 398 processes
[1559211226] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.2% (3275056 kB) free.
[1559211462] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (1012228 kB) free!
[1559211759] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9341476 kB) free!
[1559211826] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.5% (2345060 kB) free!
[1559211879] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7541336 kB) free!
[1559211946] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.8% (3374760 kB) free.
[1559211999] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.5% (5153516 kB) free!
[1559212319] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 423 processes
[1559212384] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.5% (3096296 kB) free!
[1559212439] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 420 processes
[1559212504] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.9% (2930428 kB) free!
[1559212546] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.7% (2701752 kB) free!
[1559212559] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 425 processes
[1559212599] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1559212599] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12853432 kB) free.
[1559212624] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.3% (2789956 kB) free!
[1559212660] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4883716 kB) free!
[1559212667] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.2% (3282268 kB) free.
[1559212780] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.1% (4200324 kB) free!
[1559212900] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.2% (3479552 kB) free!
[1559212901] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 305.78, 214.90, 98.90
[1559213021] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 122.85, 193.84, 106.51
[1559213136] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1559213141] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15.78, 127.75, 93.16
[1559213224] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2339484 kB) free!
[1559213266] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (163216 kB) free!
[1559213386] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (958680 kB) free!
[1559213500] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 23.1% (5667384 kB) free.
[1559213506] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.8% (2891164 kB) free!
[1559213633] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559213736] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559213753] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 348 processes
[1559213856] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559213874] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1559213976] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559214191] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (163052 kB) free!
[1559214311] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (161848 kB) free!
[1559214431] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (162028 kB) free!
[1559214462] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (1660556 kB) free!
[1559214706] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (163360 kB) free!
[1559214941] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.29, 6.47, 19.86
[1559214999] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9504876 kB) free!
[1559215062] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (175864 kB) free!
[1559215073] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 3 warn out of 362 processes
[1559215119] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7778816 kB) free!
[1559215193] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 1 warn out of 360 processes
[1559215239] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.4% (5637440 kB) free!
[1559215300] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (4606888 kB) free!
[1559215306] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.3% (3458032 kB) free.
[1559215313] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 1 warn out of 361 processes
[1559215376] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 12 users currently logged in
[1559215421] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.4% (4521004 kB) free!
[1559215463] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4825268 kB) free!
[1559215496] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 12 users currently logged in
[1559215541] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (4438800 kB) free!
[1559215583] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4733488 kB) free!
[1559215617] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1559215703] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.5% (2821784 kB) free!
[1559215840] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4935872 kB) free!
[1559215906] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.9% (2421836 kB) free!
[1559216026] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.6% (3659892 kB) free.
[1559216141] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1559216141] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2167472 kB) free!
[1559216187] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 450.89, 369.78, 191.80
[1559216303] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (1311632 kB) free!
[1559216307] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 351.89, 374.75, 216.06
[1559216427] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 49.10, 251.38, 190.09
[1559216440] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12447104 kB) free.
[1559216513] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1559216515] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 16.6% (1220328 kB) free!
[1559216635] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 15.0% (1105900 kB) free!
[1559216741] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.45, 0.36, 3.12
[1559216755] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 13.4% (984968 kB) free!
[1559217226] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (1134384 kB) free!
[1559217346] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.4% (2988420 kB) free!
[1559217355] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (24.9% change >= 20.0% threshold)
[1559217355] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (268968 kB) free!
[1559217467] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (154672 kB) free!
[1559217576] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1559217777] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 68.78, 169.15, 99.05
[1559217897] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 9.65, 113.33, 87.09
[1559218017] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.82, 76.02, 76.60
[1559218024] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.6% (2866044 kB) free!
[1559218176] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559218296] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559218416] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559218559] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 418 processes
[1559218840] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9857912 kB) free!
[1559218913] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 357 processes
[1559218960] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.5% (8135856 kB) free!
[1559219033] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1559219080] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.0% (6390896 kB) free!
[1559219107] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.54, 5.95, 5.24
[1559219153] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1559219159] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 434 processes
[1559219224] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1559219224] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1871196 kB) free!
[1559219227] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.54, 5.95, 5.34
[1559219279] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 450 processes
[1559219347] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.62, 5.86, 5.38
[1559219399] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 452 processes
[1559219753] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1559219753] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1559219756] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.1% (6844308 kB) free.
[1559219817] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.60, 13.22, 17.85
[1559219824] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.6% (23354920 kB) free.
[1559219999] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1559220280] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (4015712 kB) free!
[1559220353] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 358 processes
[1559220473] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1559220599] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 426 processes
[1559220719] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 431 processes
[1559220839] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 433 processes
[1559220880] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11950524 kB) free.
[1559221024] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195616 kB) free!
[1559221145] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177448 kB) free!
[1559221227] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.46, 2.57, 13.01
[1559221265] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195220 kB) free!
[1559221389] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46560 MB (20% inode=99%):
[1559221509] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46531 MB (20% inode=99%):
[1559221618] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.05, 1.27, 3.46
[1559221629] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45567 MB (20% inode=99%):
[1559221667] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.1% (2929648 kB) free!
[1559222016] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1559222267] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (176220 kB) free!
[1559222348] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.92, 4.02, 4.78
[1559222616] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559222680] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9819280 kB) free!
[1559222736] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559222800] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8027636 kB) free!
[1559222856] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559222920] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.7% (6243616 kB) free!
[1559223027] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.72, 1.82, 3.56
[1559223239] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 394 processes
[1559223548] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.18, 5.77, 5.06
[1559223668] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.27, 5.75, 5.13
[1559223788] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.71, 5.94, 5.29
[1559224067] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.0% (3404032 kB) free.
[1559224120] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4045176 kB) free!
[1559224439] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 411 processes
[1559224559] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 389 processes
[1559225110] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4900072 kB) free!
[1559225159] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 421 processes
[1559225170] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4891252 kB) free!
[1559225230] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 19.9% (4882776 kB) free!
[1559225267] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (163976 kB) free!
[1559225279] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 426 processes
[1559225290] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 19.9% (4875356 kB) free!
[1559225320] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11459808 kB) free.
[1559225387] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (1151928 kB) free!
[1559225399] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 426 processes
[1559225507] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (167328 kB) free!
[1559226107] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.1% (2619764 kB) free!
[1559226356] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 12.6% (928400 kB) free!
[1559226456] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1559226463] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559226476] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (713384 kB) free!
[1559226596] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (504212 kB) free!
[1559226707] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (164096 kB) free!
[1559227056] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559227120] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9339632 kB) free!
[1559227176] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559227196] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.4% (6798560 kB) free.
[1559227240] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7400588 kB) free!
[1559227296] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559227360] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.3% (5553040 kB) free!
[1559227541] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559227561] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.1% (23469428 kB) free.
[1559227598] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.5% (22426896 kB) free.
[1559227653] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 35.7% (5787000 kB) free.
[1559227715] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.8% (2895584 kB) free!
[1559227781] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.0% (14946000 kB) free.
[1559227822] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.1% (21829320 kB) free.
[1559227865] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.8% (15090416 kB) free.
[1559227961] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4882504 kB) free!
[1559228218] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1559228263] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (167004 kB) free!
[1559228340] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165612 kB) free!
[1559228383] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.3% (3462144 kB) free.
[1559228474] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (178832 kB) free!
[1559228594] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (178972 kB) free!
[1559228714] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (151272 kB) free!
[1559228873] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 360 processes
[1559228983] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (156700 kB) free!
[1559228993] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 360 processes
[1559229104] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.9% (3557364 kB) free.
[1559229113] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1559229361] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.7% (3134404 kB) free!
[1559229481] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.4% (3071032 kB) free!
[1559229601] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.2% (3013664 kB) free!
[1559229631] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (2431996 kB) free!
[1559229751] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (2364824 kB) free!
[1559229788] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.18, 4.59, 4.77
[1559229871] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2303176 kB) free!
[1559230006] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194168 kB) free!
[1559230126] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183752 kB) free!
[1559230199] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 446 processes
[1559230201] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176640 kB) free!
[1559230246] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (656156 kB) free!
[1559230303] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.6% (2201400 kB) free!
[1559230423] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (176196 kB) free!
[1559230542] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 361.18, 266.39, 126.72
[1559230543] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.0% (3731928 kB) free.
[1559230662] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 81.37, 204.28, 121.97
[1559230782] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 11.36, 136.79, 107.22
[1559230799] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 405 processes
[1559230919] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 409 processes
[1559230961] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10794528 kB) free.
[1559231039] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 427 processes
[1559231143] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (160764 kB) free!
[1559231263] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.1% (2937752 kB) free!
[1559231383] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (170404 kB) free!
[1559231514] SERVICE FLAPPING ALERT: webservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559232096] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559232239] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 440 processes
[1559232582] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.38, 0.71, 15.84
[1559232583] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.1% (2930812 kB) free!
[1559232696] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559232761] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9531108 kB) free!
[1559232816] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559232839] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 393 processes
[1559232881] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7858076 kB) free!
[1559232937] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559232959] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 392 processes
[1559233001] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.6% (6206172 kB) free!
[1559233183] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (164588 kB) free!
[1559233782] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.18, 0.26, 4.54
[1559234159] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 394 processes
[1559234279] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 397 processes
[1559234383] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.3% (3457872 kB) free.
[1559234400] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 401 processes
[1559234802] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3401708 kB) free!
[1559234982] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 468.37, 293.55, 128.12
[1559235000] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 386 processes
[1559235102] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 471.91, 352.74, 169.86
[1559235222] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1559235222] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 69.34, 241.43, 151.54
[1559235271] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.5% (2815044 kB) free!
[1559235402] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.5% (9598196 kB) free!
[1559235583] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (161544 kB) free!
[1559235600] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 392 processes
[1559235703] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.3% (2813140 kB) free!
[1559235720] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 391 processes
[1559235823] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (158900 kB) free!
[1559236313] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 361 processes
[1559236320] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 395 processes
[1559236434] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1559236440] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 400 processes
[1559236537] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1559236560] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 400 processes
[1559236827] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 466.64, 376.71, 192.50
[1559236947] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 429.14, 398.65, 223.21
[1559237023] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.2% (1984376 kB) free!
[1559237067] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 59.17, 267.24, 196.33
[1559237138] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559237258] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559237378] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559237622] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.25, 0.30, 11.63
[1559237623] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (235088 kB) free!
[1559238596] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559238822] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.43, 0.25, 3.36
[1559239423] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.9% (1930288 kB) free!
[1559239467] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.51, 1.62, 16.30
[1559239560] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1559240160] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 378 processes
[1559240280] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 378 processes
[1559240400] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 378 processes
[1559240818] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 66.17, 153.63, 89.28
[1559240890] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 45.7% (11208056 kB) free.
[1559240938] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 9.17, 102.90, 78.49