[1567623347] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1567625506] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 33.1% (8164192 kB) free.
[1567625748] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 293 processes
[1567626100] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (249372 kB) free!
[1567626106] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.1% (3235320 kB) free!
[1567626214] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (390008 kB) free!
[1567626220] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (239844 kB) free!
[1567626226] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (2360876 kB) free!
[1567626334] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (859392 kB) free!
[1567626340] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (238936 kB) free!
[1567626347] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1813552 kB) free!
[1567626454] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397324 kB) free!
[1567626948] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 289 processes
[1567627548] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.90, 5.75, 5.09
[1567627668] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.59, 5.43, 5.06
[1567627788] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.54, 5.04, 4.95
[1567630332] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (213684 kB) free!
[1567631532] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.4% (45058936 kB) free.
[1567633023] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 42.7% (10477592 kB) free.
[1567633194] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 47.6% (11656120 kB) free.
[1567643740] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 95.7% (33524976 kB) free.
[1567643854] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.5% (15196164 kB) free!
[1567644454] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (4269684 kB) free!
[1567644947] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 46.4% (11451148 kB) free.
[1567646854] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 58.0% (76488092 kB) free.
[1567648623] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.3% (2758196 kB) free!
[1567648743] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186264 kB) free!
[1567648794] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188312 kB) free!
[1567648863] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192868 kB) free!
[1567648914] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170892 kB) free!
[1567649034] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189396 kB) free!
[1567649147] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (209336 kB) free!
[1567649267] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (202940 kB) free!
[1567649387] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189484 kB) free!
[1567656587] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.2% (22756124 kB) free.
[1567656636] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.5% (22416488 kB) free.
[1567656691] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.7% (40752428 kB) free.
[1567656752] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 96.0% (23533508 kB) free.
[1567659596] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4868216 kB) free!
[1567659716] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.1% (2996460 kB) free!
[1567659836] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.6% (2852864 kB) free!
[1567660436] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (1223896 kB) free!
[1567661867] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1567661987] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1567662107] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1567663298] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 315 processes
[1567663336] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 30.34, 66.94, 42.13
[1567663418] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1567663456] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 5.42, 45.39, 37.25
[1567663576] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.63, 30.73, 32.86
[1567664176] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.43, 5.40, 17.91
[1567665455] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 28422 MB (20% inode=99%):
[1567665575] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 28422 MB (20% inode=99%):
[1567665696] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 28422 MB (20% inode=99%):
[1567665976] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.36, 1.62, 3.85
[1567667497] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.6% (7681000 kB) free!
[1567667617] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 14.2% (7014392 kB) free!
[1567667737] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.1% (6431608 kB) free!
[1567668107] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1567669537] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4747264 kB) free!
[1567675193] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.7% (40762336 kB) free.
[1567675234] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.5% (23323124 kB) free.
[1567675896] SERVICE ALERT: sig2-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 125275 MB (90% inode=99%):
[1567677044] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.2% (3759372 kB) free!
[1567677104] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1567677164] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (1832112 kB) free!
[1567677224] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 402 processes
[1567677284] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1789028 kB) free!
[1567677344] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 404 processes
[1567677944] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1567678932] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9337452 kB) free!
[1567679018] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1567679052] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.6% (9172768 kB) free!
[1567679138] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1567679172] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.3% (9013276 kB) free!
[1567679258] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1567680647] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4804576 kB) free!
[1567680768] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4766412 kB) free!
[1567680888] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (4726488 kB) free!
[1567681197] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9674652 kB) free!
[1567681317] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8854344 kB) free!
[1567681437] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.4% (8072468 kB) free!
[1567681720] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.8% (23147000 kB) free.
[1567682037] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4879780 kB) free!
[1567682772] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4117764 kB) free!
[1567683888] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1698784 kB) free!
[1567684665] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.35, 7.43, 5.71
[1567684785] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.29, 6.42, 5.58
[1567684905] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 8.79, 6.74, 5.75
[1567685323] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4753580 kB) free!
[1567685443] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (1699028 kB) free!
[1567685563] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (850172 kB) free!
[1567687777] SERVICE ALERT: sig5-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.08, 5.45, 5.08
[1567687897] SERVICE ALERT: sig5-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.72, 5.77, 5.24
[1567688018] SERVICE ALERT: sig5-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.24, 5.56, 5.23
[1567689159] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4747008 kB) free!
[1567689279] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4630252 kB) free!
[1567689399] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.5% (4529424 kB) free!
[1567691988] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.51, 6.75, 5.45
[1567692108] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.95, 6.83, 5.63
[1567692216] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.5% (23069136 kB) free.
[1567692218] SERVICE ALERT: sig5-vm2;CPU Load;OK;HARD;1;OK - load average: 3.78, 4.32, 4.83
[1567692228] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.10, 6.84, 5.78
[1567692252] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 95.4% (23379276 kB) free.
[1567692285] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 71.5% (35257372 kB) free.
[1567692326] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 46.0% (22655232 kB) free.
[1567692399] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2235080 kB) free!
[1567692823] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.4% (2560664 kB) free!
[1567692828] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.72, 3.44, 4.62
[1567692943] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (1818476 kB) free!
[1567693063] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (880396 kB) free!
[1567693905] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.92, 3.53, 4.68
[1567695292] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9511752 kB) free!
[1567695344] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1567695412] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.0% (9365044 kB) free!
[1567695464] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1567695532] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.9% (8839480 kB) free!
[1567696732] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3401376 kB) free!
[1567700703] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.2% (23235164 kB) free.
[1567700736] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1567700736] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 71.5% (35229424 kB) free.
[1567700771] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.8% (23231752 kB) free.
[1567701464] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1567701584] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1567702506] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.2% (3759956 kB) free!
[1567702626] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.9% (3671620 kB) free!
[1567702746] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.9% (3674488 kB) free!
[1567703265] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 448 processes
[1567703325] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 456 processes
[1567703346] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1567703346] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183156 kB) free!
[1567703385] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;3;CPU WARNING: 2 warn out of 442 processes
[1567703446] SERVICE ALERT: webserver2;CPU_Procs;WARNING;HARD;4;CPU WARNING: 2 warn out of 443 processes
[1567706746] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1567707346] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 441 processes