[1568140557] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1568140680] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1568143917] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 436 processes
[1568143977] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 434 processes
[1568144148] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 294 processes
[1568144748] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 296 processes
[1568144976] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (9968808 kB) free!
[1568145096] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (10062148 kB) free!
[1568145216] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (10058900 kB) free!
[1568150367] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.7% (45193892 kB) free.
[1568151227] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1568151347] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1568152140] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 46.6% (11430216 kB) free.
[1568152212] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 44.6% (10932656 kB) free.
[1568154816] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.1% (13256544 kB) free!
[1568155312] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.83, 4.50, 4.97
[1568155416] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (13239476 kB) free!
[1568155912] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.24, 4.97, 5.03
[1568156033] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.92, 5.33, 5.15
[1568156153] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.70, 5.16, 5.12
[1568156616] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1568156616] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.1% (13290420 kB) free!
[1568157353] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.98, 2.58, 3.73
[1568162016] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (12435304 kB) free!
[1568162153] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.53, 7.11, 5.58
[1568162273] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.45, 6.95, 5.71
[1568162393] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1568162393] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.78, 6.55, 5.72
[1568162616] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.7% (21962904 kB) free!
[1568163216] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 35.0% (46085500 kB) free.
[1568164077] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 461 processes
[1568164137] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 456 processes
[1568164212] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4667680 kB) free!
[1568164336] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181064 kB) free!
[1568164457] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173888 kB) free!
[1568164741] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192000 kB) free!
[1568164862] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191872 kB) free!
[1568164982] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190560 kB) free!
[1568169594] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1568170793] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.53, 3.16, 4.23
[1568174616] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1568177991] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1568178111] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1568178231] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1568180747] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1568180867] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1568180993] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.67, 5.70, 5.32
[1568181113] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.66, 5.11, 5.16
[1568181233] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.02, 4.08, 4.78
[1568182637] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1279716 kB) free!
[1568183031] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1568183646] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 52.8% (13039260 kB) free.
[1568183657] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 78.6% (19271196 kB) free.
[1568183683] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.9% (22758884 kB) free.
[1568183697] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.5% (14551836 kB) free.
[1568183719] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 63.1% (31079156 kB) free.
[1568183987] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 93.5% (22920108 kB) free.
[1568184099] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1568184253] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195664 kB) free!
[1568184322] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (220636 kB) free!
[1568184373] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192372 kB) free!
[1568184442] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (231340 kB) free!
[1568184493] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186348 kB) free!
[1568184562] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (248728 kB) free!
[1568185279] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.7% (39283592 kB) free.
[1568185299] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1568185336] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.4% (22062900 kB) free.
[1568185943] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192844 kB) free!
[1568186033] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.72, 6.09, 5.15
[1568186063] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191152 kB) free!
[1568186153] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.63, 6.61, 5.45
[1568186183] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1568186183] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180744 kB) free!
[1568186273] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.14, 6.72, 5.63
[1568186325] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.1% (21997188 kB) free.
[1568186489] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (212072 kB) free!
[1568186609] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (260508 kB) free!
[1568186729] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1568186729] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (214444 kB) free!
[1568186934] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191372 kB) free!
[1568187054] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177376 kB) free!
[1568187174] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190412 kB) free!
[1568187706] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.05, 5.97, 5.02
[1568187826] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.09, 5.64, 5.01
[1568187946] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.71, 5.52, 5.05
[1568188974] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.4% (19843052 kB) free.
[1568189024] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 93.94, 29.85, 11.59
[1568189129] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 68.1% (33565236 kB) free.
[1568189144] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 130.85, 65.24, 26.77
[1568189264] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 135.18, 88.57, 39.98
[1568189273] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 49.44, 16.52, 9.96
[1568189300] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.02, 7.13, 6.35
[1568189420] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.38, 5.23, 5.75
[1568189540] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.31, 3.82, 5.16
[1568189574] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177656 kB) free!
[1568189694] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181444 kB) free!
[1568189729] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (239748 kB) free!
[1568189814] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172928 kB) free!
[1568189850] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (243396 kB) free!
[1568189970] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (242948 kB) free!
[1568190140] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.86, 2.51, 3.84
[1568190740] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 237.57, 148.11, 65.72
[1568190860] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 239.18, 178.28, 86.76
[1568190980] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 238.92, 198.36, 105.22
[1568191064] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.87, 2.86, 17.55
[1568191673] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.01, 9.66, 16.83
[1568192141] SERVICE ALERT: sig5-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.97, 5.65, 5.34
[1568192261] SERVICE ALERT: sig5-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.06, 5.94, 5.47
[1568192265] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.24, 5.54, 5.06
[1568192381] SERVICE ALERT: sig5-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.32, 5.31, 5.30
[1568192385] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 4.29, 5.19, 4.99
[1568192780] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1568192780] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.33, 2.24, 18.14
[1568192864] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.12, 1.48, 3.86
[1568192969] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9269244 kB) free!
[1568192981] SERVICE ALERT: sig5-vm2;CPU Load;OK;HARD;1;OK - load average: 4.04, 4.75, 5.00
[1568193089] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.8% (8796508 kB) free!
[1568193209] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.9% (8349260 kB) free!
[1568193705] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 17.7% (4160148 kB) free!
[1568193825] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 16.3% (3834252 kB) free!
[1568193941] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 333 processes
[1568193946] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 14.7% (3455588 kB) free!
[1568193950] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 12.9% (4511652 kB) free!
[1568194061] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1568194070] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 11.3% (3953424 kB) free!
[1568194181] SERVICE ALERT: sig5-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.11, 5.18, 5.12
[1568194191] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (3507820 kB) free!
[1568194301] SERVICE ALERT: sig5-vm2;CPU Load;OK;HARD;1;OK - load average: 4.51, 4.85, 5.00
[1568194409] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3839180 kB) free!
[1568194546] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 97.3% (22929780 kB) free.
[1568194791] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 91.7% (32112348 kB) free.
[1568195180] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.72, 1.88, 3.80
[1568195262] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 333 processes
[1568195382] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1568195444] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 65.7% (16224208 kB) free.
[1568195482] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 50.4% (24844784 kB) free.
[1568195511] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 30.7% (15128032 kB) free.
[1568196056] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191884 kB) free!
[1568196086] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (1697056 kB) free!
[1568196103] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1568196176] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197532 kB) free!
[1568196206] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (264316 kB) free!
[1568196223] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 341 processes
[1568196287] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (4782212 kB) free!
[1568196296] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194912 kB) free!
[1568196326] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (223460 kB) free!
[1568196343] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 341 processes
[1568196407] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.7% (4595544 kB) free!
[1568196527] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.8% (4361996 kB) free!
[1568198327] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1781000 kB) free!
[1568198637] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 484 processes
[1568198697] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 492 processes
[1568198982] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1568199102] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1568199267] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4684612 kB) free!
[1568199387] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (4548504 kB) free!
[1568199507] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (4370052 kB) free!
[1568200506] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.1% (1797788 kB) free!
[1568200626] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.1% (1797572 kB) free!
[1568200667] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1568200746] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (1800340 kB) free!
[1568200788] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1568201307] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190924 kB) free!
[1568201746] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.31, 4.60, 4.91
[1568204780] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1568205225] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.81, 5.70, 5.29
[1568205345] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.17, 5.86, 5.40
[1568205465] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.02, 5.85, 5.44
[1568205896] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1568207126] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1568207515] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9465632 kB) free!
[1568207635] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.8% (9288928 kB) free!
[1568207755] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.4% (9085836 kB) free!
[1568209065] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.05, 3.33, 4.36
[1568209543] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 345 processes
[1568209787] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 364 processes
[1568210387] SERVICE ALERT: secondary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 364 processes
[1568210503] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 15.4% (3785500 kB) free!
[1568210623] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.4% (3778304 kB) free!
[1568210743] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 31.9% (7814296 kB) free.
[1568210755] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4618568 kB) free!
[1568210932] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.8% (20686696 kB) free.
[1568210987] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 363 processes
[1568210996] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.0% (22309220 kB) free.
[1568211032] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 56.6% (27883908 kB) free.
[1568211063] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 40.9% (20174520 kB) free.
[1568211088] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.1% (22581380 kB) free.
[1568211536] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165748 kB) free!
[1568211656] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185084 kB) free!
[1568211776] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174200 kB) free!
[1568211943] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (991124 kB) free!
[1568212063] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189788 kB) free!
[1568212183] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184888 kB) free!
[1568212187] SERVICE FLAPPING ALERT: secondary-server;CPU_Procs;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1568212187] SERVICE ALERT: secondary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 365 processes
[1568212235] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (223500 kB) free!
[1568212355] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (231000 kB) free!
[1568212475] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (223300 kB) free!
[1568213143] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1568213383] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 23.7% (5819088 kB) free.
[1568213743] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 348 processes
[1568213983] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179840 kB) free!
[1568213987] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 368 processes
[1568214103] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196600 kB) free!
[1568214224] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187608 kB) free!
[1568214343] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1568214343] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 341 processes
[1568214587] SERVICE ALERT: secondary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 362 processes
[1568214824] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1568214824] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 27.8% (6823808 kB) free.
[1568214944] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 346 processes
[1568215645] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.1% (22478508 kB) free.
[1568215684] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 65.7% (32395568 kB) free.
[1568215787] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 365 processes
[1568215791] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 20.0% (4913656 kB) free!
[1568215911] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 11.0% (2688492 kB) free!
[1568216024] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (1885992 kB) free!
[1568216031] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 11.0% (2687440 kB) free!
[1568216144] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1568216144] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 29.2% (7157928 kB) free.
[1568216256] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169416 kB) free!
[1568216376] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180164 kB) free!
[1568216387] SERVICE ALERT: secondary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 361 processes
[1568216496] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1568216496] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192588 kB) free!
[1568216632] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 25.1% (6149168 kB) free.
[1568216893] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (218344 kB) free!
[1568217013] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (235616 kB) free!
[1568217133] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1568217133] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224864 kB) free!
[1568217699] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1568218299] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1568220003] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (212120 kB) free!
[1568220096] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189748 kB) free!
[1568220123] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (213724 kB) free!
[1568220216] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176144 kB) free!
[1568220243] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194556 kB) free!
[1568220336] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196900 kB) free!
[1568221673] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.60, 2.66, 4.02
[1568223587] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1568223707] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1568225873] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.76, 5.67, 5.15
[1568225993] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.48, 5.33, 5.10
[1568226096] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1568226113] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.06, 4.81, 4.93
[1568226344] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)