[1590690784] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (3198744 kB) free!
[1590691347] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590691384] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.5% (4637384 kB) free.
[1590691984] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.0% (2595148 kB) free!
[1590692104] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.5% (2847840 kB) free!
[1590692224] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (3034472 kB) free!
[1590692825] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.2% (4088432 kB) free.
[1590693425] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.2% (1981004 kB) free!
[1590693545] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.5% (1871616 kB) free!
[1590693665] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.1% (1811632 kB) free!
[1590693990] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29470560 kB) free.
[1590694265] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1392928 kB) free!
[1590694438] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (193940 kB) free!
[1590694591] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (4068972 kB) free!
[1590694711] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (391032 kB) free!
[1590694832] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410124 kB) free!
[1590695465] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.9% (2102816 kB) free!
[1590695638] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 84.9% (20815076 kB) free.
[1590697232] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27391164 kB) free.
[1590697832] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4574796 kB) free!
[1590697952] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (703480 kB) free!
[1590698072] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400692 kB) free!
[1590699160] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.06, 5.12, 5.03
[1590699272] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27954676 kB) free.
[1590699280] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.74, 4.24, 4.72
[1590699872] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5415760 kB) free!
[1590699992] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1618504 kB) free!
[1590700112] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389524 kB) free!
[1590700724] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.6% (3338952 kB) free!
[1590700915] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.7% (44180764 kB) free.
[1590701063] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 89.0% (43879972 kB) free.
[1590701312] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24683696 kB) free!
[1590701490] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2732) < 203.212.70.170 [29/May/2020:03:00:44 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1590701610] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(66) < 182.18.144.246 [29/May/2020:03:01:43 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1590701627] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1590701730] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590701771] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(482) < 2020/05/29 03:01:32 [error] 22134#0: *91668730 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1590701891] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590701912] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (3154492 kB) free!
[1590702000] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(365) < 49.50.86.152 [29/May/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1590702065] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3639580 kB) free.
[1590702121] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590702512] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15755420 kB) free!
[1590703113] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31444152 kB) free.
[1590703713] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8179108 kB) free!
[1590703833] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4374988 kB) free!
[1590703953] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (838656 kB) free!
[1590704679] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.9% (10268572 kB) free.
[1590705153] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26851188 kB) free.
[1590705279] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (556448 kB) free!
[1590705399] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (245644 kB) free!
[1590705519] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (248616 kB) free!
[1590705524] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184964 kB) free!
[1590705753] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (2479212 kB) free!
[1590705873] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (400688 kB) free!
[1590705993] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32089068 kB) free.
[1590706593] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5529132 kB) free!
[1590706713] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1697276 kB) free!
[1590706833] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (884640 kB) free!
[1590707433] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (15964560 kB) free!
[1590707919] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 38.9% (9541840 kB) free.
[1590708033] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26874112 kB) free.
[1590708065] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1590708633] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6890160 kB) free!
[1590708754] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35888484 kB) free.
[1590709119] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186592 kB) free!
[1590709240] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193728 kB) free!
[1590709353] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (14999148 kB) free!
[1590709360] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1590709360] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (145096 kB) free!
[1590709473] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (11191192 kB) free!
[1590709593] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36072360 kB) free.
[1590710193] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16502668 kB) free!
[1590710283] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 42.1% (10386152 kB) free.
[1590710313] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (12753076 kB) free!
[1590710433] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (9047388 kB) free!
[1590710564] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 70.8% (17353308 kB) free.
[1590710924] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 41.2% (10095280 kB) free.
[1590711033] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16839948 kB) free!
[1590712233] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35237360 kB) free.
[1590712833] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12729660 kB) free!
[1590712953] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8758176 kB) free!
[1590713073] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34747480 kB) free.
[1590713283] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (882532 kB) free!
[1590713403] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173024 kB) free!
[1590713523] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181460 kB) free!
[1590713673] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14829616 kB) free!
[1590713793] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (10931588 kB) free!
[1590713913] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (7069972 kB) free!
[1590714238] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 13.9% (3403584 kB) free!
[1590714298] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 13.9% (3398056 kB) free!
[1590714358] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 13.9% (3397660 kB) free!
[1590714419] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 13.9% (3402908 kB) free!
[1590714514] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (15960156 kB) free!
[1590715714] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33686628 kB) free.
[1590716314] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14185028 kB) free!
[1590716434] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (10356872 kB) free!
[1590716554] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (6622388 kB) free!
[1590717792] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (1671384 kB) free!
[1590718092] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/29 07:34:04 [error] 22135#0: *91699176 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590718202] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1590718212] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590718322] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1590718354] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (18408220 kB) free!
[1590718954] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (433172 kB) free!
[1590720154] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19969744 kB) free!
[1590720160] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1590720164] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176112 kB) free!
[1590720210] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1590720228] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1590720285] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186752 kB) free!
[1590720331] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1590720348] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1590720405] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190808 kB) free!
[1590720451] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1590720468] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1590720754] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (439156 kB) free!
[1590721051] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1590721068] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1590721124] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (4520324 kB) free!
[1590721244] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4404420 kB) free!
[1590721364] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.4% (4266760 kB) free!
[1590721954] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18671136 kB) free!
[1590722268] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 344 processes
[1590722388] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 343 processes
[1590722508] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1590722554] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389604 kB) free!
[1590723164] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2152344 kB) free!
[1590723280] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.75, 5.99, 5.29
[1590723400] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 12.55, 8.25, 6.18
[1590723520] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.28, 7.74, 6.25
[1590723708] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1590723708] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 349 processes
[1590723739] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 61.9% (15175264 kB) free.
[1590723754] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.9% (20944352 kB) free!
[1590724354] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (452616 kB) free!
[1590725554] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.8% (24818324 kB) free!
[1590726065] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (1103068 kB) free!
[1590726154] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1893724 kB) free!
[1590726185] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.5% (2036300 kB) free!
[1590726305] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.1% (4234392 kB) free.
[1590726754] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14377936 kB) free!
[1590727308] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1590727354] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29303808 kB) free.
[1590727954] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (5049832 kB) free!
[1590728074] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (579704 kB) free!
[1590728194] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (423016 kB) free!
[1590728705] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.9% (2744068 kB) free!
[1590728825] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.3% (3454908 kB) free.
[1590729394] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.8% (22153548 kB) free!
[1590729994] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (428780 kB) free!
[1590730354] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.9% (3784104 kB) free!
[1590730364] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.7% (3116660 kB) free!
[1590730474] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.5% (3252556 kB) free!
[1590730594] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.1% (2716252 kB) free!
[1590730594] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.7% (22011896 kB) free!
[1590730964] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1590730964] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1618792 kB) free!
[1590731131] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(632) < 103.255.146.164 [29/May/2020:11:10:16 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590731194] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (307108 kB) free!
[1590731195] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (443620 kB) free!
[1590731225] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.7% (2883260 kB) free!
[1590731251] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590731345] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.9% (2252480 kB) free!
[1590731465] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (3015680 kB) free!
[1590731795] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16610344 kB) free!
[1590732395] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (424992 kB) free!
[1590732708] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1590733051] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(44) < 103.255.146.164 [29/May/2020:11:46:28 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590733171] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590733692] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 17.0% (5953180 kB) free!
[1590733812] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 91.4% (32022780 kB) free.
[1590733865] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.1% (3918716 kB) free.
[1590734194] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.0% (15935748 kB) free.
[1590734795] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29116500 kB) free.
[1590735395] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (2423612 kB) free!
[1590735515] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (390732 kB) free!
[1590735635] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (423484 kB) free!
[1590735665] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.4% (2835836 kB) free!
[1590735785] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (3241204 kB) free!
[1590735905] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (3135908 kB) free!
[1590736315] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8797152 kB) free!
[1590736339] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.8% (4105420 kB) free!
[1590736435] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.3% (8540656 kB) free!
[1590736459] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.7% (3837452 kB) free!
[1590736555] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.8% (8265992 kB) free!
[1590736579] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.5% (3541956 kB) free!
[1590737063] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9384808 kB) free!
[1590737105] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1590737105] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1625040 kB) free!
[1590737179] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1877908 kB) free!
[1590737183] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (9145748 kB) free!
[1590737303] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (8904004 kB) free!
[1590737972] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(468) < 103.255.146.164 [29/May/2020:13:01:29 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590738092] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590738192] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (1518252 kB) free!
[1590738305] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.6% (2694092 kB) free!
[1590738355] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (3535940 kB) free!
[1590738792] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.8% (1761708 kB) free!
[1590739308] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1590739704] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4163092 kB) free!
[1590740105] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1301624 kB) free!
[1590740705] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.1% (2458892 kB) free!
[1590740794] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 11.6% (2611068 kB) free!
[1590740914] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (2042772 kB) free!
[1590741034] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1444828 kB) free!
[1590742364] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590743105] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1528088 kB) free!
[1590743412] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 16.3% (5718648 kB) free!
[1590743528] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 20.40, 24.38, 20.69
[1590743532] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 14.6% (5116304 kB) free!
[1590743652] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 12.8% (4492012 kB) free!
[1590743705] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (3101644 kB) free!
[1590743920] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 26.60, 26.53, 20.84
[1590744252] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (1679848 kB) free!
[1590744520] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.66, 11.23, 16.19
[1590744634] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.8% (16120428 kB) free.
[1590744635] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15821660 kB) free!
[1590744692] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(153) < 185.255.8.58 [29/May/2020:15:00:04 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1590744729] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.98, 11.67, 17.71
[1590744813] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590744853] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 91.3% (31991392 kB) free.
[1590744905] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1253596 kB) free!
[1590745235] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (434464 kB) free!
[1590745508] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 18.4% (4540228 kB) free!
[1590745629] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4420860 kB) free!
[1590745749] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 43.4% (10711300 kB) free.
[1590745908] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1590745922] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1590746012] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(33) < 103.13.97.241 [29/May/2020:15:22:11 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1590746042] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1590746132] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590747708] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1590747828] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1590747948] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1590748235] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15222804 kB) free!
[1590748548] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1590748835] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (455164 kB) free!
[1590749105] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3372672 kB) free.
[1590749435] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16757336 kB) free!
[1590749705] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (3092176 kB) free!
[1590749825] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (408780 kB) free!
[1590749945] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1308400 kB) free!
[1590750035] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (446328 kB) free!
[1590750545] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.8% (2572468 kB) free!
[1590750635] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.7% (20701728 kB) free!
[1590751235] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (463204 kB) free!
[1590751546] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.0% (25035112 kB) free!
[1590751666] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 29.6% (38960932 kB) free.
[1590751745] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (470660 kB) free!
[1590751812] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/29 16:58:25 [error] 22135#0: *91786173 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590751834] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.6% (3060068 kB) free!
[1590751835] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (25733772 kB) free!
[1590751932] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590751955] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.0% (2694152 kB) free!
[1590752075] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.1% (2258820 kB) free!
[1590752345] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.4% (2184872 kB) free!
[1590752435] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (430660 kB) free!
[1590752642] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1590752675] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1590752675] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (179044 kB) free!
[1590752762] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1590752945] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.3% (3456792 kB) free.
[1590753035] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30677648 kB) free.
[1590753132] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/29 17:15:26 [error] 22135#0: *91790288 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590753252] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590753546] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (2952236 kB) free!
[1590753635] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (761576 kB) free!
[1590753666] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3523440 kB) free.
[1590753755] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (401020 kB) free!
[1590753875] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400484 kB) free!
[1590754266] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.4% (1686500 kB) free!
[1590754386] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3346244 kB) free.
[1590754475] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.4% (24312080 kB) free!
[1590754986] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (770380 kB) free!
[1590755053] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 18.2% (6362148 kB) free!
[1590755075] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (414908 kB) free!
[1590755106] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.2% (2467416 kB) free!
[1590755173] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 17.0% (5958584 kB) free!
[1590755226] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3646476 kB) free.
[1590755293] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 15.9% (5551272 kB) free!
[1590755675] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.1% (16193652 kB) free.
[1590755676] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32332112 kB) free.
[1590755826] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (1618508 kB) free!
[1590755893] SERVICE FLAPPING ALERT: sig5-vm2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1590755893] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 91.3% (31985828 kB) free.
[1590755946] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3295104 kB) free.
[1590756276] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (3865852 kB) free!
[1590756396] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (423864 kB) free!
[1590756516] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (424800 kB) free!
[1590756546] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (806540 kB) free!
[1590756666] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.8% (2723068 kB) free!
[1590756786] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.2% (3763428 kB) free.
[1590757116] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26063104 kB) free!
[1590757387] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.1% (2297908 kB) free!
[1590757507] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.5% (3816144 kB) free.
[1590757716] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (418172 kB) free!
[1590757992] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1627752 kB) free!
[1590758107] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.9% (2093764 kB) free!
[1590758227] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.4% (3960720 kB) free.
[1590758316] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29893780 kB) free.
[1590758592] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 11.0% (1783668 kB) free!
[1590758762] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 339 processes
[1590758827] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (155140 kB) free!
[1590758882] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1590758916] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (404856 kB) free!
[1590758947] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (155588 kB) free!
[1590759036] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (441084 kB) free!
[1590759067] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (864584 kB) free!
[1590759156] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (418876 kB) free!
[1590759756] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.4% (24245596 kB) free!
[1590760267] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.4% (4456996 kB) free.
[1590760356] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409048 kB) free!
[1590760867] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3246208 kB) free!
[1590760987] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.7% (4505800 kB) free.
[1590761266] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.9% (18277780 kB) free!
[1590761386] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 28.5% (37631796 kB) free.
[1590762110] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 23081 MB (17% inode=99%):
[1590762188] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.4% (2667636 kB) free!
[1590762230] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 22985 MB (17% inode=99%):
[1590762275] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.2% (3866844 kB) free!
[1590762308] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.6% (4651992 kB) free.
[1590762350] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 22977 MB (17% inode=99%):
[1590762395] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.6% (3509596 kB) free!
[1590762482] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 336 processes
[1590762515] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.2% (3186324 kB) free!
[1590762586] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.2% (24023836 kB) free!
[1590762602] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 342 processes
[1590762706] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 28.1% (37070216 kB) free.
[1590762722] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1590762756] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13764328 kB) free!
[1590763115] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1860000 kB) free!
[1590763356] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393276 kB) free!
[1590763392] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (1573284 kB) free!
[1590763508] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.5% (2191992 kB) free!
[1590763628] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.0% (4554788 kB) free.
[1590763992] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1590763992] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (1814776 kB) free!
[1590764228] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (1392376 kB) free!
[1590764348] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (166720 kB) free!
[1590764468] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.5% (2524448 kB) free!
[1590764506] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.3% (24068120 kB) free!
[1590764626] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 26.3% (34659640 kB) free.
[1590765122] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 3 warn out of 342 processes
[1590765242] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 3 warn out of 338 processes
[1590765362] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 338 processes
[1590765668] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (891812 kB) free!
[1590765756] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.8% (23502880 kB) free!
[1590765962] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1590766268] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.0% (4548044 kB) free.
[1590766356] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405388 kB) free!
[1590766716] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.9% (16142884 kB) free.
[1590766720] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.94, 3.67, 4.57
[1590767151] SERVICE ALERT: sig4-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 115118 MB (89% inode=99%):
[1590767293] SERVICE FLAPPING ALERT: sig5-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590767469] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (2845812 kB) free!
[1590767589] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.0% (4396604 kB) free.
[1590768189] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.4% (2835728 kB) free!
[1590768309] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.8% (4356880 kB) free.
[1590768756] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32760796 kB) free.
[1590768909] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.5% (2358724 kB) free!
[1590769029] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.6% (4324052 kB) free.
[1590769357] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (1725348 kB) free!
[1590769477] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (410140 kB) free!
[1590769597] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390184 kB) free!
[1590769629] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.1% (1962152 kB) free!
[1590769749] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.2% (4101284 kB) free.
[1590770197] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21608144 kB) free!
[1590770349] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (1487676 kB) free!
[1590770469] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3352160 kB) free.
[1590770797] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387620 kB) free!
[1590771069] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (1509452 kB) free!
[1590771189] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.6% (2047628 kB) free!
[1590771310] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.2% (3767560 kB) free.
[1590771397] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31535252 kB) free.
[1590771909] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (575968 kB) free!
[1590771997] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (2005220 kB) free!
[1590772029] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.3% (2164844 kB) free!
[1590772117] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (409696 kB) free!
[1590772149] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.5% (3661432 kB) free.
[1590772237] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392872 kB) free!
[1590772837] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.1% (23870996 kB) free!
[1590773349] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.7% (2072116 kB) free!
[1590773437] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398812 kB) free!
[1590773469] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.2% (3445260 kB) free.
[1590774037] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29908644 kB) free.
[1590774637] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (3091140 kB) free!
[1590774669] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (1366440 kB) free!
[1590774757] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (386992 kB) free!
[1590774789] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.5% (2685000 kB) free!
[1590774877] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406524 kB) free!
[1590774909] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.3% (3629276 kB) free.
[1590775116] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1590775392] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590775478] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20516808 kB) free!
[1590776078] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395676 kB) free!
[1590776678] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.9% (24854040 kB) free!