[1581186638] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21758064 kB) free!
[1581186758] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.8% (43203192 kB) free.
[1581187009] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.5% (3820732 kB) free!
[1581187049] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.9% (29543844 kB) free.
[1581187080] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165348 kB) free!
[1581187109] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (202028 kB) free!
[1581187129] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (203084 kB) free!
[1581187183] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(55) < 2020/02/09 00:02:46 [error] 10218#0: *51002664 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581187201] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174172 kB) free!
[1581187229] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (203500 kB) free!
[1581187249] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184488 kB) free!
[1581187303] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581187321] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182336 kB) free!
[1581187349] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175448 kB) free!
[1581187358] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20579812 kB) free!
[1581187478] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41581460 kB) free.
[1581188678] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22163716 kB) free!
[1581188798] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.7% (41732128 kB) free.
[1581189912] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1581189998] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19656432 kB) free!
[1581190032] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1581190118] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.3% (43919796 kB) free.
[1581190718] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19608708 kB) free!
[1581190838] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.3% (43913412 kB) free.
[1581191438] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18109420 kB) free!
[1581191558] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.0% (43559128 kB) free.
[1581191660] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 12.4% (6110100 kB) free!
[1581191780] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 10.8% (5298676 kB) free!
[1581191900] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4574116 kB) free!
[1581192158] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16167076 kB) free!
[1581192279] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41090520 kB) free.
[1581193478] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20669044 kB) free!
[1581193598] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43333936 kB) free.
[1581194798] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23520248 kB) free!
[1581194918] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.2% (43785196 kB) free.
[1581195450] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (1490436 kB) free!
[1581195570] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (229784 kB) free!
[1581195690] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (238372 kB) free!
[1581195721] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 47.9% (11736012 kB) free.
[1581196118] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (25111748 kB) free!
[1581196238] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.9% (44745864 kB) free.
[1581196290] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 83.9% (41338760 kB) free.
[1581197503] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(263) < 2020/02/09 03:01:32 [error] 10220#0: *51047966 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581197623] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581197753] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2417) < 203.212.70.170 [09/Feb/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1581197786] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(205) < 103.16.101.72 [09/Feb/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1581197873] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581197907] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581201223] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(15) < 2020/02/09 04:00:42 [error] 10218#0: *51062130 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581201343] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581204721] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177408 kB) free!
[1581204842] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189788 kB) free!
[1581204962] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163772 kB) free!
[1581204987] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 982.96, 337.53, 121.13
[1581205107] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 199.33, 283.40, 130.20
[1581205226] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 752 processes
[1581205249] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 47.3% (11674036 kB) free.
[1581205346] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 364 processes
[1581205349] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 84.5% (20704028 kB) free.
[1581205673] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4) < 172.105.59.173 [09/Feb/2020:05:08:01 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1581205793] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581206143] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(302) < 2020/02/09 05:22:37 [error] 10215#0: *51084380 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1581206263] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581206306] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(147) < 103.16.101.72 [09/Feb/2020:05:20:53 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1581206393] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(14) < 49.50.64.122 [09/Feb/2020:05:22:41 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1581206426] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581206513] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581206864] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/09 05:34:45 [error] 10208#0: *51088050 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1581206984] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581209050] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 88.4% (116560528 kB) free.
[1581209100] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.8% (39801424 kB) free.
[1581209142] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.3% (21889504 kB) free.
[1581210650] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (2130736 kB) free!
[1581210770] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (2129316 kB) free!
[1581210890] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2134320 kB) free!
[1581211632] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1581211752] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1581212190] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.7% (21887236 kB) free.
[1581215352] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1581215472] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1581227197] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.2% (3022840 kB) free!
[1581227317] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.6% (2869712 kB) free!
[1581227438] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2736424 kB) free!
[1581227549] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (1161668 kB) free!
[1581227660] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1581227669] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178524 kB) free!
[1581227708] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9069360 kB) free!
[1581227749] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194768 kB) free!
[1581227780] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 350 processes
[1581227789] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187896 kB) free!
[1581227828] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (2448512 kB) free!
[1581227869] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166868 kB) free!
[1581227901] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 351 processes
[1581227948] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224012 kB) free!
[1581227989] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185944 kB) free!
[1581228038] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177344 kB) free!
[1581228152] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.7% (3970856 kB) free!
[1581228272] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.4% (3688728 kB) free!
[1581228350] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183116 kB) free!
[1581228392] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.1% (3385876 kB) free!
[1581228470] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175204 kB) free!
[1581228501] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1581228548] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 45.4% (22352604 kB) free.
[1581228589] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 59.7% (14622792 kB) free.
[1581228590] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178404 kB) free!
[1581228638] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 21.0% (5169872 kB) free.
[1581228638] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (23539840 kB) free!
[1581228758] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27825872 kB) free.
[1581228993] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1765716 kB) free!
[1581229238] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175508 kB) free!
[1581229290] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (1924668 kB) free!
[1581229358] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174472 kB) free!
[1581229410] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (226784 kB) free!
[1581229478] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1581229478] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174396 kB) free!
[1581229490] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (238600 kB) free!
[1581229531] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (215084 kB) free!
[1581229610] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (223056 kB) free!
[1581229730] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 59.5% (29343992 kB) free.
[1581229913] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4) < 172.105.59.173 [09/Feb/2020:12:01:47 +0530] TCP 502 0 0 60.001 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.001, 0.000"
[1581230033] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(19) < 172.105.59.173 [09/Feb/2020:12:02:57 +0530] TCP 502 0 0 120.000 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1581230131] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 51.9% (25574968 kB) free.
[1581230153] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(2) < 172.105.59.173 [09/Feb/2020:12:03:57 +0530] TCP 502 0 0 180.001 "10.147.212.143:3001, 10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.001, 60.000, 60.000, 0.000"
[1581230193] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.7% (15653508 kB) free.
[1581230289] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.4% (21081592 kB) free.
[1581230338] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.2% (21623312 kB) free.
[1581230384] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(29) < 2020/02/09 12:00:48 [error] 10213#0: *51205728 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581230408] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.2% (21131064 kB) free.
[1581230504] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581230753] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581230899] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171204 kB) free!
[1581230901] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1581231020] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168764 kB) free!
[1581231140] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169560 kB) free!
[1581231672] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1581231792] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1581233447] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1581233567] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1581233775] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.6% (21122228 kB) free.
[1581234379] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (1473248 kB) free!
[1581234499] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164332 kB) free!
[1581234619] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166648 kB) free!
[1581236558] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (22012760 kB) free!
[1581236678] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.6% (45619424 kB) free.
[1581237278] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20077968 kB) free!
[1581237398] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44495420 kB) free.
[1581237998] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17958776 kB) free!
[1581238118] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43376168 kB) free.
[1581238190] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.4% (4502512 kB) free!
[1581238310] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.9% (4393836 kB) free!
[1581238430] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.7% (4336264 kB) free!
[1581238718] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (18944448 kB) free!
[1581238838] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42829472 kB) free.
[1581238967] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1581239087] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1581239438] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18452616 kB) free!
[1581239558] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41519320 kB) free.
[1581240158] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17614628 kB) free!
[1581240278] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39828596 kB) free.
[1581240953] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(152) < 92.42.104.166 [09/Feb/2020:15:00:16 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1581241073] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581241479] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18151536 kB) free!
[1581241599] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41376308 kB) free.
[1581241748] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9755620 kB) free!
[1581241868] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.8% (9756292 kB) free!
[1581241988] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.8% (9755524 kB) free!
[1581242199] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17256088 kB) free!
[1581242319] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40674364 kB) free.
[1581242919] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15782184 kB) free!
[1581243039] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38471072 kB) free.
[1581243230] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2173096 kB) free!
[1581243619] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 32.4% (7983512 kB) free.
[1581243931] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9791628 kB) free!
[1581243989] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 83.7% (20504476 kB) free.
[1581244051] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.8% (9733784 kB) free!
[1581244171] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.6% (9672180 kB) free!
[1581244239] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17680716 kB) free!
[1581244304] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(32) < 2020/02/09 16:00:40 [error] 10213#0: *51280899 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581244359] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39912020 kB) free.
[1581244424] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581244959] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15868148 kB) free!
[1581245079] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38524580 kB) free.
[1581246279] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17299968 kB) free!
[1581246399] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37975772 kB) free.
[1581247148] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4800700 kB) free!
[1581247219] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (4800688 kB) free!
[1581247268] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4699488 kB) free!
[1581247339] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.4% (4784848 kB) free!
[1581247388] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (4618332 kB) free!
[1581247459] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.3% (4771556 kB) free!
[1581247600] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18114776 kB) free!
[1581247720] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38867808 kB) free.
[1581248320] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14284360 kB) free!
[1581248440] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37778080 kB) free.
[1581249040] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15231820 kB) free!
[1581249160] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36147616 kB) free.
[1581250360] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14770956 kB) free!
[1581250480] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.9% (48692096 kB) free.
[1581250988] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2264856 kB) free!
[1581251659] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2136476 kB) free!
[1581252188] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4804604 kB) free!
[1581252571] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4923260 kB) free!
[1581257680] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25574264 kB) free!
[1581257800] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.1% (48946208 kB) free.
[1581258192] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1581258272] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 63.8% (15734992 kB) free.
[1581258302] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 56.6% (27880960 kB) free.
[1581258312] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1581258333] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 30.1% (14843944 kB) free.
[1581258363] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 71.7% (17575440 kB) free.
[1581258400] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (24631308 kB) free!
[1581258520] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.2% (46384448 kB) free.
[1581258824] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(42) < 2020/02/09 20:01:01 [error] 10216#0: *51356771 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581258944] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581260678] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.3% (2553820 kB) free!
[1581260798] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (2324932 kB) free!
[1581260918] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1831116 kB) free!
[1581264887] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1581265007] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1581269318] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1581271649] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (3186852 kB) free!
[1581271770] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (227876 kB) free!
[1581271890] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (222940 kB) free!