[1556908294] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10469640 kB) free.
[1556908894] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8965028 kB) free!
[1556909014] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.1% (7454624 kB) free!
[1556909134] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11404700 kB) free.
[1556910335] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9342776 kB) free!
[1556910455] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7868240 kB) free!
[1556910548] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 276.26, 164.62, 69.78
[1556910575] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11421184 kB) free.
[1556910668] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 279.09, 202.63, 95.21
[1556910788] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 52.63, 148.56, 89.02
[1556911175] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.3% (6042152 kB) free!
[1556911295] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10649228 kB) free.
[1556911563] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1556911748] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1556911895] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9369932 kB) free!
[1556912015] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7881108 kB) free!
[1556912135] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11416024 kB) free.
[1556912588] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.32, 0.69, 13.14
[1556912735] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6686604 kB) free!
[1556912855] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11420436 kB) free.
[1556912948] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556913069] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556913189] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556913457] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1556913788] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.75, 0.53, 3.95
[1556914300] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 249.15, 126.22, 50.30
[1556914420] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 257.06, 169.93, 75.52
[1556914540] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 111.92, 159.33, 84.13
[1556914656] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6092696 kB) free!
[1556914707] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 20.4% (4989516 kB) free.
[1556914776] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10580476 kB) free.
[1556916340] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.15, 0.61, 12.33
[1556916508] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4892252 kB) free!
[1556916628] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (4884628 kB) free!
[1556916748] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (4880856 kB) free!
[1556916789] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1556917121] SERVICE ALERT: sigqscluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1556917176] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7029624 kB) free!
[1556917242] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1556917296] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11431232 kB) free.
[1556917540] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.10, 0.15, 3.50
[1556917896] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6487376 kB) free!
[1556918016] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11439608 kB) free.
[1556918589] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556918616] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1556918709] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556918829] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556920417] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9265052 kB) free!
[1556920537] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7762964 kB) free!
[1556920657] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11441936 kB) free.
[1556921257] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7336188 kB) free!
[1556921377] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11444404 kB) free.
[1556921977] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6805488 kB) free!
[1556922097] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11444760 kB) free.
[1556922188] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 17.17, 147.79, 107.30
[1556922308] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 2.48, 98.94, 94.31
[1556922428] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.42, 66.26, 82.90
[1556923297] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7569332 kB) free!
[1556923417] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11457840 kB) free.
[1556923629] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1556924017] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7198480 kB) free!
[1556924138] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11446888 kB) free.
[1556924228] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.30, 0.46, 12.15
[1556924737] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7003088 kB) free!
[1556924857] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11460332 kB) free.
[1556925428] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1556925428] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.26, 0.23, 3.52
[1556925429] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1556925457] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6344976 kB) free!
[1556925549] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1556925577] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11463004 kB) free.
[1556926749] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1556926777] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7246096 kB) free!
[1556926869] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1556926897] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11460324 kB) free.
[1556927469] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1556927497] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6649732 kB) free!
[1556927589] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1556927617] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11468572 kB) free.
[1556928189] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556928309] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1556928817] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7439372 kB) free!
[1556928938] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11463076 kB) free.
[1556929509] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1556929537] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6786200 kB) free!
[1556929563] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 261.73, 220.69, 118.05
[1556929629] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1556929657] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11486432 kB) free.
[1556929683] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 50.65, 160.27, 108.91
[1556929803] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 8.64, 107.94, 95.96
[1556930229] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556930349] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556930469] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556931457] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8411508 kB) free!
[1556931577] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.8% (6824036 kB) free!
[1556931603] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.64, 1.90, 15.22
[1556931669] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1556931697] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10999980 kB) free.
[1556932270] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1556932297] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9113676 kB) free!
[1556932390] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1556932417] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7647844 kB) free!
[1556932510] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1556932537] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11472344 kB) free.
[1556933110] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556933228] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 65.31, 14.71, 5.06
[1556933230] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556933348] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 263.36, 107.22, 40.13
[1556933350] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1556933403] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.36, 1.54, 3.56
[1556933468] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 290.33, 169.20, 70.98
[1556933738] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (8001284 kB) free!
[1556933740] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 241.90, 148.70, 63.81
[1556933858] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.3% (6047288 kB) free!
[1556933860] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 244.13, 180.34, 85.69
[1556933978] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10567748 kB) free.
[1556933980] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 33.09, 120.76, 75.33
[1556934578] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9051688 kB) free!
[1556934698] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.1% (7425900 kB) free!
[1556934818] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11455700 kB) free.
[1556935268] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.46, 1.31, 18.24
[1556935418] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9785172 kB) free!
[1556935538] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8074876 kB) free!
[1556935658] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.1% (6448480 kB) free!
[1556935750] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1556935780] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.52, 0.71, 11.12
[1556936258] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10928228 kB) free.
[1556936350] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556936470] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1556936590] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1556936858] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8846980 kB) free!
[1556936978] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7169104 kB) free!
[1556936980] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.11, 0.21, 3.23
[1556937068] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.61, 0.45, 2.98
[1556937099] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11357728 kB) free.
[1556937698] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9806144 kB) free!
[1556937818] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.9% (8343344 kB) free!
[1556937938] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.6% (6695904 kB) free!
[1556938538] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1556938538] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10792852 kB) free.
[1556939138] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9141900 kB) free!
[1556939258] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7647216 kB) free!
[1556939378] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11327004 kB) free.
[1556939710] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47071 MB (20% inode=99%):
[1556939830] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47039 MB (20% inode=99%):
[1556939950] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47007 MB (20% inode=99%):
[1556939978] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9803880 kB) free!
[1556940098] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.0% (8353536 kB) free!
[1556940218] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.5% (6662036 kB) free!
[1556940818] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.1% (10384084 kB) free.
[1556941419] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8368060 kB) free!
[1556941539] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6827172 kB) free!
[1556941659] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11294520 kB) free.
[1556941948] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2442484 kB) free!
[1556942259] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9409900 kB) free!
[1556942379] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7930580 kB) free!
[1556942499] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.7% (6279832 kB) free!
[1556942590] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1556943099] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.1% (10390716 kB) free.
[1556943190] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556943311] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1556943431] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556943699] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8392824 kB) free!
[1556943787] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.7% (3204080 kB) free!
[1556943819] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.7% (6760748 kB) free!
[1556943907] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.2% (3128404 kB) free!
[1556943939] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11217180 kB) free.
[1556944027] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.8% (3056832 kB) free!
[1556944447] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.3% (14509020 kB) free.
[1556944499] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.9% (22042484 kB) free.
[1556944539] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (8997816 kB) free!
[1556944598] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.6% (20250036 kB) free.
[1556944660] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.9% (7357212 kB) free!
[1556944779] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.9% (5385036 kB) free!
[1556945231] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1556945498] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4900224 kB) free!
[1556945618] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (4828112 kB) free!
[1556945738] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (4745888 kB) free!
[1556945831] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556945951] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1556946072] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1556946672] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1556947179] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10463712 kB) free.
[1556947268] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.99, 14.96, 8.83
[1556947272] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1556947388] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.47, 10.05, 7.77
[1556947392] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556947509] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.39, 6.87, 6.89
[1556947512] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1556947779] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8260352 kB) free!
[1556947899] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.4% (6620824 kB) free!
[1556948019] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.3% (10990900 kB) free.
[1556948109] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.78, 1.40, 3.84
[1556948619] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9232788 kB) free!
[1556948657] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (993784 kB) free!
[1556948739] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7650720 kB) free!
[1556948777] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (179208 kB) free!
[1556948859] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.0% (5916364 kB) free!
[1556948897] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156652 kB) free!
[1556949406] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (832096 kB) free!
[1556949459] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10039220 kB) free.
[1556949497] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1556949497] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 74.5% (12112156 kB) free.
[1556949526] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (923476 kB) free!
[1556949550] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 70462 MB (31% inode=99%):
[1556949646] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (864244 kB) free!
[1556949938] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (1410420 kB) free!
[1556950060] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8199488 kB) free!
[1556950180] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.3% (6573852 kB) free!
[1556950204] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 367.87, 336.13, 184.64
[1556950300] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10870740 kB) free.
[1556950324] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 50.84, 225.38, 162.43
[1556950444] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 8.40, 151.43, 142.96
[1556950539] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.0% (3177848 kB) free!
[1556950846] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 13.0% (3188904 kB) free!
[1556950900] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9056176 kB) free!
[1556951020] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.2% (7479772 kB) free!
[1556951140] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.5% (5650788 kB) free!
[1556951740] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9953940 kB) free.
[1556952340] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7865760 kB) free!
[1556952460] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.7% (6235600 kB) free!
[1556952580] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10736004 kB) free.
[1556952909] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 274.73, 195.86, 90.96
[1556953029] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 66.95, 154.11, 89.28
[1556953149] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 9.69, 103.34, 78.53
[1556953180] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8711444 kB) free!
[1556953246] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1556953246] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1564360 kB) free!
[1556953300] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7149648 kB) free!
[1556953421] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.1% (5482856 kB) free!
[1556953512] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1556953539] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1556953539] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2075380 kB) free!
[1556953780] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 11.28, 98.93, 72.45
[1556953900] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.90, 66.35, 63.72
[1556953967] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 70.7% (17322672 kB) free.
[1556954020] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.52, 44.49, 56.03
[1556954021] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9951500 kB) free.
[1556954022] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.5% (22428660 kB) free.
[1556954044] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.75, 1.93, 15.20
[1556954112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1556954232] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556954353] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556954621] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8193268 kB) free!
[1556954741] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.2% (6507944 kB) free!
[1556954861] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10554688 kB) free.
[1556954949] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.41, 0.71, 11.75
[1556955220] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.45, 1.18, 15.63
[1556955461] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9116404 kB) free!
[1556955581] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7537192 kB) free!
[1556955641] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1556955701] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.0% (5901952 kB) free!
[1556955761] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1556956149] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.53, 0.49, 3.60
[1556956276] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.60, 5.77, 5.30
[1556956301] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (10007796 kB) free.
[1556956396] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.70, 5.29, 5.18
[1556956420] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.35, 0.38, 4.56
[1556956444] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.74, 1.90, 4.12
[1556956516] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.68, 5.02, 5.09
[1556956901] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7910252 kB) free!
[1556957021] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.7% (6240144 kB) free!
[1556957117] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.07, 4.19, 4.54
[1556957141] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.1% (10402248 kB) free.
[1556957297] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 15.7% (2558452 kB) free!
[1556957417] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 14.6% (2369412 kB) free!
[1556957537] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 13.5% (2191112 kB) free!
[1556957742] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8822496 kB) free!
[1556957862] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7232264 kB) free!
[1556957982] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.2% (5513840 kB) free!
[1556958137] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1404176 kB) free!
[1556958553] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1556958582] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10309536 kB) free.
[1556959154] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556959182] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8189936 kB) free!
[1556959274] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1556959302] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.2% (6516176 kB) free!
[1556959394] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556959422] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4588096 kB) free!
[1556959978] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (298500 kB) free!
[1556960022] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.5% (9101904 kB) free!
[1556960098] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183256 kB) free!
[1556960218] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164772 kB) free!
[1556960409] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1556961009] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1556961161] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1556961281] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1556961418] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 20.7% (5077660 kB) free.
[1556961609] SERVICE ALERT: sservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1556961794] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1556961822] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9930140 kB) free.
[1556961961] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.2% (14660532 kB) free.
[1556962018] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.1% (2731828 kB) free!
[1556962139] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.7% (2630124 kB) free!
[1556962259] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 21.5% (5265824 kB) free.
[1556962394] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1556962422] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7823248 kB) free!
[1556962515] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1556962543] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.9% (6343936 kB) free!
[1556962635] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1556962663] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4539024 kB) free!
[1556962859] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (4214692 kB) free!
[1556962979] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.5% (2827372 kB) free!
[1556963099] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1277656 kB) free!
[1556963263] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.2% (8949252 kB) free!
[1556964299] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 27.8% (6819236 kB) free.
[1556964827] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1556964899] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.1% (3695796 kB) free!
[1556965019] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.7% (2874888 kB) free!
[1556965035] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1556965139] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1822744 kB) free!
[1556965329] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.7% (4611568 kB) free!
[1556965449] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.6% (4588564 kB) free!
[1556965569] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.5% (4566232 kB) free!
[1556965635] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1556965755] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556965876] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1556966169] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182348 kB) free!
[1556966476] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1556966626] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (200856 kB) free!
[1556966708] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170268 kB) free!
[1556966746] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198780 kB) free!
[1556966828] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196564 kB) free!
[1556966859] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.7% (2881388 kB) free!
[1556966863] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4668140 kB) free!
[1556966866] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (183948 kB) free!
[1556966940] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 68.6% (16817700 kB) free.
[1556966948] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175988 kB) free!
[1556966949] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1556966979] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.0% (2756784 kB) free!
[1556967076] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1556967099] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.4% (2668032 kB) free!
[1556967196] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556967316] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1556967463] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.2% (9470080 kB) free!
[1556967609] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1556967609] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1556967916] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1556968516] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1556968636] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1556968740] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (900928 kB) free!
[1556968757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1556968860] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (412588 kB) free!
[1556968899] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1395540 kB) free!
[1556968980] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (718688 kB) free!
[1556969167] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1556971063] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4622472 kB) free!
[1556971150] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46926 MB (20% inode=99%):
[1556971157] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1556971270] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46863 MB (20% inode=99%):
[1556971390] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46750 MB (20% inode=99%):
[1556971420] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 273.18, 209.69, 102.45
[1556971540] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 40.18, 143.47, 91.36
[1556971660] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 5.79, 96.13, 80.33
[1556971663] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.2% (8980064 kB) free!
[1556971748] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2718016 kB) free!
[1556971757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556971877] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1556971997] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1556972949] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 265.66, 156.78, 66.20
[1556973069] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 269.17, 194.18, 90.89
[1556973099] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 87.4% (14203292 kB) free.
[1556973181] SERVICE ALERT: sig1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1556973184] HOST ALERT: sig1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.44)
[1556973189] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 60.30, 149.11, 87.69
[1556973246] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig1' was out of bounds)
[1556973258] SERVICE ALERT: sig1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig1' was out of bounds)
[1556973258] SERVICE ALERT: sig1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig1' was out of bounds)
[1556973258] SERVICE ALERT: sig1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sig1' was out of bounds)
[1556973298] SERVICE ALERT: sig1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig1' was out of bounds)
[1556973304] SERVICE ALERT: sig1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig1' was out of bounds)
[1556973304] SERVICE ALERT: sig1;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig1' was out of bounds)
[1556973369] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.5% (2842816 kB) free!
[1556973460] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.12, 0.49, 11.78
[1556973548] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2237860 kB) free!
[1556973609] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1556973612] SERVICE ALERT: sig1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig1' was out of bounds)
[1556973745] HOST ALERT: sig1;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.44)
[1556973967] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.1% (2949932 kB) free!
[1556974087] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.8% (2890536 kB) free!
[1556974208] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.4% (2825504 kB) free!
[1556974397] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1556974989] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.67, 0.92, 13.09
[1556974997] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1556975117] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1556975169] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1556975169] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2417304 kB) free!
[1556975237] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1556975260] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.19, 0.35, 2.89
[1556975263] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3794204 kB) free!
[1556975863] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.4% (8551416 kB) free!
[1556976189] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.75, 0.67, 4.04
[1556976449] HOST ALERT: sig1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1556976608] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (1370808 kB) free!
[1556977048] HOST ALERT: sig1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.44)
[1556977238] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1556978169] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (2542724 kB) free!
[1556978769] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2353500 kB) free!
[1556979180] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1556979463] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3946344 kB) free!
[1556980063] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.6% (8197568 kB) free!
[1556981263] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4437984 kB) free!
[1556981863] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.3% (8503848 kB) free!
[1556983063] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4545024 kB) free!
[1556983637] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1556983663] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.0% (8878700 kB) free!
[1556984237] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1556984357] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556984477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1556985180] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.2% (2489492 kB) free!
[1556986877] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1556987263] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4141392 kB) free!
[1556987477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1556987597] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556987717] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1556987863] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.2% (8471632 kB) free!
[1556990044] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 254.15, 104.25, 39.85
[1556990118] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1556990164] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 279.75, 163.63, 69.42
[1556990169] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1556990284] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 283.07, 203.35, 95.40
[1556990718] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1556990838] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556990958] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1556992684] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.54, 1.64, 12.02
[1556993263] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4487804 kB) free!
[1556993359] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1556993863] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.8% (8780388 kB) free!
[1556993884] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.38, 1.52, 4.43
[1556993959] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1556994079] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1556994199] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes