[1567968013] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176212 kB) free!
[1567968125] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (875884 kB) free!
[1567968176] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.5% (3064428 kB) free!
[1567968245] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (870980 kB) free!
[1567968296] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.5% (3057160 kB) free!
[1567968365] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187432 kB) free!
[1567968417] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.5% (3055332 kB) free!
[1567972426] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1567972546] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1567972687] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (2253348 kB) free!
[1567973270] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.04, 6.04, 5.62
[1567973390] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.84, 6.04, 5.67
[1567973510] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.32, 5.37, 5.47
[1567974710] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.89, 3.56, 4.56
[1567975565] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 21.2% (5219352 kB) free.
[1567975636] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 10.9% (5351168 kB) free!
[1567975757] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (239572 kB) free!
[1567975877] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240288 kB) free!
[1567976165] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181020 kB) free!
[1567976285] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196124 kB) free!
[1567976405] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198376 kB) free!
[1567977077] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.8% (43769184 kB) free.
[1567978617] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180664 kB) free!
[1567978813] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 46.2% (11324012 kB) free.
[1567979217] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 48.8% (11969988 kB) free.
[1567979259] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1567979379] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1567991405] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 50.5% (12459544 kB) free.
[1567994406] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180816 kB) free!
[1567994413] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (599488 kB) free!
[1567994526] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187932 kB) free!
[1567994533] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (564724 kB) free!
[1567994646] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (208728 kB) free!
[1567994653] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (520684 kB) free!
[1567998579] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 327 processes
[1567998699] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1567998946] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1567999017] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4761268 kB) free!
[1567999066] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1567999137] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4700716 kB) free!
[1567999257] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.9% (4634176 kB) free!
[1567999819] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 22.06, 5.47, 2.48
[1567999940] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 14.07, 13.13, 6.10
[1568000060] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.67, 9.07, 5.46
[1568000660] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.03, 2.26, 3.46
[1568003502] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 138.35, 79.32, 33.55
[1568003622] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 140.45, 99.67, 46.55
[1568003742] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 140.64, 113.19, 57.95
[1568005460] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 132.99, 102.77, 50.89
[1568005542] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.44, 1.50, 12.56
[1568005580] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 75.08, 98.32, 56.01
[1568005700] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 11.36, 66.21, 49.37
[1568006225] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.7% (23377932 kB) free.
[1568006278] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.7% (23220952 kB) free.
[1568006345] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.4% (40613700 kB) free.
[1568006508] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1568006629] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1568006742] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.93, 0.96, 4.11
[1568006749] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1568006900] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1568006900] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.23, 2.88, 14.78
[1568007057] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2298252 kB) free!
[1568008636] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.0% (2726356 kB) free!
[1568008700] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.60, 1.94, 3.69
[1568008756] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (2350600 kB) free!
[1568008876] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1660252 kB) free!
[1568009499] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1568009619] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1568013557] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.0% (8872024 kB) free!
[1568013677] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.4% (8592356 kB) free!
[1568013797] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.4% (8103932 kB) free!
[1568014910] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.19, 8.73, 6.63
[1568015030] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.96, 9.45, 7.13
[1568015150] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.74, 9.55, 7.46
[1568015750] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.15, 2.42, 4.63
[1568016197] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (3370264 kB) free!
[1568016464] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.8% (2413420 kB) free!
[1568017278] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8776584 kB) free!
[1568017398] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.2% (8490860 kB) free!
[1568017518] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.6% (8200500 kB) free!
[1568018150] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 17.79, 15.56, 9.97
[1568018270] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 17.27, 16.31, 10.93
[1568018283] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196328 kB) free!
[1568018390] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 17.41, 16.96, 11.84
[1568018403] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187068 kB) free!
[1568018523] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190728 kB) free!
[1568018900] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1568019057] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 44.8% (10984488 kB) free.
[1568019123] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 54.7% (13401848 kB) free.
[1568019318] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (3975788 kB) free!
[1568019950] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1568020791] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1568020791] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.50, 3.25, 4.37
[1568021037] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 572 processes
[1568021097] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 573 processes
[1568021391] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.89, 10.23, 6.88
[1568021511] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.98, 7.80, 6.42
[1568021631] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.84, 5.43, 5.72
[1568022231] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 6.45, 3.85, 4.30
[1568024019] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1568024031] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.85, 9.31, 7.29
[1568024139] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1568024151] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.26, 9.21, 7.50
[1568024271] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.97, 8.79, 7.56
[1568025657] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (4767888 kB) free!
[1568025777] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.8% (4615356 kB) free!
[1568025897] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (4444208 kB) free!
[1568026323] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.0% (4160512 kB) free!
[1568026443] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.2% (3959524 kB) free!
[1568026563] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.3% (3757860 kB) free!
[1568026666] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.71, 5.49, 5.19
[1568026786] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.33, 5.85, 5.36
[1568026906] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.31, 5.66, 5.35
[1568027506] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.59, 3.39, 4.40
[1568028297] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2107772 kB) free!
[1568028363] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1568028363] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1799068 kB) free!
[1568029550] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1568029670] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1568029790] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1568031471] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.24, 3.15, 4.08
[1568034339] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1568034459] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1568035906] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.85, 5.71, 5.25
[1568036026] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.96, 5.94, 5.40
[1568036097] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.5% (3069112 kB) free!
[1568036146] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.36, 5.75, 5.40
[1568036271] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.47, 8.51, 6.05
[1568036391] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.27, 8.23, 6.23
[1568036511] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 12.09, 9.49, 6.92
[1568036697] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1568036697] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1765936 kB) free!
[1568036746] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1568036746] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.08, 4.38, 4.76
[1568037823] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.33, 6.82, 5.49
[1568037944] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.56, 6.33, 5.48
[1568037946] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.37, 5.39, 5.09
[1568038064] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.06, 5.84, 5.40
[1568038066] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.46, 5.14, 5.04
[1568038186] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.14, 4.89, 4.96
[1568038563] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1568038664] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 3.67, 3.56, 4.32
[1568041311] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.27, 2.99, 4.59
[1568042390] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1568043343] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 68.43, 49.01, 23.21
[1568043463] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 68.31, 55.42, 28.69
[1568043500] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 25.54, 6.79, 3.32
[1568043583] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 68.78, 59.87, 33.56
[1568043620] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 79.19, 33.49, 13.52
[1568043740] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 86.70, 51.47, 22.53
[1568044183] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.23, 9.25, 18.47
[1568044940] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.14, 4.21, 14.29
[1568045984] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.01, 1.24, 3.70
[1568046099] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 14 users currently logged in
[1568046140] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.94, 1.03, 4.65
[1568048097] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1568048386] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1568048511] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1568051017] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 294 processes
[1568051137] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1568052466] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1568052586] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 402 processes
[1568052706] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 402 processes
[1568053307] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes