[1574361015] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 214469 MB (20% inode=99%):
[1574361239] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 93.0% (45842612 kB) free.
[1574361282] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.1% (2954276 kB) free!
[1574361422] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1574361544] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 374 processes
[1574361603] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184060 kB) free!
[1574361723] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182500 kB) free!
[1574361726] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193600 kB) free!
[1574361844] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181644 kB) free!
[1574361846] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174440 kB) free!
[1574361883] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1574361883] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168972 kB) free!
[1574361966] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175752 kB) free!
[1574362180] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(48) < 2019/11/22 00:19:14 [error] 24132#0: *8379559 recv() failed (104: Connection reset by peer) while proxying connection, client: 54.254.180.107, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:50/0, bytes from/to upstream:0/50
[1574362300] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574362792] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1574362839] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1574362900] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(41) < 2019/11/22 00:30:52 [error] 24131#0: *8381722 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574362912] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1574362960] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1574363021] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574364820] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(50) < 2019/11/22 00:55:31 [error] 24130#0: *8386338 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574364874] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 24.1% (5939156 kB) free.
[1574364940] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(21) < 2019/11/22 01:05:39 [error] 24130#0: *8388275 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574365060] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(5) < 2019/11/22 01:05:55 [error] 24130#0: *8388370 connect() failed (111: Connection refused) while connecting to upstream, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574365474] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167976 kB) free!
[1574365594] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187056 kB) free!
[1574365660] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574365674] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.4% (13700612 kB) free!
[1574365714] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192688 kB) free!
[1574366465] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.9% (7331400 kB) free!
[1574366585] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 12.7% (6278268 kB) free!
[1574366705] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (5508824 kB) free!
[1574366761] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/11/22 01:31:24 [error] 479#0: *7909238 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:14001, upstream: "10.147.212.38:5001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574366860] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(48) < 2019/11/22 01:30:40 [error] 24130#0: *8393047 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574366881] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574366937] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 103.255.146.164 [22/Nov/2019:01:31:24 +0530] TCP 502 0 0 0.004 "10.147.212.37:5001, 10.147.212.49:5001, 10.147.212.38:5001" "0, 0, 0" "0, 0, 0" "0.003, 0.000, 0.001"
[1574366980] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574367058] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574367305] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (1070624 kB) free!
[1574367580] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(51) < 2019/11/22 01:42:30 [error] 24130#0: *8395163 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574367700] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574367713] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1574367833] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1574368300] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2019/11/22 02:00:53 [error] 24123#0: *8398712 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574368420] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574368646] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 368 processes
[1574368767] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 372 processes
[1574369021] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(39) < 2019/11/22 02:07:44 [error] 24130#0: *8399924 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574369141] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574369741] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(33) < 2019/11/22 02:17:03 [error] 24126#0: *8401736 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574369744] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (525912 kB) free!
[1574369755] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (2674140 kB) free!
[1574369861] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574369864] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (521244 kB) free!
[1574369875] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (216564 kB) free!
[1574369984] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 76.9% (37889036 kB) free.
[1574369995] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (211936 kB) free!
[1574370595] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 71.0% (34973840 kB) free.
[1574370681] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1574370759] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1574370801] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 374 processes
[1574370879] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1574371061] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(45) < 2019/11/22 02:42:13 [error] 24134#0: *8406460 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574371114] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 48.2% (11884244 kB) free.
[1574371181] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574371782] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(30) < 2019/11/22 02:53:00 [error] 24130#0: *8408496 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574371902] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574373102] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(40) < 2019/11/22 03:19:33 [error] 24130#0: *8413543 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574373222] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574373283] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574373822] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(30) < 2019/11/22 03:27:38 [error] 24130#0: *8415105 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574373833] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1574373942] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574373953] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1574375142] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(21) < 2019/11/22 03:54:49 [error] 24130#0: *8420244 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574375262] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574375862] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(35) < 2019/11/22 04:03:50 [error] 24130#0: *8421907 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574375982] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574377714] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1574377783] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(64) < 2019/11/22 04:34:34 [error] 24130#0: *8427765 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.218.150, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574377903] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(21) < 2019/11/22 04:39:58 [error] 24130#0: *8428753 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1574378023] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574379279] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1574379399] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1574379823] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(62) < 2019/11/22 05:10:25 [error] 24130#0: *8434517 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1574379943] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(27) < 2019/11/22 05:14:38 [error] 24130#0: *8435312 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574380063] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574381863] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(67) < 2019/11/22 05:46:28 [error] 24130#0: *8441353 connect() failed (111: Connection refused) while connecting to upstream, client: 185.255.8.58, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574381983] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574382583] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(25) < 2019/11/22 05:51:50 [error] 24130#0: *8442339 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574382703] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574383114] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169752 kB) free!
[1574383234] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195840 kB) free!
[1574383355] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180548 kB) free!
[1574383566] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.8% (22017904 kB) free.
[1574383675] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 37.5% (49429776 kB) free.
[1574383815] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 243013 MB (22% inode=99%):
[1574383904] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(59) < 2019/11/22 06:21:22 [error] 24130#0: *8447909 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574384024] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574384044] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 86.6% (21236932 kB) free.
[1574384083] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 65.2% (15992856 kB) free.
[1574384624] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(28) < 2019/11/22 06:27:17 [error] 24130#0: *8448994 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574384744] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574385944] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(46) < 2019/11/22 06:54:32 [error] 24130#0: *8454111 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574386064] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574386664] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(45) < 2019/11/22 07:03:40 [error] 24130#0: *8455854 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574386785] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574387799] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1574387919] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1574387984] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(22) < 2019/11/22 07:24:40 [error] 24130#0: *8459793 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574388104] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(21) < 2019/11/22 07:30:30 [error] 24130#0: *8460911 connect() failed (111: Connection refused) while connecting to upstream, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574388224] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(22) < 2019/11/22 07:33:25 [error] 24132#0: *8461465 connect() failed (111: Connection refused) while connecting to upstream, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574389424] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574390024] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(20) < 2019/11/22 07:59:08 [error] 24132#0: *8466276 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574390145] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574390744] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(70) < 2019/11/22 08:14:36 [error] 24130#0: *8469202 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574390864] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574392064] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(23) < 2019/11/22 08:35:12 [error] 24130#0: *8473093 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.218.150, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:38/0, bytes from/to upstream:0/38
[1574392184] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574392784] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(46) < 2019/11/22 08:42:07 [error] 24130#0: *8474389 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574392904] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574393505] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(26) < 2019/11/22 08:51:53 [error] 24130#0: *8476231 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574393625] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574394225] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(23) < 2019/11/22 09:11:26 [error] 24129#0: *8479901 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574394345] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(25) < 2019/11/22 09:15:37 [error] 24132#0: *8480737 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1574394465] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(30) < 2019/11/22 09:16:28 [error] 24126#0: *8480880 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574395665] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1574395665] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574396266] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(25) < 2019/11/22 09:46:45 [error] 24130#0: *8486588 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.218.150, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574396386] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574396986] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(49) < 2019/11/22 09:52:08 [error] 24133#0: *8487550 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574397106] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(24) < 2019/11/22 10:01:42 [error] 24131#0: *8489434 connect() failed (111: Connection refused) while connecting to upstream, client: 185.255.8.58, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574397226] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574398426] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(21) < 2019/11/22 10:22:08 [error] 24130#0: *8493271 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574398546] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574399147] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(27) < 2019/11/22 10:27:13 [error] 24130#0: *8494278 connect() failed (111: Connection refused) while connecting to upstream, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574399267] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(27) < 2019/11/22 10:36:25 [error] 24134#0: *8496050 connect() failed (111: Connection refused) while connecting to upstream, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574399387] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574399820] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.0% (22711356 kB) free.
[1574399935] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.7% (39789336 kB) free.
[1574399971] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.1% (14475880 kB) free.
[1574400587] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(25) < 2019/11/22 10:57:12 [error] 24130#0: *8499999 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574400707] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(28) < 2019/11/22 11:00:27 [error] 24131#0: *8500613 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1574400827] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(23) < 2019/11/22 11:03:41 [error] 24131#0: *8501269 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574401938] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < webmachine error: path="/api/bindings/%2f/e/amq.direct/q/1165_n"
[1574402027] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574402060] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1574402627] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(26) < 2019/11/22 11:31:20 [error] 24130#0: *8506591 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574402747] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(35) < 2019/11/22 11:33:57 [error] 24130#0: *8507080 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574402824] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181248 kB) free!
[1574402867] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574402944] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178180 kB) free!
[1574403064] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187796 kB) free!
[1574403467] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(56) < 2019/11/22 11:45:08 [error] 24130#0: *8509223 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1574403587] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574404087] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 80.4% (18059064 kB) free.
[1574404682] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1574404787] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(40) < 2019/11/22 12:08:31 [error] 24130#0: *8513727 connect() failed (111: Connection refused) while connecting to upstream, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574404861] SERVICE ALERT: sig6-vm2;OutOfMemory-killer--oom-killer;CRITICAL;SOFT;1;(1) < Nov 21 02:59:17 sig6-vm2 kernel: Killed process 5604 (test-dec) total-vm:3157740kB, anon-rss:2283604kB, file-rss:0kB, shmem-rss:0kB
[1574404907] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574404983] SERVICE ALERT: sig6-vm2;OutOfMemory-killer--oom-killer;OK;HARD;1;Log check ok - 0 pattern matches found
[1574405507] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(58) < 2019/11/22 12:18:49 [error] 24133#0: *8515704 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574405627] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574406227] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(6) < 2019/11/22 12:31:07 [error] 24131#0: *8518148 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574406283] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4833336 kB) free!
[1574406348] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574406403] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4698040 kB) free!
[1574406523] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.7% (4576080 kB) free!
[1574406947] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(43) < 2019/11/22 12:43:21 [error] 24130#0: *8520504 connect() failed (111: Connection refused) while connecting to upstream, client: 185.255.8.58, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574407067] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574407667] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(51) < 2019/11/22 12:52:44 [error] 24132#0: *8522291 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574407787] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574407795] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9488732 kB) free!
[1574407800] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/22 12:54:09 [error] 484#0: *8144113 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /favicon.ico HTTP/1.1", host: "59.91.63.148:8091", referrer: "http://59.91.63.148:8091/"
[1574407915] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.8% (9270504 kB) free!
[1574407921] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574408035] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.3% (9003136 kB) free!
[1574408842] SERVICE ALERT: db1-vm2;HDD_Home status;CRITICAL;SOFT;1;DISK CRITICAL - /dev/centos-home is not accessible: No such file or directory
[1574408901] SERVICE ALERT: db1-vm2;HDD_Root status;CRITICAL;SOFT;1;DISK CRITICAL - /dev/centos-root is not accessible: No such file or directory
[1574408923] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1848672 kB) free!
[1574408962] SERVICE ALERT: db1-vm2;HDD_Home status;CRITICAL;SOFT;2;DISK CRITICAL - /dev/centos-home is not accessible: No such file or directory
[1574408987] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(59) < 2019/11/22 13:17:29 [error] 24133#0: *8527088 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574409018] SERVICE ALERT: db1-vm2;Mariadb status;CRITICAL;SOFT;1;NRPE: Command 'check_mysql' not defined
[1574409021] SERVICE ALERT: db1-vm2;HDD_Root status;CRITICAL;SOFT;2;DISK CRITICAL - /dev/centos-root is not accessible: No such file or directory
[1574409077] SERVICE ALERT: db1-vm2;Memory;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1574409082] SERVICE ALERT: db1-vm2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - /dev/centos-home is not accessible: No such file or directory
[1574409107] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574409138] SERVICE ALERT: db1-vm2;Mariadb status;CRITICAL;SOFT;2;NRPE: Command 'check_mysql' not defined
[1574409141] SERVICE ALERT: db1-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 43580 MB (85% inode=99%):
[1574409195] SERVICE ALERT: db1-vm2;PING;CRITICAL;SOFT;1;NRPE: Command 'check_ping' not defined
[1574409197] SERVICE ALERT: db1-vm2;Memory;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1574409258] SERVICE ALERT: db1-vm2;Mariadb status;CRITICAL;HARD;3;NRPE: Command 'check_mysql' not defined
[1574409315] SERVICE ALERT: db1-vm2;PING;CRITICAL;SOFT;2;NRPE: Command 'check_ping' not defined
[1574409317] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 84.3% (20813776 kB) free.
[1574409395] SERVICE ALERT: db1-vm2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 26739 MB (66% inode=99%):
[1574409410] SERVICE ALERT: db1-vm2;PING;OK;HARD;1;PING OK - Packet loss = 0%, RTA = 0.07 ms
[1574409707] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(43) < 2019/11/22 13:28:00 [error] 24131#0: *8529076 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574409827] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574409844] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (4515392 kB) free!
[1574409964] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4416492 kB) free!
[1574410084] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (4283768 kB) free!
[1574410435] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4521600 kB) free!
[1574411027] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(43) < 2019/11/22 13:52:51 [error] 24130#0: *8533931 connect() failed (111: Connection refused) while connecting to upstream, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574411148] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574411747] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(44) < 2019/11/22 14:03:19 [error] 24130#0: *8536027 connect() failed (111: Connection refused) while connecting to upstream, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574411867] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574412325] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25507552 kB) free!
[1574412445] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.3% (45230384 kB) free.
[1574412484] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2455340 kB) free!
[1574412720] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/11/22 14:14:20 [error] 478#0: *8177439 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /favicon.ico HTTP/1.1", host: "59.91.63.148:8091", referrer: "http://59.91.63.148:8091/httpsms/bulksms.php?key=1165&username=nanital&password=iV4M03M4&source_addr=BP-TESTMS&message=HTTP%20TEST%20MSG%20FROM%20VNS&msisdn=919533213151"
[1574412840] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574413067] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(42) < 2019/11/22 14:27:11 [error] 24134#0: *8540569 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574413187] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574413440] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/11/22 14:25:36 [error] 480#0: *8182549 open() "/usr/share/nginx/html/httpsms/bulksms.php" failed (2: No such file or directory), client: 103.255.146.164, server: 148qservercluster1.com, request: "GET /httpsms/bulksms.php?key=1165 HTTP/1.1", host: "59.91.63.148:8091"
[1574413560] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2) < 2019/11/22 14:35:22 [error] 26542#0: *8187239 open() "/usr/share/nginx/html/httpsms" failed (2: No such file or directory), client: 103.255.146.164, server: 148qservercluster1.com, request: "GET /httpsms HTTP/1.1", host: "59.91.63.148:8091"
[1574413645] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22197128 kB) free!
[1574413680] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574413720] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1574413766] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.4% (44030232 kB) free.
[1574413787] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(55) < 2019/11/22 14:39:07 [error] 24130#0: *8542927 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574413840] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1574413908] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574414166] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4676204 kB) free!
[1574414280] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(18) < 2019/11/22 14:44:36 [error] 26838#0: *8191257 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 59.91.63.148, request: "GET /favicon.ico HTTP/1.1", upstream: "http://10.147.212.142:80/favicon.ico", host: "59.91.63.148", referrer: "http://59.91.63.148/httpsms/"
[1574414286] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4515760 kB) free!
[1574414365] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20277508 kB) free!
[1574414401] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2) < 2019/11/22 14:48:57 [error] 26847#0: *8193223 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 59.91.63.148, request: "GET /favicon.ico HTTP/1.1", upstream: "http://10.147.212.142:80/favicon.ico", host: "59.91.63.148", referrer: "http://59.91.63.148/httpsms/"
[1574414406] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.0% (4407600 kB) free!
[1574414412] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 378 processes
[1574414485] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40756756 kB) free.
[1574414521] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;HARD;3;(2) < 2019/11/22 14:52:00 [error] 26849#0: *8194661 connect() failed (111: Connection refused) while connecting to upstream, client: 183.82.105.122, server: 59.91.63.148, request: "GET /favicon.ico HTTP/1.1", upstream: "http://10.147.212.142:80/favicon.ico", host: "59.91.63.148", referrer: "http://59.91.63.148/httpsms/"
[1574414532] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 378 processes
[1574415108] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(29) < 2019/11/22 14:57:23 [error] 24131#0: *8546429 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574415228] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(22) < 2019/11/22 15:03:24 [error] 24129#0: *8547610 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574415348] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(4) < 2019/11/22 15:04:33 [error] 24130#0: *8547883 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574415361] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.1% (19758264 kB) free.
[1574415420] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 34.4% (16947288 kB) free.
[1574415462] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 73.0% (17898220 kB) free.
[1574415484] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 35.7% (8738672 kB) free.
[1574415685] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19574976 kB) free!
[1574415721] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574415805] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40929772 kB) free.
[1574416142] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9851192 kB) free!
[1574416206] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2395584 kB) free!
[1574416262] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (9850180 kB) free!
[1574416382] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 20.0% (9850516 kB) free!
[1574416405] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17243904 kB) free!
[1574416525] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39711948 kB) free.
[1574416549] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574417125] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14780740 kB) free!
[1574417149] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(27) < 2019/11/22 15:32:42 [error] 24129#0: *8553303 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.218.150, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574417165] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (4418416 kB) free!
[1574417245] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38100488 kB) free.
[1574417269] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(32) < 2019/11/22 15:37:48 [error] 24129#0: *8554281 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574417285] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.8% (4384304 kB) free!
[1574417341] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46995 MB (20% inode=99%):
[1574417389] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(3) < 2019/11/22 15:37:53 [error] 24131#0: *8554274 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574417405] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.4% (4284684 kB) free!
[1574417461] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46995 MB (20% inode=99%):
[1574417581] SERVICE ALERT: primary-server;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46995 MB (20% inode=99%):
[1574417845] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13630220 kB) free!
[1574417965] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36366972 kB) free.
[1574418565] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12079160 kB) free!
[1574418584] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9595432 kB) free!
[1574418685] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35224472 kB) free.
[1574418704] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (9463572 kB) free!
[1574418781] SERVICE ALERT: primary-server;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47215 MB (21% inode=99%):
[1574418824] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (9247180 kB) free!
[1574419153] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 80.1% (19637448 kB) free.
[1574419205] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171060 kB) free!
[1574419285] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10775264 kB) free!
[1574419406] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34419360 kB) free.
[1574419870] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(69) < 169.38.76.212 [22/Nov/2019:16:12:47 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1574419981] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47164 MB (20% inode=99%):
[1574419990] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574420005] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10690328 kB) free!
[1574420101] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47164 MB (20% inode=99%):
[1574420125] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31125784 kB) free.
[1574420221] SERVICE ALERT: primary-server;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47164 MB (20% inode=99%):
[1574420389] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1574420389] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574420583] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (212968 kB) free!
[1574420725] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6456812 kB) free!
[1574420821] SERVICE FLAPPING ALERT: primary-server;HDD_Home status;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1574420821] SERVICE ALERT: primary-server;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 65867 MB (29% inode=99%):
[1574420845] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.1% (42252372 kB) free.
[1574421010] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 353 processes
[1574421130] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 353 processes
[1574421250] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1574421445] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (19030476 kB) free!
[1574421565] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.1% (42272204 kB) free.
[1574421850] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 356 processes
[1574421970] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1574422090] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1574422127] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170004 kB) free!
[1574422165] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19246096 kB) free!
[1574422189] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(94) < 2019/11/22 16:51:06 [error] 24132#0: *8568343 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574422247] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167132 kB) free!
[1574422285] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (42924848 kB) free.
[1574422309] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574422367] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166100 kB) free!
[1574422424] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4595236 kB) free!
[1574422690] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 358 processes
[1574422810] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 357 processes
[1574422885] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17412844 kB) free!
[1574422930] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 359 processes
[1574422967] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 23.5% (5756616 kB) free.
[1574423005] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41564832 kB) free.
[1574423605] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17206768 kB) free!
[1574423725] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (13232004 kB) free!
[1574423845] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38093980 kB) free.
[1574424110] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(100) < 2019/11/22 17:27:17 [error] 24134#0: *8575423 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574424167] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183004 kB) free!
[1574424231] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574424287] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191612 kB) free!
[1574424407] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191468 kB) free!
[1574424446] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13784312 kB) free!
[1574424566] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38636876 kB) free.
[1574424730] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1574425092] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4414376 kB) free!
[1574425166] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13468664 kB) free!
[1574425212] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4137460 kB) free!
[1574425286] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (9558688 kB) free!
[1574425330] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1574425332] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.3% (3886300 kB) free!
[1574425383] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.7% (39790844 kB) free.
[1574425406] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34730264 kB) free.
[1574425417] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1574425417] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 90.5% (22183092 kB) free.
[1574425450] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 358 processes
[1574425458] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.8% (22897100 kB) free.
[1574425544] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 81.9% (40339464 kB) free.
[1574425570] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 357 processes
[1574425738] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 85.8% (19262308 kB) free.
[1574426006] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10312472 kB) free!
[1574426030] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(83) < 2019/11/22 18:02:45 [error] 24130#0: *8582277 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574426126] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35859124 kB) free.
[1574426150] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574426223] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9089148 kB) free!
[1574426265] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.5% (4300680 kB) free!
[1574426343] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.8% (8781200 kB) free!
[1574426385] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.9% (4150012 kB) free!
[1574426463] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.2% (8462844 kB) free!
[1574426505] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.3% (3993456 kB) free!
[1574426664] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.3% (3778756 kB) free!
[1574426726] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (5048140 kB) free!
[1574426770] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1574426770] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1574426784] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174660 kB) free!
[1574426846] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1207612 kB) free!
[1574426904] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1574426904] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185332 kB) free!
[1574426966] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29342308 kB) free.
[1574427566] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (771720 kB) free!
[1574427686] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30697204 kB) free.
[1574427950] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(22) < 2019/11/22 18:34:03 [error] 24130#0: *8588423 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574427971] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 350 processes
[1574428070] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(55) < 2019/11/22 18:37:18 [error] 24132#0: *8588971 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574428091] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 351 processes
[1574428190] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(28) < 2019/11/22 18:38:08 [error] 24132#0: *8589140 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574428211] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 351 processes
[1574428263] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4563648 kB) free!
[1574428286] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5544708 kB) free!
[1574428305] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1588604 kB) free!
[1574428406] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1716404 kB) free!
[1574428424] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170512 kB) free!
[1574428526] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34110528 kB) free.
[1574428544] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168232 kB) free!
[1574428664] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (205772 kB) free!
[1574428811] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1574429127] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (3156568 kB) free!
[1574429247] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393024 kB) free!
[1574429367] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35703852 kB) free.
[1574429390] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574429864] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 69.6% (17060156 kB) free.
[1574429924] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(8) < {writer,send_failed,{error,timeout}}
[1574429967] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (2310988 kB) free!
[1574430011] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 351 processes
[1574430044] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1574430087] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (381744 kB) free!
[1574430131] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 352 processes
[1574430184] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.3% (8011864 kB) free!
[1574430207] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398896 kB) free!
[1574430251] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 352 processes
[1574430304] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.2% (7998452 kB) free!
[1574430424] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.5% (8603328 kB) free!
[1574430464] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173604 kB) free!
[1574430584] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193056 kB) free!
[1574430591] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(85) < 2019/11/22 19:14:37 [error] 24129#0: *8596073 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574430704] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193340 kB) free!
[1574430711] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574430851] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1574431159] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 92.5% (22675332 kB) free.
[1574431190] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.6% (21117852 kB) free.
[1574431220] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.9% (44625796 kB) free.
[1574431244] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(10) < {writer,send_failed,{error,timeout}}
[1574431258] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 28.9% (14250104 kB) free.
[1574431297] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 72.8% (17849368 kB) free.
[1574431364] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1574431629] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(5) < {writer,send_failed,{error,timeout}}
[1574431750] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;2;(3) < {writer,send_failed,{error,timeout}}
[1574431762] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 73.38, 53.33, 39.22
[1574431870] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1574431882] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 59.72, 53.64, 40.98
[1574431965] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(3) < {writer,send_failed,{error,timeout}}
[1574431997] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.4% (2834604 kB) free!
[1574432002] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 186.12, 108.50, 62.75
[1574432059] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(32) < {writer,send_failed,{error,timeout}}
[1574432085] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1574432117] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.1% (2941284 kB) free!
[1574432181] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1574432221] SERVICE FLAPPING ALERT: primary-server;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574432237] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (3019388 kB) free!
[1574432296] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.1% (4181748 kB) free!
[1574432416] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.3% (3983412 kB) free!
[1574432511] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(99) < 2019/11/22 19:50:39 [error] 24134#0: *8602835 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574432536] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.2% (3961792 kB) free!
[1574432632] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574432838] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.3% (4924660 kB) free.
[1574433381] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(39) < {writer,send_failed,{error,timeout}}
[1574433424] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4944012 kB) free!
[1574433502] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1574433624] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (989848 kB) free!
[1574433670] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(6) < {writer,send_failed,{error,timeout}}
[1574433744] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (384892 kB) free!
[1574433790] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1574433864] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32994512 kB) free.
[1574433885] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(36) < {writer,send_failed,{error,timeout}}
[1574434005] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;(16) < {writer,send_failed,{error,timeout}}
[1574434103] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(158) < {writer,send_failed,{error,timeout}}
[1574434125] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;HARD;3;(8) < {writer,send_failed,{error,timeout}}
[1574434144] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.5% (3926200 kB) free!
[1574434221] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1574434224] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1574434264] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.0% (3597956 kB) free!
[1574434336] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1845228 kB) free!
[1574434341] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 375 processes
[1574434384] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.6% (3285804 kB) free!
[1574434390] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(47) < {writer,send_failed,{error,timeout}}
[1574434431] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(21) < 2019/11/22 20:20:27 [error] 24134#0: *8608464 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1574434464] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3630252 kB) free!
[1574434510] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;2;(1) < {writer,send_failed,{error,timeout}}
[1574434551] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574434584] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (383552 kB) free!
[1574434630] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;HARD;3;(6) < {writer,send_failed,{error,timeout}}
[1574434704] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36708464 kB) free.
[1574434824] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(52) < {writer,send_failed,{error,timeout}}
[1574434946] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1574434984] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1574434984] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1478728 kB) free!
[1574435117] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4745224 kB) free!
[1574435151] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(69) < 2019/11/22 20:26:56 [error] 24132#0: *8609664 recv() failed (104: Connection reset by peer) while proxying connection, client: 139.59.21.208, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:41/0, bytes from/to upstream:0/41
[1574435237] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4424128 kB) free!
[1574435271] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574435274] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 77.3% (18957568 kB) free.
[1574435304] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (6299556 kB) free!
[1574435310] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1574435310] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 68.7% (33846564 kB) free.
[1574435344] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 77.7% (17446912 kB) free.
[1574435357] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 16.7% (4113836 kB) free!
[1574435386] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.5% (23031072 kB) free!
[1574435416] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 39.3% (51759180 kB) free.
[1574435545] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(24) < {writer,send_failed,{error,timeout}}
[1574435667] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1574435957] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2437964 kB) free!
[1574435996] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4935336 kB) free!
[1574436117] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.8% (4878136 kB) free!
[1574436237] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (4825196 kB) free!
[1574436352] SERVICE ALERT: sservercluster1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45716 MB (20% inode=99%):
[1574436471] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(26) < 2019/11/22 20:54:34 [error] 24129#0: *8614908 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.218.150, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574436472] SERVICE ALERT: sservercluster1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45715 MB (20% inode=99%):
[1574436525] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1574436591] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574436592] SERVICE ALERT: sservercluster1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45250 MB (20% inode=99%):
[1574437030] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1574437191] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(74) < 2019/11/22 21:03:00 [error] 24130#0: *8616421 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574437216] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22152316 kB) free!
[1574437311] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574437336] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.0% (18437980 kB) free!
[1574437437] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202864 kB) free!
[1574437456] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.1% (48877212 kB) free.
[1574437755] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4814732 kB) free!
[1574437757] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 94.7% (23372396 kB) free.
[1574437875] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (4781472 kB) free!
[1574437995] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (4756964 kB) free!
[1574438056] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24059448 kB) free!
[1574438176] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.3% (20191576 kB) free!
[1574438296] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.0% (50116452 kB) free.
[1574438511] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(21) < 2019/11/22 21:30:46 [error] 24132#0: *8621662 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574438602] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 14.90, 15.62, 39.37
[1574438632] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574438896] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (23642456 kB) free!
[1574439016] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.2% (20062092 kB) free!
[1574439137] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.7% (49653380 kB) free.
[1574439202] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 7.41, 8.61, 24.24
[1574439231] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(67) < 2019/11/22 21:40:06 [error] 24128#0: *8623354 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1574439351] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574439438] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (3051116 kB) free!
[1574439558] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (3181896 kB) free!
[1574439678] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3261100 kB) free.
[1574440303] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (4688792 kB) free!
[1574440423] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.9% (4387008 kB) free!
[1574440543] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.7% (4582000 kB) free!
[1574440922] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(66) < 2019/11/22 22:04:11 [error] 26853#0: *8383530 writev() failed (32: Broken pipe) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:840/80654, bytes from/to upstream:84765/755
[1574441042] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574441151] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(132) < 2019/11/22 22:15:35 [error] 24130#0: *8629506 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:71/33, bytes from/to upstream:33/71
[1574441271] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1574441271] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574441365] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574441872] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(35) < 2019/11/22 22:24:25 [error] 24134#0: *8631308 connect() failed (111: Connection refused) while connecting to upstream, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574441992] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574442167] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.6% (23095004 kB) free.
[1574442251] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574442729] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1574442849] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1574442944] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 20.6% (5037624 kB) free.
[1574442969] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1574443192] SERVICE ALERT: sservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 162837 MB (73% inode=99%):
[1574443192] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2019/11/22 22:45:53 [error] 24130#0: *8634901 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574443312] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574443336] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24338864 kB) free!
[1574443456] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.8% (20777820 kB) free!
[1574443544] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4724300 kB) free!
[1574443576] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.5% (46800116 kB) free.
[1574443664] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4642088 kB) free!
[1574443784] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.5% (4536364 kB) free!
[1574443912] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(81) < 2019/11/22 23:00:47 [error] 24134#0: *8637501 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1574444032] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574444176] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25307296 kB) free!
[1574444296] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.8% (22151624 kB) free!
[1574444314] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1574444416] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.4% (48011412 kB) free.
[1574444633] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(35) < 2019/11/22 23:04:23 [error] 24131#0: *8638162 connect() failed (111: Connection refused) while connecting to upstream, client: 210.212.235.210, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574444753] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574444955] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574445953] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/11/22 23:32:04 [error] 24131#0: *8642741 connect() failed (111: Connection refused) while connecting to upstream, client: 103.15.179.39, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574446073] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574446673] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(51) < 2019/11/22 23:44:03 [error] 24130#0: *8644766 connect() failed (111: Connection refused) while connecting to upstream, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1574446675] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6486944 kB) free!
[1574446793] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1574446795] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (387076 kB) free!
[1574446817] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25382200 kB) free!
[1574446915] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2318080 kB) free!
[1574446937] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.9% (22283988 kB) free!
[1574447057] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.9% (46066792 kB) free.