[1582137090] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1582137377] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 172.105.59.173 [20/Feb/2020:00:01:49 +0530] TCP 502 0 0 60.001 "10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.001, 60.000, 0.000"
[1582137497] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582137596] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(66) < 2020/02/20 00:01:34 [error] 12251#0: *56760866 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582137600] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 58.9% (29034888 kB) free.
[1582137716] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582137793] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186404 kB) free!
[1582137851] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186088 kB) free!
[1582137913] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (183948 kB) free!
[1582137971] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172576 kB) free!
[1582138033] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186188 kB) free!
[1582138091] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177012 kB) free!
[1582139074] SERVICE ALERT: webserver2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582139134] SERVICE ALERT: webserver2;Sensors Status;OK;HARD;1;SENSORS OK
[1582139143] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182084 kB) free!
[1582139203] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (200404 kB) free!
[1582139263] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 68.2% (16721200 kB) free.
[1582139297] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(5) < 49.50.64.122 [20/Feb/2020:00:37:04 +0530] TCP 502 0 0 60.001 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.001, 0.000"
[1582139417] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582140415] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582141576] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8982988 kB) free!
[1582141697] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.4% (8586820 kB) free!
[1582141817] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.5% (8107392 kB) free!
[1582142417] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (2827784 kB) free!
[1582144184] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1582144304] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1582144424] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1582144889] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1582145009] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1582145552] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (216328 kB) free!
[1582145624] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1582145744] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 705 processes
[1582145864] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 705 processes
[1582145944] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (222552 kB) free!
[1582146064] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (247644 kB) free!
[1582146184] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (221120 kB) free!
[1582146784] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 76.8% (37838804 kB) free.
[1582147033] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 48.0% (11770652 kB) free.
[1582147352] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 82.0% (40432204 kB) free.
[1582147664] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 697 processes
[1582147817] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1759) < 49.50.86.152 [20/Feb/2020:03:00:17 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1582147916] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(256) < 2020/02/20 03:00:21 [error] 12257#0: *56795079 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582147937] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(233) < 203.212.70.170 [20/Feb/2020:03:01:32 +0530] TCP 502 0 0 60.000 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1582148036] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582148057] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(2) < 172.105.59.173 [20/Feb/2020:03:02:22 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1582148205] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(215) < 103.16.101.72 [20/Feb/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1582148327] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582148657] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582149464] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1582149584] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1582155434] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172932 kB) free!
[1582155554] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (200596 kB) free!
[1582155674] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183600 kB) free!
[1582156330] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 75.3% (18456224 kB) free.
[1582156691] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 48.9% (12072176 kB) free.
[1582160217] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.3% (32013700 kB) free.
[1582160384] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1582160504] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1582161236] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/20 06:42:42 [error] 12256#0: *56831014 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1582161356] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582161575] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.1% (39485588 kB) free.
[1582161600] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.0% (21815172 kB) free.
[1582161809] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1582161929] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1582163291] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179728 kB) free!
[1582163411] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175304 kB) free!
[1582163531] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164336 kB) free!
[1582166156] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(14) < 2020/02/20 08:01:02 [error] 12254#0: *56849291 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582166276] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582166740] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.3% (22039552 kB) free.
[1582169504] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1582169624] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1582170416] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25405580 kB) free!
[1582170536] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.9% (44709716 kB) free.
[1582171136] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25406316 kB) free!
[1582171256] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.9% (44691796 kB) free.
[1582171856] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24097588 kB) free!
[1582171911] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.5% (4789092 kB) free!
[1582171976] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.2% (43732316 kB) free.
[1582172031] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.1% (4685620 kB) free!
[1582172151] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.7% (4579540 kB) free!
[1582172576] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21893756 kB) free!
[1582172696] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42523560 kB) free.
[1582172729] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1582172850] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1582173225] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1582173296] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19891672 kB) free!
[1582173345] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1582173416] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42839392 kB) free.
[1582173952] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1973232 kB) free!
[1582174016] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18863432 kB) free!
[1582174136] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.3% (14934732 kB) free!
[1582174257] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39135564 kB) free.
[1582174545] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1582174665] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 703 processes
[1582174776] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.4% (21677536 kB) free.
[1582174785] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1582174856] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15878380 kB) free!
[1582174976] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28817364 kB) free.
[1582175385] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 692 processes
[1582175576] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16934208 kB) free!
[1582175696] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (12806060 kB) free!
[1582175816] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36746992 kB) free.
[1582176416] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5463048 kB) free!
[1582176536] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1528164 kB) free!
[1582176656] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36207192 kB) free.
[1582177062] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.9% (4235564 kB) free!
[1582177182] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.9% (3798568 kB) free!
[1582177256] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12161500 kB) free!
[1582177302] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.0% (3366948 kB) free!
[1582177376] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (8163924 kB) free!
[1582177496] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (4278600 kB) free!
[1582177902] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (1295120 kB) free!
[1582178057] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(585) < 103.255.146.164 [20/Feb/2020:11:19:21 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1582178177] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582178696] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16756888 kB) free!
[1582178748] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172400 kB) free!
[1582178868] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168240 kB) free!
[1582178975] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8829188 kB) free!
[1582178988] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169360 kB) free!
[1582179000] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8076096 kB) free!
[1582179007] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.5% (4045156 kB) free!
[1582179096] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.9% (8828840 kB) free!
[1582179120] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7545052 kB) free!
[1582179127] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.6% (3824596 kB) free!
[1582179130] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4787152 kB) free!
[1582179152] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4906152 kB) free!
[1582179216] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.9% (8829088 kB) free!
[1582179240] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.3% (7048412 kB) free!
[1582179247] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.8% (3631796 kB) free!
[1582179250] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (4615168 kB) free!
[1582179272] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4539464 kB) free!
[1582179297] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39359412 kB) free.
[1582179370] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.0% (4399904 kB) free!
[1582179392] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 16.5% (4082876 kB) free!
[1582179840] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4144568 kB) free!
[1582179847] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2306016 kB) free!
[1582179897] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14455784 kB) free!
[1582179977] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(551) < 103.255.146.164 [20/Feb/2020:11:48:56 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1582179992] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1981668 kB) free!
[1582180017] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (10504480 kB) free!
[1582180049] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1582180097] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582180137] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39341576 kB) free.
[1582180169] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1582180302] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.5% (16061268 kB) free.
[1582180441] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.0% (30083944 kB) free.
[1582180570] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2308152 kB) free!
[1582180676] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(96) < 2020/02/20 12:01:48 [error] 12245#0: *56938545 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582180697] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 203.212.70.170 [20/Feb/2020:12:03:09 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1582180737] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14265220 kB) free!
[1582180797] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582180817] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582180857] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (9988568 kB) free!
[1582180967] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.7% (21650900 kB) free.
[1582180977] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1582180977] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (5981424 kB) free!
[1582181078] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.5% (22843416 kB) free.
[1582181127] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 57.9% (76320284 kB) free.
[1582181327] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.1% (39469624 kB) free.
[1582181398] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.0% (21559288 kB) free.
[1582181450] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.5% (21435752 kB) free.
[1582181552] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9454824 kB) free!
[1582181672] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.3% (9042040 kB) free!
[1582181792] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.6% (8667580 kB) free!
[1582182178] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (2463012 kB) free!
[1582182298] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165588 kB) free!
[1582182418] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175328 kB) free!
[1582182938] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21404412 kB) free!
[1582183058] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 40.8% (53788052 kB) free.
[1582183185] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.3% (4737880 kB) free!
[1582183185] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1582183305] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.3% (4473808 kB) free!
[1582183305] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 703 processes
[1582183425] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.1% (4198548 kB) free!
[1582183425] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 703 processes
[1582183592] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2939856 kB) free!
[1582184625] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1714240 kB) free!
[1582184626] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1582184626] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1582184858] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16813976 kB) free!
[1582184978] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.6% (49523884 kB) free.
[1582185226] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1582185346] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1582185466] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 703 processes
[1582185538] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9832536 kB) free!
[1582185578] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23247172 kB) free!
[1582185658] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.2% (8961936 kB) free!
[1582185698] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.4% (18972124 kB) free!
[1582185778] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.8% (7790452 kB) free!
[1582185818] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.3% (14923068 kB) free!
[1582185855] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.4% (21675888 kB) free.
[1582185912] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1582185912] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.7% (21633324 kB) free.
[1582185957] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 58.9% (77670540 kB) free.
[1582186009] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1582186009] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.5% (38705016 kB) free.
[1582186060] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.7% (42706820 kB) free.
[1582186984] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9335568 kB) free!
[1582187104] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (9175544 kB) free!
[1582187115] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (4396820 kB) free!
[1582187224] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (9026308 kB) free!
[1582187235] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.4% (4291968 kB) free!
[1582187355] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (4175604 kB) free!
[1582187728] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.8% (2397968 kB) free!
[1582187765] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (22305540 kB) free!
[1582187848] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.2% (4259120 kB) free.
[1582187866] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1582187885] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.6% (46885176 kB) free.
[1582187955] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195276 kB) free!
[1582188017] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(395) < 103.255.146.164 [20/Feb/2020:14:01:31 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1582188055] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195468 kB) free!
[1582188103] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.8% (4212880 kB) free!
[1582188137] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582188175] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189296 kB) free!
[1582188223] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.8% (3999020 kB) free!
[1582188295] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197436 kB) free!
[1582188343] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.9% (3798972 kB) free!
[1582188606] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4881396 kB) free!
[1582188726] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4635460 kB) free!
[1582188846] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (4425928 kB) free!
[1582189543] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1582189543] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1452204 kB) free!
[1582189685] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15474212 kB) free!
[1582189805] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.7% (24678876 kB) free!
[1582189925] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16801088 kB) free!
[1582190046] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1582190046] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2175588 kB) free!
[1582190334] SERVICE ALERT: smscapp1;Current Users;WARNING;SOFT;1;USERS WARNING - 12 users currently logged in
[1582190454] SERVICE ALERT: smscapp1;Current Users;WARNING;SOFT;2;USERS WARNING - 12 users currently logged in
[1582190525] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390148 kB) free!
[1582190574] SERVICE ALERT: smscapp1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1582190804] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1582190824] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4669332 kB) free!
[1582190924] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1582191137] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(159) < 185.255.8.58 [20/Feb/2020:15:00:24 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1582191257] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582191343] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.9% (16147688 kB) free.
[1582191448] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.6% (2373624 kB) free!
[1582191524] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 352 processes
[1582191568] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.1% (2297924 kB) free!
[1582191644] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1582191688] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.6% (2365248 kB) free!
[1582191857] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(501) < 103.255.146.164 [20/Feb/2020:15:05:28 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1582191885] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 16.8% (4142392 kB) free!
[1582191978] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582192005] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.3% (22781412 kB) free.
[1582192066] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 709 processes
[1582192186] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 708 processes
[1582192288] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.9% (4372484 kB) free.
[1582192306] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 708 processes
[1582192325] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14272364 kB) free!
[1582192491] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1582192611] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1582192844] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1582192925] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.1% (44917356 kB) free.
[1582192964] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1582193177] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(495) < 103.255.146.164 [20/Feb/2020:15:30:58 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1582193216] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.9% (8352892 kB) free!
[1582193297] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582193336] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.9% (8344524 kB) free!
[1582193456] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.9% (8339724 kB) free!
[1582193488] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.1% (2289712 kB) free!
[1582193526] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11667500 kB) free!
[1582193564] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582193608] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.9% (2588032 kB) free!
[1582193646] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7276272 kB) free!
[1582193648] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.2% (22012168 kB) free.
[1582193685] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1582193693] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 84.9% (41833000 kB) free.
[1582193728] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.7% (2708232 kB) free!
[1582193766] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.9% (48704000 kB) free.
[1582193801] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.5% (21927360 kB) free.
[1582193888] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.6% (21232240 kB) free.
[1582194328] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1582194328] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3666652 kB) free.
[1582194495] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174768 kB) free!
[1582194615] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181180 kB) free!
[1582194735] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188096 kB) free!
[1582194928] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (1358992 kB) free!
[1582194977] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26386740 kB) free!
[1582195048] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.3% (1676828 kB) free!
[1582195098] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.8% (51169996 kB) free.
[1582195098] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(21) < 203.212.70.170 [20/Feb/2020:16:02:45 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1582195168] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.0% (1946620 kB) free!
[1582195196] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(188) < 2020/02/20 16:02:26 [error] 12257#0: *57080746 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582195218] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582195316] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582195484] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582195604] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1582195697] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20796192 kB) free!
[1582195724] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1582195817] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.4% (16332160 kB) free!
[1582195937] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (12110400 kB) free!
[1582196324] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582196344] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 50.3% (66290860 kB) free.
[1582196369] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.9% (3722300 kB) free.
[1582196444] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1582196456] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4344736 kB) free!
[1582196564] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582196574] SERVICE ALERT: smscapp1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1582196969] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.1% (2621796 kB) free!
[1582197089] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (3239776 kB) free!
[1582197165] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1582197209] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3167592 kB) free!
[1582197258] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.1% (4220464 kB) free!
[1582197378] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.2% (3992528 kB) free!
[1582197411] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1582197498] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.2% (3755488 kB) free!
[1582197531] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1582197706] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1582197765] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582197885] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 350 processes
[1582198005] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1582198148] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13659972 kB) free!
[1582198268] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9757956 kB) free!
[1582198388] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.7% (45790096 kB) free.
[1582198409] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (881144 kB) free!
[1582198605] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582198698] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (1149656 kB) free!
[1582198725] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1582198988] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20850864 kB) free!
[1582199009] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3348604 kB) free.
[1582199108] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.3% (16161400 kB) free!
[1582199143] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.6% (3047140 kB) free!
[1582199228] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (11737308 kB) free!
[1582199264] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.1% (2724496 kB) free!
[1582199325] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1582199384] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.0% (2479772 kB) free!
[1582199445] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1582199609] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.3% (2157816 kB) free!
[1582199655] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.3% (4483248 kB) free!
[1582199729] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (2952728 kB) free!
[1582199775] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.8% (4367656 kB) free!
[1582199828] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30924260 kB) free.
[1582199849] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.2% (3768612 kB) free.
[1582199895] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.2% (4222380 kB) free!
[1582199984] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (1182292 kB) free!
[1582200045] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1582200165] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1582200286] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1582200428] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6499316 kB) free!
[1582200449] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.7% (2559540 kB) free!
[1582200462] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1582200468] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9219292 kB) free!
[1582200548] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (2465412 kB) free!
[1582200569] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.3% (3466272 kB) free.
[1582200588] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.3% (9015200 kB) free!
[1582200668] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (43032360 kB) free.
[1582200708] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.8% (8780132 kB) free!
[1582200841] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9600288 kB) free!
[1582200961] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.7% (9227472 kB) free!
[1582201081] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.0% (8872496 kB) free!
[1582201169] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.3% (2319352 kB) free!
[1582201268] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12460936 kB) free!
[1582201290] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3375348 kB) free.
[1582201306] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 703 processes
[1582201388] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7890536 kB) free!
[1582201426] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1582201485] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1582201508] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (3523468 kB) free!
[1582201546] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 702 processes
[1582201605] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1582201725] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1582201889] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.6% (2370356 kB) free!
[1582202009] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.2% (3444256 kB) free.
[1582202109] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18332276 kB) free!
[1582202205] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4704216 kB) free!
[1582202262] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.7% (4347552 kB) free!
[1582202282] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 58.8% (28990400 kB) free.
[1582202295] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1971164 kB) free!
[1582202325] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1582202325] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4450084 kB) free!
[1582202382] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.1% (4201476 kB) free!
[1582202384] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.1% (16179456 kB) free.
[1582202445] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1582202446] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 16.9% (4170560 kB) free!
[1582202502] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.4% (4016940 kB) free!
[1582202609] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.8% (2410776 kB) free!
[1582202709] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39051524 kB) free.
[1582202729] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3541164 kB) free.
[1582202736] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.5% (21679144 kB) free.
[1582202746] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1582202772] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.8% (21671648 kB) free.
[1582202800] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.1% (39461252 kB) free.
[1582202825] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 77.3% (18945732 kB) free.
[1582203045] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1582203046] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.2% (22769848 kB) free.
[1582203166] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1582203286] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582203309] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9359124 kB) free!
[1582203329] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.6% (2532336 kB) free!
[1582203426] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165932 kB) free!
[1582203429] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4837644 kB) free!
[1582203449] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.1% (2448920 kB) free!
[1582203546] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170772 kB) free!
[1582203549] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400080 kB) free!
[1582203569] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.6% (2373256 kB) free!
[1582203600] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (3081528 kB) free!
[1582203666] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1582203666] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167012 kB) free!
[1582203703] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 20.1% (4923348 kB) free.
[1582203708] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (4010400 kB) free!
[1582203720] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (3073224 kB) free!
[1582203840] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 18.9% (3072236 kB) free!
[1582203886] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1582203886] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 351 processes
[1582203977] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193476 kB) free!
[1582204097] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166128 kB) free!
[1582204169] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.9% (4699968 kB) free.
[1582204217] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164400 kB) free!
[1582204303] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.3% (4241980 kB) free!
[1582204423] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.8% (4107544 kB) free!
[1582204486] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582204543] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1582204543] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.2% (3970208 kB) free!
[1582205086] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1582205369] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3036952 kB) free!
[1582205489] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.2% (3605312 kB) free.
[1582205743] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2235892 kB) free!
[1582205949] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (42980232 kB) free.
[1582206089] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (1389856 kB) free!
[1582206209] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.7% (2070100 kB) free!
[1582206329] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.7% (2546648 kB) free!
[1582206405] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8983224 kB) free!
[1582206525] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.2% (8982744 kB) free!
[1582206549] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13418324 kB) free!
[1582206645] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.2% (8983492 kB) free!
[1582206669] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8840040 kB) free!
[1582206789] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (4622016 kB) free!
[1582206929] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3690600 kB) free.
[1582207389] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.6% (24529176 kB) free!
[1582207486] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582207530] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.9% (2748248 kB) free!
[1582207606] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1582207650] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.2% (3443724 kB) free.
[1582207989] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1388424 kB) free!
[1582208206] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1582208250] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (853124 kB) free!
[1582208326] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582208370] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.0% (4059424 kB) free.
[1582208446] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1582208590] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.6% (21818980 kB) free!
[1582209046] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1582209116] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(153) < 2020/02/20 20:01:48 [error] 12254#0: *57194397 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:2075/17, bytes from/to upstream:17/2075
[1582209166] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 374 processes
[1582209236] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(7) < 2020/02/20 20:01:58 [error] 12246#0: *57194908 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582209287] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 368 processes
[1582209356] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582209570] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.0% (2933628 kB) free!
[1582209618] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 203.212.70.170 [20/Feb/2020:20:02:51 +0530] TCP 502 0 0 120.015 "10.147.212.154:3001, 10.147.212.142:3001, 10.147.212.155:3001, 10.147.212.143:3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 60.007, 0.000, 60.007"
[1582209645] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4703756 kB) free!
[1582209690] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.0% (3406680 kB) free.
[1582209738] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582209886] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582210006] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582210126] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1582210288] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.2739.2312> (10.147.212.53:46368 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1582210390] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44325224 kB) free.
[1582210408] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1582210890] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.0% (2596468 kB) free!
[1582210990] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16115800 kB) free!
[1582211010] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3145632 kB) free!
[1582211110] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (11375264 kB) free!
[1582211130] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.2% (3608948 kB) free.
[1582211230] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (7536168 kB) free!
[1582211830] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30542468 kB) free.
[1582211984] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1582212330] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.8% (2723564 kB) free!
[1582212430] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8670544 kB) free!
[1582212450] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (3012260 kB) free!
[1582212550] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (4956748 kB) free!
[1582212571] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3333160 kB) free.
[1582212670] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41385256 kB) free.
[1582213126] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1582213270] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16861440 kB) free!
[1582213390] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (12795064 kB) free!
[1582213510] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (9105404 kB) free!
[1582213546] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1582213726] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582213817] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1582213822] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1582213943] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1582213963] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 16.3% (3995492 kB) free!
[1582214023] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 16.3% (3994956 kB) free!
[1582214063] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1582214083] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 16.3% (3986220 kB) free!
[1582214110] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26042284 kB) free!
[1582214143] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 16.3% (3988312 kB) free!
[1582214444] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 24.9% (6110676 kB) free.
[1582214711] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (5767616 kB) free!
[1582214744] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 17.5% (4279676 kB) free!
[1582214804] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 10.1% (2464324 kB) free!
[1582214864] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 10.1% (2465588 kB) free!
[1582214924] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 10.1% (2465608 kB) free!
[1582215066] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582215263] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1582215311] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21586908 kB) free!
[1582215526] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1582215911] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42118324 kB) free.
[1582216124] SERVICE FLAPPING ALERT: webserver2;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1582216124] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 1.8% (431304 kB) free!
[1582216280] HOST ALERT: sservercluster1;DOWN;SOFT;1;(Host check timed out after 35.22 seconds)
[1582216286] HOST ALERT: sservercluster1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1582216511] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16762596 kB) free!
[1582216543] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1582216631] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (12671692 kB) free!
[1582216726] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582216751] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (8977256 kB) free!
[1582216846] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1582217351] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23059512 kB) free!
[1582217951] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (11691208 kB) free!
[1582218139] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(147) < 103.255.146.164 [20/Feb/2020:22:23:03 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1582218259] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582218551] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (18500368 kB) free!
[1582218645] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1582219151] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37419200 kB) free.
[1582219170] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582219751] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14529268 kB) free!
[1582219871] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (10817052 kB) free!
[1582219991] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (7397844 kB) free!
[1582220591] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.6% (21867064 kB) free!
[1582220924] SERVICE FLAPPING ALERT: webserver2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1582221191] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41103852 kB) free.
[1582221791] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16478244 kB) free!
[1582221912] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (12663700 kB) free!
[1582222032] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (9229660 kB) free!
[1582222084] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (237944 kB) free!
[1582222204] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (242032 kB) free!
[1582222324] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (238392 kB) free!
[1582222618] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (431888 kB) free!
[1582222632] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.4% (22974140 kB) free!
[1582222738] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (385036 kB) free!
[1582222858] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403792 kB) free!
[1582223232] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41085496 kB) free.