[1569004486] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.94, 1.10, 4.93
[1569004667] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569004733] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177624 kB) free!
[1569004854] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (203008 kB) free!
[1569004974] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183828 kB) free!
[1569005066] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (228244 kB) free!
[1569005186] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183264 kB) free!
[1569005306] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180192 kB) free!
[1569005867] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569006467] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569007667] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569008150] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 279 processes
[1569008267] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569008398] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (790460 kB) free!
[1569008490] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (241324 kB) free!
[1569008507] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 454 processes
[1569008518] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (378120 kB) free!
[1569008611] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (243784 kB) free!
[1569008625] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 222730 MB (20% inode=99%):
[1569008638] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (3347936 kB) free!
[1569008731] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (239368 kB) free!
[1569008745] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 213246 MB (19% inode=99%):
[1569008865] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 206655 MB (19% inode=99%):
[1569009350] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 274 processes
[1569009353] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.0% (8846308 kB) free!
[1569009467] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569009473] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 14.5% (7165888 kB) free!
[1569009593] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.0% (6912480 kB) free!
[1569010068] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 599 processes
[1569011268] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569011868] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569012468] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 607 processes
[1569012609] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (230776 kB) free!
[1569012729] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (228156 kB) free!
[1569012849] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (230700 kB) free!
[1569014049] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.2% (42496992 kB) free.
[1569014686] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1569015286] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 2.31, 32.90, 23.99
[1569015406] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.00, 22.31, 21.19
[1569015526] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.15, 15.32, 18.77
[1569015716] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 50.1% (12283752 kB) free.
[1569015774] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 50.0% (12248612 kB) free.
[1569017326] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 60.13, 33.75, 16.00
[1569017868] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569017926] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.84, 6.22, 9.84
[1569018468] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569019068] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569019126] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1569019126] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.92, 0.98, 3.28
[1569020268] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569020277] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1569020397] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1569021468] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569022068] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569022669] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569025197] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 403 processes
[1569025318] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1569025669] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 604 processes
[1569026132] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 94.7% (33178584 kB) free.
[1569026270] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569026638] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.7% (22026380 kB) free!
[1569026870] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569026954] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 485 processes
[1569027015] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 481 processes
[1569027239] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (11228992 kB) free!
[1569028070] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569028439] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 48.4% (63772392 kB) free.
[1569028665] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 279822 MB (26% inode=99%):
[1569028706] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 65.2% (16100864 kB) free.
[1569028917] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1569029037] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1569029270] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569029306] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181736 kB) free!
[1569029426] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192932 kB) free!
[1569029518] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197596 kB) free!
[1569029546] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187816 kB) free!
[1569029575] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165112 kB) free!
[1569029640] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171384 kB) free!
[1569029695] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193604 kB) free!
[1569029761] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180080 kB) free!
[1569029815] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187804 kB) free!
[1569029871] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569029927] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1569030837] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1569030957] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1569031071] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569031671] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569032871] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569033472] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569034672] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569035272] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569036472] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569037672] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569038272] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569038872] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569040073] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569040673] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569041873] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 605 processes
[1569042227] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.39, 9.73, 7.91
[1569042347] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.31, 7.44, 7.29
[1569042467] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.82, 6.22, 6.85
[1569042593] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4736900 kB) free!
[1569046907] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 560 processes
[1569047028] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 560 processes
[1569047148] SERVICE ALERT: smscapp2;Total Processes;WARNING;HARD;3;PROCS WARNING: 561 processes
[1569048249] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9354896 kB) free!
[1569048369] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8875924 kB) free!
[1569048489] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.1% (8432856 kB) free!
[1569049689] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4484016 kB) free!
[1569051473] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569052470] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.0% (23196860 kB) free.
[1569052518] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.7% (22730424 kB) free.
[1569052567] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.0% (39906060 kB) free.
[1569052623] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 96.0% (47327688 kB) free.
[1569052667] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.0% (23040024 kB) free.
[1569052673] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 604 processes
[1569053873] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569054473] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569055474] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.5% (3567388 kB) free!
[1569055594] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.3% (4023056 kB) free!
[1569055714] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.0% (3959432 kB) free!
[1569056314] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2400612 kB) free!
[1569056657] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1569056657] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.7% (23111480 kB) free.
[1569056874] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569057476] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4583660 kB) free!
[1569057597] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4328112 kB) free!
[1569057701] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 26602 MB (19% inode=99%):
[1569057717] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 17.2% (4063696 kB) free!
[1569057821] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 18400 MB (13% inode=99%):
[1569057941] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 18400 MB (13% inode=99%):
[1569058074] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569058317] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 97.0% (22872800 kB) free.
[1569058541] SERVICE ALERT: sig2-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 58217 MB (42% inode=99%):
[1569058674] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569059274] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569059867] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.84, 4.36, 4.94
[1569060475] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569061075] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569062267] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 14.29, 9.26, 5.68
[1569062387] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.21, 7.91, 5.62
[1569062507] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.61, 6.41, 5.35
[1569062664] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1569063107] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.61, 2.96, 3.95
[1569063117] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 57.73, 16.67, 6.21
[1569063237] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 90.38, 42.85, 17.08
[1569063357] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 95.05, 60.38, 26.63
[1569064557] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.83, 3.12, 12.84
[1569066357] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.63, 1.36, 4.22
[1569072475] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569073672] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4831352 kB) free!
[1569073792] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (4797396 kB) free!
[1569073913] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.4% (4761440 kB) free!
[1569074064] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4935640 kB) free!
[1569074127] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4773724 kB) free!
[1569074184] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.2% (4483556 kB) free!
[1569074247] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (4714720 kB) free!
[1569074304] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (4454948 kB) free!
[1569074367] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (4650820 kB) free!
[1569076075] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569076675] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569077748] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 347 processes
[1569077823] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9754352 kB) free!
[1569077943] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.7% (9706892 kB) free!
[1569078063] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.6% (9656716 kB) free!
[1569078348] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1569078505] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2383312 kB) free!
[1569080149] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 345 processes
[1569081113] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2395212 kB) free!
[1569081328] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.3% (23017868 kB) free.
[1569081349] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1569081349] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1569081387] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.7% (22964532 kB) free.
[1569081949] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 345 processes
[1569082549] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1569083968] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2382252 kB) free!
[1569084475] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569085726] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.79, 6.93, 7.04
[1569085846] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.66, 4.85, 6.26
[1569085966] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.63, 3.49, 5.60
[1569086566] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.62, 0.98, 3.21
[1569086875] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569089275] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569090476] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1569090476] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes