[1553193293] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553193413] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553193590] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11692988 kB) free.
[1553194013] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553194133] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553194190] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4255980 kB) free!
[1553194253] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553194310] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11698916 kB) free.
[1553194910] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6475288 kB) free!
[1553195030] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.5% (5171420 kB) free!
[1553195150] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3405884 kB) free!
[1553195453] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553195750] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.8% (7283500 kB) free!
[1553196053] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553196173] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1553196350] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11694472 kB) free.
[1553196774] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553196894] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553196950] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4070260 kB) free!
[1553197014] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553197070] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11692280 kB) free.
[1553197671] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6431268 kB) free!
[1553197791] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.5% (5154700 kB) free!
[1553197911] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (3436064 kB) free!
[1553198214] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553198511] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.8% (7291696 kB) free!
[1553198606] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 22.30, 11.88, 4.77
[1553198726] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.10, 7.98, 4.20
[1553198743] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 331 processes
[1553198814] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553198934] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1553199037] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 335 processes
[1553199111] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11697536 kB) free.
[1553199534] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553199654] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553199711] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3955536 kB) free!
[1553199774] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553199831] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11703116 kB) free.
[1553200431] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6391316 kB) free!
[1553200551] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.3% (5095352 kB) free!
[1553200671] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (3293028 kB) free!
[1553200974] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1553201271] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.8% (7307836 kB) free!
[1553201574] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553201694] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1553201872] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11702428 kB) free.
[1553202294] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553202414] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553202472] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (4208520 kB) free!
[1553202534] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553202592] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11696020 kB) free.
[1553203192] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6476900 kB) free!
[1553203312] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.5% (5199044 kB) free!
[1553203432] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (3492216 kB) free!
[1553203734] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1553204027] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 47.80, 21.14, 9.00
[1553204032] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.0% (7390572 kB) free!
[1553204147] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 51.82, 31.53, 14.28
[1553204267] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 52.11, 38.37, 18.88
[1553204334] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553204432] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5204 MB (10% inode=97%):
[1553204454] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553204632] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11698632 kB) free.
[1553204867] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.74, 11.94, 16.16
[1553205054] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553205174] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553205232] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4293660 kB) free!
[1553205294] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553205352] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11700528 kB) free.
[1553205393] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 42.59, 29.50, 13.74
[1553205513] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 42.17, 33.70, 17.19
[1553205531] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.4% (3053744 kB) free!
[1553205634] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.63, 23.18, 15.36
[1553205926] SERVICE ALERT: sig4-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (687488 kB) free!
[1553205952] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6629464 kB) free!
[1553206072] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.9% (5372600 kB) free!
[1553206193] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3825944 kB) free!
[1553206495] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1553206667] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.16, 1.37, 3.55
[1553206793] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.4% (7589440 kB) free!
[1553206834] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.71, 0.96, 4.66
[1553207095] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553207215] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1553207393] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11707624 kB) free.
[1553207815] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1553207935] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553207993] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4742100 kB) free!
[1553208055] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1553208064] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 18.43, 23.61, 13.08
[1553208113] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11639756 kB) free.
[1553208185] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.73, 15.89, 11.53
[1553208305] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.79, 10.83, 10.21
[1553208713] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7052472 kB) free!
[1553208833] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.7% (5787996 kB) free!
[1553208953] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4490460 kB) free!
[1553209255] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553209505] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.37, 0.55, 3.06
[1553209553] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.9% (8327884 kB) free!
[1553209837] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 334 processes
[1553209856] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553209976] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553210153] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11692488 kB) free.
[1553210576] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553210696] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553210726] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 33.88, 15.21, 6.00
[1553210744] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 382 processes
[1553210753] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.2% (5038484 kB) free!
[1553210816] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553210846] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 36.00, 22.20, 9.68
[1553210874] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (3222116 kB) free!
[1553210966] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 33.46, 26.30, 12.73
[1553210994] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11703224 kB) free.
[1553211344] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 321 processes
[1553211566] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.15, 3.70, 6.76
[1553211594] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5984524 kB) free!
[1553211714] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (4690700 kB) free!
[1553211834] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11701796 kB) free.
[1553212016] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553212166] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.18, 0.67, 3.64
[1553212434] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6857580 kB) free!
[1553212554] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.3% (5551640 kB) free!
[1553212616] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553212674] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4070188 kB) free!
[1553212737] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553212857] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1553213144] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 319 processes
[1553213274] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.7% (7743452 kB) free!
[1553213457] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553213577] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553213697] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553213874] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11698316 kB) free.
[1553214297] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553214417] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553214474] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4705876 kB) free!
[1553214537] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553214595] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11705164 kB) free.
[1553215194] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.0% (6890704 kB) free!
[1553215314] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.3% (5578852 kB) free!
[1553215434] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4128712 kB) free!
[1553216034] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.4% (5149252 kB) free!
[1553216337] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553216937] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553217057] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553217177] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553217234] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11695088 kB) free.
[1553217777] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553217834] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.3% (5577720 kB) free!
[1553217897] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553217954] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (4087992 kB) free!
[1553218017] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553218074] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11701156 kB) free.
[1553218674] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6308764 kB) free!
[1553218794] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.1% (4995272 kB) free!
[1553218915] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (3069060 kB) free!
[1553219217] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553219515] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.2% (7022192 kB) free!
[1553219817] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553219937] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1553220115] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11699948 kB) free.
[1553220537] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553220657] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553220715] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3737804 kB) free!
[1553220777] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553220835] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11694096 kB) free.
[1553221377] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553221435] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.6% (6214788 kB) free!
[1553221498] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553221555] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (4906588 kB) free!
[1553221618] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553221675] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2981152 kB) free!
[1553222275] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.4% (7086656 kB) free!
[1553222818] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553222875] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11691400 kB) free.
[1553223418] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553223475] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (3787240 kB) free!
[1553223538] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1553223595] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11696724 kB) free.
[1553224138] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553224196] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6255912 kB) free!
[1553224258] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553224316] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.1% (4957136 kB) free!
[1553224378] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553224436] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (3045516 kB) free!
[1553225036] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.5% (7167248 kB) free!
[1553225636] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10198420 kB) free.
[1553226178] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1553226778] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553226898] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1553227018] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1553227436] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4563576 kB) free!
[1553227556] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11683164 kB) free.
[1553228156] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3960832 kB) free!
[1553228276] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11684084 kB) free.
[1553228877] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (3632716 kB) free!
[1553228997] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11588764 kB) free.
[1553229597] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (2979984 kB) free!
[1553229717] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11297596 kB) free.
[1553230317] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (3993476 kB) free!
[1553230437] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11655544 kB) free.
[1553230763] SERVICE ALERT: sig3-vm1;Memory;OK;HARD;1;OK - 95.5% (7025524 kB) free.
[1553230814] SERVICE ALERT: sig4-vm1;Memory;OK;HARD;1;OK - 95.1% (6997056 kB) free.
[1553230861] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 95.8% (7046124 kB) free.
[1553230906] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.5% (21589884 kB) free.
[1553230951] SERVICE ALERT: sig6-vm1;Memory;OK;HARD;1;OK - 94.9% (6977656 kB) free.
[1553231037] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3853304 kB) free!
[1553231044] SERVICE ALERT: sig3-vm1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 281 processes
[1553231157] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11651256 kB) free.
[1553231164] SERVICE ALERT: sig3-vm1;CPU_Procs;OK;HARD;1;CPU OK: 276 processes
[1553231371] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (158668 kB) free!
[1553231421] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;1;WARNING - 11.8% (870752 kB) free!
[1553231464] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (146160 kB) free!
[1553231491] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (155212 kB) free!
[1553231541] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;2;WARNING - 11.8% (869264 kB) free!
[1553231584] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (140868 kB) free!
[1553231611] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (150144 kB) free!
[1553231650] SERVICE ALERT: sig3-vm1;Memory;OK;HARD;1;OK - 95.8% (7045984 kB) free.
[1553231661] SERVICE ALERT: sig4-vm1;Memory;WARNING;HARD;3;WARNING - 11.8% (867784 kB) free!
[1553231705] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 96.0% (7062004 kB) free.
[1553231747] SERVICE ALERT: sig4-vm1;Memory;OK;HARD;1;OK - 95.1% (6996924 kB) free.
[1553231757] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (2881056 kB) free!
[1553231797] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 43.4% (21382980 kB) free.
[1553231818] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553232418] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553232538] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553232658] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553233004] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.8% (5311464 kB) free!
[1553233124] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12526108 kB) free.
[1553233413] SERVICE ALERT: webservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: / 23581 MB (46% inode=99%):
[1553233724] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4486756 kB) free!
[1553233844] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12437224 kB) free.
[1553234445] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4252396 kB) free!
[1553234565] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12509832 kB) free.
[1553235165] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4605656 kB) free!
[1553235285] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12540516 kB) free.
[1553235885] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4076744 kB) free!
[1553236005] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12153360 kB) free.
[1553237205] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (4819876 kB) free!
[1553237325] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1553237325] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12512932 kB) free.
[1553237925] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3939556 kB) free!
[1553238045] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12246456 kB) free.
[1553238344] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 322 processes
[1553238632] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 7147 MB (13% inode=98%):
[1553238637] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 335 processes
[1553238645] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (3423716 kB) free!
[1553238659] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1553238765] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11669636 kB) free.
[1553239259] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553239379] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553239965] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4932936 kB) free!
[1553239979] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553240085] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12469428 kB) free.
[1553240099] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553240685] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4312688 kB) free!
[1553240699] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553240744] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 323 processes
[1553240805] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12381840 kB) free.
[1553240819] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553240939] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553241037] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 338 processes
[1553241405] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3616932 kB) free!
[1553241525] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11807768 kB) free.
[1553242139] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1553242725] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.5% (5175120 kB) free!
[1553242739] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1553242845] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12447220 kB) free.
[1553242859] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1553243446] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (4450124 kB) free!
[1553243459] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553243566] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12379248 kB) free.
[1553243579] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553243699] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553244166] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4085016 kB) free!
[1553244286] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12166936 kB) free.
[1553244886] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4251148 kB) free!
[1553245006] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12301620 kB) free.
[1553245548] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47115 MB (20% inode=99%):
[1553245606] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4273076 kB) free!
[1553245668] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 57291 MB (25% inode=99%):
[1553245726] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12295000 kB) free.
[1553246326] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (3684692 kB) free!
[1553246446] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11699732 kB) free.
[1553246861] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3237160 kB) free!
[1553246981] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.8% (3220816 kB) free!
[1553247046] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3433876 kB) free!
[1553247101] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.5% (3171564 kB) free!
[1553247166] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11684504 kB) free.
[1553247767] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4156928 kB) free!
[1553247887] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12220604 kB) free.
[1553248487] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3766556 kB) free!
[1553248607] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11868576 kB) free.
[1553249339] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 89.2% (14498552 kB) free.
[1553249807] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (4851532 kB) free!
[1553249927] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12399620 kB) free.
[1553250527] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4308148 kB) free!
[1553250647] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12235140 kB) free.
[1553251247] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3551948 kB) free!
[1553251367] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11604556 kB) free.
[1553252568] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (4447604 kB) free!
[1553252688] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12357240 kB) free.
[1553253288] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (3741684 kB) free!
[1553253408] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11881020 kB) free.
[1553254499] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1553254608] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (4776824 kB) free!
[1553254728] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12349672 kB) free.
[1553255328] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4080868 kB) free!
[1553255448] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12182824 kB) free.
[1553256048] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (3970228 kB) free!
[1553256168] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12112788 kB) free.
[1553256769] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (3354412 kB) free!
[1553256889] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11526824 kB) free.
[1553257499] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1553258089] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4132576 kB) free!
[1553258099] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553258209] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.7% (12182180 kB) free.
[1553258219] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553258339] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553258809] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (3387868 kB) free!
[1553258929] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11748840 kB) free.
[1553259529] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6483980 kB) free!
[1553259539] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553259649] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11565660 kB) free.
[1553260140] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553260143] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 12.4% (2022004 kB) free!
[1553260260] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1553260263] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 12.4% (2015080 kB) free!
[1553260383] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 12.4% (2007656 kB) free!
[1553260602] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1553260722] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1553260843] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1553260849] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (228412 kB) free!
[1553260860] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1553260970] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (214204 kB) free!
[1553260980] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553260984] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161472 kB) free!
[1553261090] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12508656 kB) free.
[1553261100] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1553261100] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553261443] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 7 users currently logged in
[1553261690] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3948264 kB) free!
[1553261810] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11690432 kB) free.
[1553262299] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 587 processes
[1553262359] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 590 processes
[1553263010] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4098504 kB) free!
[1553263130] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.6% (12108648 kB) free.
[1553264330] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (4399972 kB) free!
[1553264450] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12323768 kB) free.
[1553265050] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3465484 kB) free!
[1553265170] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11475620 kB) free.
[1553266370] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4139708 kB) free!
[1553266490] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12387496 kB) free.
[1553266500] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1553267090] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5594404 kB) free!
[1553267100] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553267210] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (3594352 kB) free!
[1553267220] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553267330] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12387864 kB) free.
[1553267340] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553267930] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6415176 kB) free!
[1553267941] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553268050] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.2% (5043020 kB) free!
[1553268061] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553268170] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12384296 kB) free.
[1553268181] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1553268770] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7317804 kB) free!
[1553268781] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553268890] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6086684 kB) free!
[1553268901] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553269011] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4487220 kB) free!
[1553269021] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553269611] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.5% (8132328 kB) free!
[1553270211] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12379004 kB) free.
[1553270221] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1553270811] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (4784232 kB) free!
[1553270821] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553270931] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12385036 kB) free.
[1553270942] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1553271531] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7008196 kB) free!
[1553271541] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553271651] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.7% (5754988 kB) free!
[1553271661] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553271771] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1553271771] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3836616 kB) free!
[1553271781] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553272371] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.7% (7750916 kB) free!
[1553272971] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12381104 kB) free.
[1553272981] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553273571] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4259452 kB) free!
[1553273581] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553273691] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12377216 kB) free.
[1553273701] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1553274292] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6671296 kB) free!
[1553274302] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553274412] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.9% (5391092 kB) free!
[1553274422] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553274532] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12377980 kB) free.
[1553274542] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1553275132] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7393756 kB) free!
[1553275142] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553275252] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.4% (6107796 kB) free!
[1553275260] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 585 processes
[1553275262] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553275320] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 589 processes
[1553275372] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4522660 kB) free!
[1553275382] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553275673] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1553275793] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1553275913] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4919840 kB) free!
[1553275913] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1553275972] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.2% (6527360 kB) free!
[1553276033] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (4919180 kB) free!
[1553276153] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (4918496 kB) free!
[1553277353] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 20.3% (5021344 kB) free.
[1553277772] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10576952 kB) free.
[1553278913] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1553278972] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (3184156 kB) free!
[1553278982] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1553279093] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11521368 kB) free.
[1553279216] SERVICE ALERT: sig1;Memory;WARNING;SOFT;1;WARNING - 20.0% (3257412 kB) free!
[1553279336] SERVICE ALERT: sig1;Memory;WARNING;SOFT;2;WARNING - 20.0% (3257792 kB) free!