[1547756364] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.02, 9.99, 10.04
[1547756485] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 9.44, 9.85, 9.98
[1547757084] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.21, 10.44, 10.20
[1547757204] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.07, 10.39, 10.22
[1547757324] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.31, 10.35, 10.22
[1547759001] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 20.9% (5127480 kB) free.
[1547759111] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.04, 11.75, 6.95
[1547759132] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 7.78, 8.79, 9.68
[1547759231] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.70, 8.29, 6.27
[1547759289] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 337 processes
[1547759351] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.41, 6.03, 5.68
[1547759951] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.37, 1.92, 3.58
[1547763489] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 337 processes
[1547777601] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4823676 kB) free!
[1547777722] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (4823464 kB) free!
[1547777842] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (4822912 kB) free!
[1547781549] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547781669] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1547781789] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1547781932] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.10, 10.62, 10.38
[1547782052] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.68, 10.26, 10.28
[1547782172] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.47, 10.34, 10.30
[1547786589] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547786709] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1547786889] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 6 warn out of 345 processes
[1547786951] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.68, 6.02, 5.70
[1547787071] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.50, 4.51, 5.19
[1547787191] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.70, 3.58, 4.77
[1547787572] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 6.85, 7.41, 9.40
[1547787909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547788029] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1547788629] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547788749] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547788870] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547788971] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.2% (11932408 kB) free.
[1547789355] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.7% (5258244 kB) free!
[1547790290] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.9% (2908088 kB) free!
[1547790410] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.1% (2779144 kB) free!
[1547790530] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.2% (2634596 kB) free!
[1547791642] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (1998348 kB) free!
[1547792242] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.1% (3466436 kB) free!
[1547792289] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 346 processes
[1547792930] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1077764 kB) free!
[1547793443] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2284644 kB) free!
[1547793530] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 12.1% (1961220 kB) free!
[1547793563] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547793683] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1547794155] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (3189560 kB) free!
[1547794772] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.59, 10.83, 10.02
[1547794870] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1547794892] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 10.10, 10.34, 9.93
[1547795330] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1547795330] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1594028 kB) free!
[1547795492] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.71, 10.54, 10.29
[1547795571] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9644576 kB) free!
[1547795591] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46813 MB (20% inode=99%):
[1547795612] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.24, 10.88, 10.45
[1547795691] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 19.5% (9589656 kB) free!
[1547795711] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46797 MB (20% inode=99%):
[1547795732] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.32, 10.65, 10.42
[1547795811] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.4% (9537180 kB) free!
[1547795831] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46797 MB (20% inode=99%):
[1547796004] HOST ALERT: sig1;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1547796046] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1547796070] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547796190] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547796310] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547796489] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 352 processes
[1547797089] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1547797089] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 3 warn out of 351 processes
[1547797643] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1547797643] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.5% (2579280 kB) free!
[1547798243] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2132064 kB) free!
[1547799683] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547799803] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1547799910] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1547799910] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1547800510] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547800576] SERVICE ALERT: sig3;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547800576] SERVICE ALERT: sig3;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547800606] HOST ALERT: sig3;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1547800630] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547800680] SERVICE ALERT: sig3;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig3' was out of bounds)
[1547800683] SERVICE ALERT: sig3;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig3' was out of bounds)
[1547800696] SERVICE ALERT: sig3;HDD_Root status;OK;HARD;1;DISK OK - free space: / 37910 MB (74% inode=98%):
[1547800696] SERVICE ALERT: sig3;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1547800700] HOST ALERT: sig3;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1547800750] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547800755] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6084756 kB) free!
[1547800800] SERVICE ALERT: sig3;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1547800803] SERVICE ALERT: sig3;Total Processes;OK;HARD;1;PROCS OK: 313 processes
[1547801132] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 9.71, 9.51, 9.99
[1547801355] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3839676 kB) free!
[1547801956] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1547801956] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.5% (5684648 kB) free!
[1547802556] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (3469812 kB) free!
[1547802701] SERVICE ALERT: sig1;Memory;CRITICAL;SOFT;1;(Return code of 255 for service 'Memory' on host 'sig1' was out of bounds)
[1547802723] SERVICE ALERT: sig1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547802753] HOST ALERT: sig1;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1547802808] SERVICE ALERT: sig2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547802808] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547802808] SERVICE ALERT: sig2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547802811] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU_Procs' on host 'sig2' was out of bounds)
[1547802817] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 39.42 ms
[1547802821] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 96.5% (15685452 kB) free.
[1547802829] HOST ALERT: sig2;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1547802837] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1547802843] SERVICE ALERT: sig1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 37145 MB (72% inode=98%):
[1547802928] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1547802928] SERVICE ALERT: sig2;Total Processes;OK;HARD;1;PROCS OK: 323 processes
[1547802928] SERVICE ALERT: sig2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38716 MB (75% inode=98%):
[1547802931] SERVICE ALERT: sig2;CPU_Procs;OK;HARD;1;CPU OK: 319 processes
[1547803532] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.76, 10.43, 10.22
[1547803652] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.56, 10.37, 10.22
[1547803772] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 11.03, 10.60, 10.32
[1547804356] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.7% (5748188 kB) free!
[1547804552] SERVICE ALERT: sig6;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (1288388 kB) free!
[1547804672] SERVICE ALERT: sig6;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (1288196 kB) free!
[1547804792] SERVICE ALERT: sig6;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1288912 kB) free!
[1547804956] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3405708 kB) free!
[1547805530] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1547805804] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1547805924] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1547806156] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.4% (5598924 kB) free!
[1547807950] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1547808489] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1547808551] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547808556] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4492948 kB) free!
[1547808671] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547808791] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1547809156] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.6% (6192412 kB) free!
[1547809643] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1547810356] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (3318764 kB) free!
[1547810738] SERVICE ALERT: sig2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547810738] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547810738] SERVICE ALERT: sig2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547810741] HOST ALERT: sig2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.150)
[1547810741] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig2' was out of bounds)
[1547810742] SERVICE ALERT: sig1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547810744] HOST ALERT: sig1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.44)
[1547810744] SERVICE ALERT: sig1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig1' was out of bounds)
[1547810761] SERVICE ALERT: sig1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sig1' was out of bounds)
[1547810858] SERVICE ALERT: sig2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38719 MB (75% inode=98%):
[1547810858] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1547810858] SERVICE ALERT: sig2;Total Processes;OK;HARD;1;PROCS OK: 325 processes
[1547810861] SERVICE ALERT: sig2;CPU_Procs;WARNING;HARD;1;CPU WARNING: 2 warn out of 326 processes
[1547810862] SERVICE ALERT: sig1;CPU_Procs;WARNING;HARD;1;CPU WARNING: 1 warn out of 308 processes
[1547810862] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1547810864] SERVICE ALERT: sig1;Sensors Status;OK;HARD;1;SENSORS OK
[1547810868] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1547810881] SERVICE ALERT: sig1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 196952 MB (87% inode=99%):
[1547811461] SERVICE ALERT: sig2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1547811462] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1547811556] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6173256 kB) free!
[1547812156] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (3994764 kB) free!
[1547812756] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.6% (6185532 kB) free!
[1547812845] SERVICE ALERT: sig1;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'sig1' was out of bounds)
[1547812845] SERVICE ALERT: sig1;CPU Load;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU Load' on host 'sig1' was out of bounds)
[1547812848] HOST ALERT: sig1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.44)
[1547812855] SERVICE ALERT: sig1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig1' was out of bounds)
[1547812912] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1547812965] SERVICE ALERT: sig1;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1547812965] SERVICE ALERT: sig1;CPU Load;OK;HARD;1;OK - load average: 0.57, 0.27, 0.10
[1547812975] SERVICE ALERT: sig1;Total Processes;OK;HARD;1;PROCS OK: 324 processes
[1547813356] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4452676 kB) free!
[1547813956] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6175824 kB) free!
[1547814191] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1547814556] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4752328 kB) free!
[1547815156] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.6% (6188004 kB) free!
[1547815991] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547816111] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547816231] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1547816906] SERVICE ALERT: sig3;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547816906] SERVICE ALERT: sig3;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547816925] HOST ALERT: sig3;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1547816989] HOST ALERT: sig3;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1547817000] SERVICE ALERT: sig3;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 325 processes
[1547817026] SERVICE ALERT: sig3;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1547817026] SERVICE ALERT: sig3;HDD_Root status;OK;HARD;1;DISK OK - free space: / 37903 MB (74% inode=98%):
[1547817120] SERVICE ALERT: sig3;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1547817191] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 29.84, 21.77, 9.65
[1547817312] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.56, 15.15, 8.69
[1547817432] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.68, 10.85, 7.90
[1547818062] SERVICE ALERT: sig1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1547818172] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 8.01, 8.56, 9.86
[1547818182] SERVICE ALERT: sig1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1547818302] SERVICE ALERT: sig1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 320 processes
[1547818632] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.05, 2.14, 3.61
[1547820765] SERVICE ALERT: sig1;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'sig1' was out of bounds)
[1547820765] SERVICE ALERT: sig1;CPU Load;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU Load' on host 'sig1' was out of bounds)
[1547820885] SERVICE ALERT: sig1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1547820885] SERVICE ALERT: sig1;CPU Load;OK;HARD;1;OK - load average: 0.21, 0.19, 0.08
[1547821172] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.89, 10.80, 10.14
[1547821292] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.60, 10.36, 10.06
[1547821302] SERVICE FLAPPING ALERT: sig1;CPU_Procs;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1547821302] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1547821412] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.61, 10.39, 10.10
[1547821632] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547821752] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547821756] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4312480 kB) free!
[1547821872] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547822356] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.9% (5879116 kB) free!
[1547822724] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1547822844] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1547822956] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4502884 kB) free!
[1547823295] SERVICE ALERT: sig1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547823295] SERVICE ALERT: sig1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547823325] HOST ALERT: sig1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1547823415] SERVICE ALERT: sig1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1547823415] SERVICE ALERT: sig1;CPU Load;OK;HARD;1;OK - load average: 0.78, 0.20, 0.07
[1547823419] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1547823556] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6163476 kB) free!
[1547824157] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4270948 kB) free!
[1547824272] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1547824757] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6160936 kB) free!
[1547825244] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1547825364] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1547825502] SERVICE ALERT: sig1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 314 processes
[1547825622] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1547825957] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4642480 kB) free!
[1547826073] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547826193] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1547826557] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6177648 kB) free!
[1547827393] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547827513] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1547828290] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 6 warn out of 352 processes
[1547828890] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 4 warn out of 361 processes
[1547830422] SERVICE FLAPPING ALERT: sig1;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1547830758] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4292560 kB) free!
[1547831269] SERVICE ALERT: smscapp1;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1547831358] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.8% (5324708 kB) free!
[1547832002] SERVICE ALERT: sig2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547832005] HOST ALERT: sig2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.150)
[1547832009] SERVICE ALERT: sig2;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sig2' was out of bounds)
[1547832122] SERVICE ALERT: sig2;CPU Load;OK;HARD;1;OK - load average: 0.50, 0.17, 0.06
[1547832126] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1547832129] SERVICE ALERT: sig2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 211837 MB (94% inode=99%):
[1547833545] SERVICE ALERT: sig3;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547833551] SERVICE ALERT: sig3;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547833575] HOST ALERT: sig3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1547833605] HOST ALERT: sig1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.44)
[1547833618] SERVICE ALERT: sig1;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig1' was out of bounds)
[1547833618] SERVICE ALERT: sig1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig1' was out of bounds)
[1547833639] HOST ALERT: sig3;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1547833652] HOST ALERT: sig2;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1547833662] SERVICE ALERT: sig2;CPU_Procs;WARNING;HARD;1;CPU WARNING: 1 warn out of 331 processes
[1547833662] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1547833665] SERVICE ALERT: sig3;Memory;OK;HARD;1;OK - 94.8% (15401332 kB) free.
[1547833668] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1547833671] SERVICE ALERT: sig3;CPU Load;OK;HARD;1;OK - load average: 0.36, 0.14, 0.05
[1547833738] SERVICE ALERT: sig1;CPU Load;OK;HARD;1;OK - load average: 0.23, 0.15, 0.06
[1547833738] SERVICE ALERT: sig1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1547833782] SERVICE ALERT: sig2;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1547834258] SERVICE ALERT: sig2;Total Processes;CRITICAL;SOFT;1;(Return code of 255 for service 'Total Processes' on host 'sig2' was out of bounds)
[1547834258] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig2' was out of bounds)
[1547834274] SERVICE ALERT: sig1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547834284] SERVICE ALERT: sig1;HDD_Home status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Home status' on host 'sig1' was out of bounds)
[1547834294] HOST ALERT: sig1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1547834338] SERVICE ALERT: sig1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig1' was out of bounds)
[1547834338] SERVICE ALERT: sig1;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig1' was out of bounds)
[1547834358] HOST FLAPPING ALERT: sig1;STARTED; Host appears to have started flapping (20.1% change > 20.0% threshold)
[1547834358] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1547834381] SERVICE ALERT: sig2;Total Processes;CRITICAL;SOFT;2;(Return code of 255 for service 'Total Processes' on host 'sig2' was out of bounds)
[1547834381] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig2' was out of bounds)
[1547834383] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU_Procs' on host 'sig2' was out of bounds)
[1547834394] SERVICE ALERT: sig1;Sensors Status;OK;HARD;1;SENSORS OK
[1547834404] SERVICE ALERT: sig1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 195330 MB (86% inode=99%):
[1547834501] SERVICE ALERT: sig2;Total Processes;OK;HARD;1;PROCS OK: 327 processes
[1547834501] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1547834503] SERVICE ALERT: sig2;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1547835004] SERVICE ALERT: sig1;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547835010] HOST ALERT: sig1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.44)
[1547835013] HOST ALERT: sig2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1547835014] SERVICE ALERT: sig1;HDD_Home status;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547835077] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1547835124] SERVICE ALERT: sig1;Sensors Status;OK;HARD;1;SENSORS OK
[1547835128] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1547835134] SERVICE ALERT: sig1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 195330 MB (86% inode=99%):
[1547835313] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547835433] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1547835496] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 327 processes
[1547835658] SERVICE FLAPPING ALERT: sig1;Current Users;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1547835658] SERVICE ALERT: sig1;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1547835658] SERVICE FLAPPING ALERT: sig1;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1547835658] SERVICE ALERT: sig1;CPU Load;OK;HARD;1;OK - load average: 0.00, 0.03, 0.05
[1547835775] HOST FLAPPING ALERT: sig1;STOPPED; Host appears to have stopped flapping (4.5% change < 5.0% threshold)