[1547664081] SERVICE ALERT: sig1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1547664083] SERVICE FLAPPING ALERT: db1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1547664202] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1547664919] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 39.7% (9742776 kB) free.
[1547665089] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 343 processes
[1547666224] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1547668156] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1547670489] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 342 processes
[1547672450] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1547690930] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.0% (6424836 kB) free!
[1547690955] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4499008 kB) free!
[1547691551] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48332 MB (21% inode=99%):
[1547693331] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.2% (9966868 kB) free.
[1547693355] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.2% (5024024 kB) free!
[1547693719] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.2% (2980440 kB) free!
[1547693839] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.0% (2950340 kB) free!
[1547693959] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.0% (2948804 kB) free!
[1547696825] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547696945] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547696955] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (3374848 kB) free!
[1547697065] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1547697555] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1547697555] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.2% (5028548 kB) free!
[1547698031] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.83, 10.48, 7.71
[1547698036] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 7934 MB (15% inode=98%):
[1547698089] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 340 processes
[1547698151] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.22, 7.39, 6.92
[1547698159] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 36.6% (8966132 kB) free.
[1547698271] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.39, 5.41, 6.25
[1547698581] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1547698755] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3713608 kB) free!
[1547698865] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547698871] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.27, 1.83, 3.88
[1547698985] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 298 processes
[1547699105] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1547699955] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.6% (5248132 kB) free!
[1547701505] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547701625] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1547701745] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1547702356] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4726692 kB) free!
[1547702945] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1547703065] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1547703083] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1547703186] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547703203] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1547703489] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 5 warn out of 369 processes
[1547704131] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9810952 kB) free!
[1547704251] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 19.9% (9823712 kB) free!
[1547704371] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.9% (9784152 kB) free!
[1547705360] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4809252 kB) free!
[1547705480] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (4760244 kB) free!
[1547705600] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (4663992 kB) free!
[1547706186] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1547706203] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1547706323] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1547706921] HOST ALERT: sig2;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1547707048] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1547707386] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547707506] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1547707523] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1547707643] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1547707751] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46880 MB (20% inode=99%):
[1547707871] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46872 MB (20% inode=99%):
[1547707991] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47108 MB (20% inode=99%):
[1547708053] SERVICE ALERT: sig4;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 308 processes
[1547708173] SERVICE ALERT: sig4;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1547708773] SERVICE ALERT: sig4;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1547708893] SERVICE ALERT: sig4;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1547709200] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2371468 kB) free!
[1547709306] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547709391] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547709411] HOST ALERT: sig1;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1547709426] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547709459] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1547709511] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1547709546] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547709631] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1547709649] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (3231440 kB) free!
[1547709769] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.8% (3211188 kB) free!
[1547709889] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.6% (3188452 kB) free!
[1547711243] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1547711363] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1547711956] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1547711963] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1547712083] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1547712458] SERVICE ALERT: sig1;Memory;CRITICAL;SOFT;1;(Return code of 255 for service 'Memory' on host 'sig1' was out of bounds)
[1547712458] HOST ALERT: sig1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.44)
[1547712522] SERVICE ALERT: sig1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig1' was out of bounds)
[1547712578] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 94.7% (15389212 kB) free.
[1547712582] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1547712642] SERVICE ALERT: sig1;CPU Load;OK;HARD;1;OK - load average: 0.14, 0.12, 0.05
[1547713146] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1547713146] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1547713156] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.6% (5242420 kB) free!
[1547713746] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1547713756] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (3101792 kB) free!
[1547713781] SERVICE ALERT: sig1;Memory;CRITICAL;SOFT;1;(Return code of 255 for service 'Memory' on host 'sig1' was out of bounds)
[1547713781] HOST ALERT: sig1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.44)
[1547713793] SERVICE ALERT: sig1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig1' was out of bounds)
[1547713846] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1547713866] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1547713901] SERVICE ALERT: sig1;Memory;OK;HARD;1;OK - 94.9% (15429804 kB) free.
[1547713913] SERVICE ALERT: sig1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 37147 MB (72% inode=98%):
[1547714466] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1547714586] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1547715083] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1547715186] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1547715203] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1547715306] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1547715889] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (1424216 kB) free!
[1547715907] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547716027] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547716147] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547716415] SERVICE ALERT: sig3;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547716421] SERVICE ALERT: sig3;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547716445] HOST ALERT: sig3;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1547716446] SERVICE ALERT: sig3;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547716446] SERVICE ALERT: sig3;HDD_Root status;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547716535] SERVICE ALERT: sig3;Memory;OK;HARD;1;OK - 94.5% (15364068 kB) free.
[1547716539] HOST ALERT: sig3;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1547716541] SERVICE ALERT: sig3;CPU Load;OK;HARD;1;OK - load average: 0.42, 0.13, 0.05
[1547716546] HOST ALERT: sig1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.44)
[1547716566] SERVICE ALERT: sig3;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1547716566] SERVICE ALERT: sig3;HDD_Root status;OK;HARD;1;DISK OK - free space: / 37912 MB (74% inode=98%):
[1547716610] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1547717003] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547717123] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1547718211] SERVICE ALERT: sig1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547718214] HOST ALERT: sig1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.44)
[1547718218] SERVICE ALERT: sig1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig1' was out of bounds)
[1547718238] SERVICE ALERT: sig1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sig1' was out of bounds)
[1547718278] HOST ALERT: sig1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1547718332] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1547718338] SERVICE ALERT: sig1;Sensors Status;OK;HARD;1;SENSORS OK
[1547718358] SERVICE ALERT: sig1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 197201 MB (87% inode=99%):
[1547720348] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1547720948] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1547721068] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547721188] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1547723588] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547723708] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1547724200] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (4182480 kB) free!
[1547724229] SERVICE ALERT: smscapp1;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1547724908] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547725028] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547725132] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.00, 0.49, 3.48
[1547725148] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547725748] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1547726295] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 332 processes
[1547726948] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547727068] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547727189] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547727800] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2224720 kB) free!
[1547728389] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1547728401] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (4800040 kB) free!
[1547728429] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1547728523] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1547728549] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1547728643] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1547728669] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1547731696] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 330 processes
[1547732289] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 348 processes
[1547732296] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 328 processes
[1547732364] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 6.29, 7.51, 9.54
[1547732471] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.22, 6.33, 7.42
[1547732591] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.50, 4.78, 6.72
[1547732711] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.38, 3.65, 6.07
[1547733311] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.34, 1.69, 3.85
[1547734878] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547734878] SERVICE ALERT: sig2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547734878] SERVICE ALERT: sig2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547734881] HOST ALERT: sig2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.150)
[1547734888] SERVICE ALERT: sig2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig2' was out of bounds)
[1547734908] SERVICE ALERT: sig2;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig2' was out of bounds)
[1547734919] SERVICE ALERT: sig2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig2' was out of bounds)
[1547734998] SERVICE ALERT: sig2;Total Processes;OK;HARD;1;PROCS OK: 322 processes
[1547734998] SERVICE ALERT: sig2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1547734998] SERVICE ALERT: sig2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38716 MB (75% inode=98%):
[1547735002] HOST ALERT: sig2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1547735028] SERVICE ALERT: sig2;Sensors Status;OK;HARD;1;SENSORS OK
[1547735040] SERVICE ALERT: sig2;Memory;OK;HARD;1;OK - 94.4% (15339068 kB) free.
[1547735589] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1547735609] SERVICE ALERT: sig2;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1547735709] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1547736555] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.1% (4962212 kB) free!
[1547736909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1547737029] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1547737089] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 342 processes
[1547737155] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4899076 kB) free!
[1547738843] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1547738964] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1547739429] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1547739791] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53345 MB (23% inode=99%):
[1547739889] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 22.1% (3599264 kB) free.
[1547740629] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547740750] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes