[1587753528] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7481676 kB) free!
[1587753648] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (3935824 kB) free!
[1587753768] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34432264 kB) free.
[1587754368] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10919128 kB) free!
[1587754489] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (7169636 kB) free!
[1587754609] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29545340 kB) free.
[1587755111] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (178144 kB) free!
[1587755209] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14328684 kB) free!
[1587755329] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (10360196 kB) free!
[1587755449] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (6629660 kB) free!
[1587755711] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 71.8% (17602528 kB) free.
[1587755993] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (1346920 kB) free!
[1587756049] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (17011728 kB) free!
[1587756113] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (892324 kB) free!
[1587756233] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (413280 kB) free!
[1587756649] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29457128 kB) free.
[1587757249] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8210908 kB) free!
[1587757369] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31332936 kB) free.
[1587757969] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14384500 kB) free!
[1587758089] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (10498076 kB) free!
[1587758209] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (6908084 kB) free!
[1587758289] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (224464 kB) free!
[1587758409] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (225708 kB) free!
[1587758452] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (1962856 kB) free!
[1587758530] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (223588 kB) free!
[1587758572] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (1971656 kB) free!
[1587758692] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (1764456 kB) free!
[1587758809] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16575192 kB) free!
[1587759409] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27387912 kB) free.
[1587760009] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7329112 kB) free!
[1587760129] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34882712 kB) free.
[1587760330] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.6% (7702488 kB) free!
[1587760730] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (12002188 kB) free!
[1587760850] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8520200 kB) free!
[1587760970] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35808140 kB) free.
[1587761570] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13785964 kB) free!
[1587761690] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (10083496 kB) free!
[1587761810] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32221224 kB) free.
[1587762233] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1587762410] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15189720 kB) free!
[1587762478] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 83.9% (41333108 kB) free.
[1587762530] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (11557396 kB) free!
[1587762650] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (8235484 kB) free!
[1587762898] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.9% (45266976 kB) free.
[1587763250] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16310312 kB) free!
[1587763874] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(316) < 49.50.86.152 [25/Apr/2020:03:00:41 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1587763925] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2460) < 203.212.70.170 [25/Apr/2020:03:01:52 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1587763931] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.5% (44611748 kB) free.
[1587763994] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587764033] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.3% (2486284 kB) free!
[1587764045] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587764047] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.5% (4038916 kB) free!
[1587764070] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1587764120] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(387) < 2020/04/25 03:01:32 [error] 22133#0: *85651798 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1587764241] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587764450] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35613952 kB) free.
[1587764647] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (299924 kB) free!
[1587765050] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (12998684 kB) free!
[1587765170] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (9492528 kB) free!
[1587765290] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35560572 kB) free.
[1587765870] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1587765883] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1587765891] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7783456 kB) free!
[1587765990] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1587766003] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 349 processes
[1587766011] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (3891636 kB) free!
[1587766110] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1587766123] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 353 processes
[1587766131] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35454700 kB) free.
[1587766731] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10732304 kB) free!
[1587766851] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (6734528 kB) free!
[1587766971] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18280284 kB) free!
[1587767048] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.9% (10016136 kB) free.
[1587767256] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.1% (21344508 kB) free.
[1587767310] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1587767323] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1587767571] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (10846504 kB) free!
[1587767648] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.6% (2594496 kB) free!
[1587767768] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (295068 kB) free!
[1587767888] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1587767888] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (290556 kB) free!
[1587768171] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20401184 kB) free!
[1587768233] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.3% (3469840 kB) free.
[1587768771] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33052408 kB) free.
[1587769371] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11701936 kB) free!
[1587769491] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11075396 kB) free!
[1587769611] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35221796 kB) free.
[1587770211] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16278984 kB) free!
[1587770288] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 37.2% (9112492 kB) free.
[1587770331] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37570956 kB) free.
[1587770931] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19333372 kB) free!
[1587771051] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (15900192 kB) free!
[1587771172] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36171008 kB) free.
[1587771488] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179876 kB) free!
[1587771608] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172936 kB) free!
[1587771728] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180788 kB) free!
[1587771771] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17660892 kB) free!
[1587771891] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37599568 kB) free.
[1587772222] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 77.3% (18943440 kB) free.
[1587772430] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 44.9% (11076900 kB) free.
[1587772491] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19575200 kB) free!
[1587772556] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 77.8% (19072148 kB) free.
[1587772611] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.4% (16392564 kB) free!
[1587772731] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38077104 kB) free.
[1587773331] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (20478876 kB) free!
[1587773451] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39702376 kB) free.
[1587774051] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (22661252 kB) free!
[1587774171] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.4% (18976404 kB) free!
[1587774291] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38572988 kB) free.
[1587774891] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20291252 kB) free!
[1587775011] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39532136 kB) free.
[1587775430] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (2439152 kB) free!
[1587775550] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (2422920 kB) free!
[1587775611] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21743860 kB) free!
[1587775670] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2395696 kB) free!
[1587775731] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.6% (17948768 kB) free!
[1587775852] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38351324 kB) free.
[1587776092] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 33.2% (43827348 kB) free.
[1587776451] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19841368 kB) free!
[1587776571] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39310668 kB) free.
[1587777171] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19392564 kB) free!
[1587777291] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39191780 kB) free.
[1587777891] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19271884 kB) free!
[1587778011] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39063700 kB) free.
[1587778611] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18589084 kB) free!
[1587778731] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38836668 kB) free.
[1587779331] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20556756 kB) free!
[1587779451] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.9% (17011116 kB) free!
[1587779571] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37513928 kB) free.
[1587780171] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19060520 kB) free!
[1587780291] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38591032 kB) free.
[1587780891] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20315804 kB) free!
[1587781011] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32868988 kB) free.
[1587781612] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21354040 kB) free!
[1587781732] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.4% (17657880 kB) free!
[1587781852] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37281488 kB) free.
[1587782452] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15564276 kB) free!
[1587782528] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1587782572] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37917028 kB) free.
[1587783172] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (17999096 kB) free!
[1587783292] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1587783292] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37607016 kB) free.
[1587783892] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17636272 kB) free!
[1587784012] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.4% (13751692 kB) free!
[1587784132] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35255932 kB) free.
[1587784732] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14303528 kB) free!
[1587784852] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35995656 kB) free.
[1587785452] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14481144 kB) free!
[1587785502] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.24, 6.60, 5.22
[1587785572] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35989192 kB) free.
[1587785622] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.68, 7.26, 5.63
[1587785742] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.03, 7.94, 6.05
[1587786172] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13766204 kB) free!
[1587786292] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9742044 kB) free!
[1587786356] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.7% (3602444 kB) free!
[1587786413] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33800384 kB) free.
[1587786476] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.3% (3517024 kB) free!
[1587786596] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.0% (3428568 kB) free!
[1587787012] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10850572 kB) free!
[1587787132] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34146092 kB) free.
[1587787732] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11802840 kB) free!
[1587787852] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7705340 kB) free!
[1587787972] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31386952 kB) free.
[1587788033] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3229160 kB) free!
[1587788153] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (2987264 kB) free!
[1587788274] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (3195488 kB) free!
[1587788396] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2039176 kB) free!
[1587788572] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8757312 kB) free!
[1587788692] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32365868 kB) free.
[1587789292] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8815972 kB) free!
[1587789412] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4285468 kB) free!
[1587789532] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30516088 kB) free.
[1587790132] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6537656 kB) free!
[1587790252] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2253456 kB) free!
[1587790372] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27870412 kB) free.
[1587790540] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 20.0% (4902572 kB) free!
[1587790660] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 20.0% (4900724 kB) free!
[1587790780] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 20.0% (4900472 kB) free!
[1587790972] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3247576 kB) free!
[1587791093] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28895632 kB) free.
[1587791692] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4178312 kB) free!
[1587791720] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1587791812] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (399116 kB) free!
[1587791840] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1587791932] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27653292 kB) free.
[1587792532] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (3212828 kB) free!
[1587792652] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (405256 kB) free!
[1587792772] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.8% (23501404 kB) free!
[1587793074] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (193048 kB) free!
[1587793238] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.4% (4141724 kB) free!
[1587793358] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.2% (3859196 kB) free!
[1587793372] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (404532 kB) free!
[1587793478] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.9% (3579448 kB) free!
[1587793674] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (2530980 kB) free!
[1587793972] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16268468 kB) free!
[1587794045] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(359) < 103.255.146.164 [25/Apr/2020:11:17:31 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587794078] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2031488 kB) free!
[1587794165] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587794256] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.6% (4319856 kB) free!
[1587794376] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.0% (4165884 kB) free!
[1587794496] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.5% (4043236 kB) free!
[1587794846] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.10296.5420> (10.147.212.53:34724 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1587794966] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1587795172] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1587795172] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28080016 kB) free.
[1587795365] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(302) < 103.255.146.164 [25/Apr/2020:11:45:40 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587795485] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587795772] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4562792 kB) free!
[1587795892] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31320956 kB) free.
[1587796297] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1953208 kB) free!
[1587796478] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.9% (16376960 kB) free.
[1587796492] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7289172 kB) free!
[1587796612] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2758544 kB) free!
[1587796733] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30883552 kB) free.
[1587797274] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1587797274] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1618124 kB) free!
[1587797332] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5441800 kB) free!
[1587797452] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (802208 kB) free!
[1587797572] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32449060 kB) free.
[1587797874] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.3% (2158448 kB) free!
[1587798172] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7804232 kB) free!
[1587798292] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (3096212 kB) free!
[1587798412] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32561056 kB) free.
[1587798474] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (1491764 kB) free!
[1587798485] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(219) < 103.255.146.164 [25/Apr/2020:12:31:17 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587798605] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587799012] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8413156 kB) free!
[1587799074] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.7% (1740200 kB) free!
[1587799132] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3713348 kB) free!
[1587799252] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383780 kB) free!
[1587799805] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(110) < 103.255.146.164 [25/Apr/2020:12:56:57 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587799925] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587800452] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26524624 kB) free.
[1587800526] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(101) < 103.255.146.164 [25/Apr/2020:13:05:33 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587800646] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587801052] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (550424 kB) free!
[1587801172] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35932068 kB) free.
[1587801292] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.2% (23956236 kB) free!
[1587801412] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.1% (27832944 kB) free.
[1587801474] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (1537540 kB) free!
[1587801772] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8143232 kB) free!
[1587801893] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2292924 kB) free!
[1587802013] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39559364 kB) free.
[1587802613] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7121540 kB) free!
[1587802733] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1882896 kB) free!
[1587802853] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39296324 kB) free.
[1587803046] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(322) < 103.255.146.164 [25/Apr/2020:13:45:21 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587803167] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587803453] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10268084 kB) free!
[1587803573] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (5536564 kB) free!
[1587803693] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38834376 kB) free.
[1587803742] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 5.33, 4.66, 4.86
[1587804293] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12913524 kB) free!
[1587804342] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.74, 5.50, 5.54
[1587804366] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(174) < 103.255.146.164 [25/Apr/2020:14:11:19 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587804413] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7940160 kB) free!
[1587804462] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.99, 7.38, 6.21
[1587804486] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587804533] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21470792 kB) free!
[1587804583] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.63, 7.94, 6.57
[1587805133] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (11935712 kB) free!
[1587805733] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.1% (22567088 kB) free!
[1587806287] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(120) < 103.255.146.164 [25/Apr/2020:14:44:01 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587806333] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36376220 kB) free.
[1587806381] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.6% (2354244 kB) free!
[1587806407] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587806678] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.2% (3857072 kB) free!
[1587806698] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9097268 kB) free!
[1587806798] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.4% (3459604 kB) free!
[1587806812] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.5% (25665248 kB) free!
[1587806818] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8854588 kB) free!
[1587806874] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.7% (2230692 kB) free!
[1587806918] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.7% (3080920 kB) free!
[1587806932] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.2% (29264908 kB) free.
[1587806933] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7376412 kB) free!
[1587806938] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.5% (8602696 kB) free!
[1587807007] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(126) < 185.255.8.58 [25/Apr/2020:15:00:04 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1587807053] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (2511676 kB) free!
[1587807127] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587807173] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391176 kB) free!
[1587807518] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.7% (16333884 kB) free.
[1587807727] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(568) < 103.255.146.164 [25/Apr/2020:15:11:33 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587807847] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587808074] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1528184 kB) free!
[1587808132] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.9% (21023484 kB) free!
[1587808252] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.3% (25471856 kB) free!
[1587808372] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.5% (28366124 kB) free.
[1587808373] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.3% (24057824 kB) free!
[1587808822] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4743872 kB) free!
[1587808942] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (4716984 kB) free!
[1587808973] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.2% (21373028 kB) free!
[1587808973] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394324 kB) free!
[1587809062] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (4689880 kB) free!
[1587809093] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.1% (22494196 kB) free!
[1587809213] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.6% (29801688 kB) free.
[1587810173] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.6% (24527848 kB) free!
[1587810247] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(148) < 103.255.146.164 [25/Apr/2020:15:47:49 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587810367] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587810538] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4804132 kB) free!
[1587810773] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34725040 kB) free.
[1587811373] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7355008 kB) free!
[1587811493] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2656032 kB) free!
[1587811613] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34773744 kB) free.
[1587811683] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.8% (2409544 kB) free!
[1587812213] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7642772 kB) free!
[1587812283] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1239648 kB) free!
[1587812333] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (3123732 kB) free!
[1587812453] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390804 kB) free!
[1587812813] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.0% (23710840 kB) free!
[1587812883] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.9% (2100328 kB) free!
[1587812933] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.9% (30137416 kB) free.
[1587813654] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13711344 kB) free!
[1587814083] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1521596 kB) free!
[1587814183] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.38, 3.58, 4.78
[1587814567] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(496) < 103.255.146.164 [25/Apr/2020:17:01:26 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1587814687] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587814854] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397260 kB) free!
[1587815283] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.0% (1795252 kB) free!
[1587815383] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.05, 6.46, 5.42
[1587815454] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15330512 kB) free!
[1587815503] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.75, 6.22, 5.46
[1587815623] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.88, 6.12, 5.52
[1587815933] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.8% (24775816 kB) free!
[1587816053] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.0% (30342496 kB) free.
[1587816054] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27553172 kB) free.
[1587816483] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (509660 kB) free!
[1587816654] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (1401620 kB) free!
[1587816774] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36353448 kB) free.
[1587817083] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (2763776 kB) free!
[1587817374] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7472708 kB) free!
[1587817494] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2828648 kB) free!
[1587817614] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386688 kB) free!
[1587818283] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1566668 kB) free!
[1587818453] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 20.0% (26414700 kB) free!
[1587818573] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.6% (27141704 kB) free.
[1587818814] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.4% (24287008 kB) free!
[1587818887] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(593) < 103.255.146.164 [25/Apr/2020:18:15:32 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1587819008] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1587819414] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36185436 kB) free.
[1587819823] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 6.37, 5.15, 4.91
[1587819862] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2408980 kB) free!
[1587820014] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8653972 kB) free!
[1587820083] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.3% (2156024 kB) free!
[1587820134] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (3917292 kB) free!
[1587820255] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (376168 kB) free!
[1587821023] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.08, 8.97, 6.56
[1587821143] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.57, 9.80, 7.16
[1587821263] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1587821263] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 12.91, 10.86, 7.86
[1587821455] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.2% (24040860 kB) free!
[1587821883] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9747024 kB) free!
[1587822003] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (9543664 kB) free!
[1587822055] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388792 kB) free!
[1587822123] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (9347788 kB) free!
[1587822173] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.5% (23053256 kB) free!
[1587822293] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.6% (27167400 kB) free.
[1587822893] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.7% (26015152 kB) free!
[1587823013] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.3% (29370732 kB) free.
[1587823255] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30130452 kB) free.
[1587823613] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.1% (21238220 kB) free!
[1587823733] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 10.9% (14319960 kB) free!
[1587823853] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.6% (20611824 kB) free!
[1587823855] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (1631144 kB) free!
[1587823975] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (376968 kB) free!
[1587824095] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38470660 kB) free.
[1587824318] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.7% (3310976 kB) free!
[1587824438] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.7% (2851484 kB) free!
[1587824453] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (9375064 kB) free!
[1587824559] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.6% (2370664 kB) free!
[1587824695] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (9028756 kB) free!
[1587824815] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (4017992 kB) free!
[1587824935] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384876 kB) free!
[1587825053] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.6% (20618060 kB) free!
[1587825159] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (175144 kB) free!
[1587825535] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16776588 kB) free!
[1587825723] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4276668 kB) free!
[1587826083] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1351132 kB) free!
[1587826135] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36953824 kB) free.
[1587826683] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3304568 kB) free.
[1587826736] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6959456 kB) free!
[1587826853] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1587826853] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (13232464 kB) free!
[1587826856] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (2167484 kB) free!
[1587826976] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403356 kB) free!
[1587827283] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (1329736 kB) free!
[1587827404] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (1599984 kB) free!
[1587827453] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.2% (27980656 kB) free.
[1587827524] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.4% (1856316 kB) free!
[1587827576] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13890324 kB) free!
[1587828053] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.8% (20809456 kB) free!
[1587828173] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.3% (26808832 kB) free.
[1587828176] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26608112 kB) free.
[1587828724] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1060404 kB) free!
[1587828760] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 73.3% (16449068 kB) free.
[1587828773] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.1% (21246916 kB) free!
[1587828776] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (384304 kB) free!
[1587828893] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.8% (26091624 kB) free!
[1587828896] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26903616 kB) free.
[1587829013] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.9% (22230124 kB) free!
[1587829324] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (2592432 kB) free!
[1587829496] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (1715360 kB) free!
[1587829616] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (402592 kB) free!
[1587829664] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.92, 3.68, 4.45
[1587829736] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28735496 kB) free.
[1587829924] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (731992 kB) free!
[1587830336] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6825784 kB) free!
[1587830456] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2691076 kB) free!
[1587830524] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.2% (2152280 kB) free!
[1587830576] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408184 kB) free!
[1587831724] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (1328908 kB) free!
[1587831776] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.0% (25016744 kB) free!
[1587832324] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.7% (2396224 kB) free!
[1587832376] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (3895276 kB) free!
[1587832976] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (14209200 kB) free!
[1587833524] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1350980 kB) free!
[1587834124] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.2% (2308600 kB) free!
[1587834176] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (5478440 kB) free!
[1587834776] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (15908176 kB) free!
[1587835324] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (473344 kB) free!
[1587835376] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26777560 kB) free.
[1587835664] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1587835976] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5212404 kB) free!
[1587836096] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1532872 kB) free!
[1587836216] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (28961956 kB) free.
[1587836524] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.6% (1729952 kB) free!
[1587836816] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7568784 kB) free!
[1587836936] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (3975616 kB) free!
[1587837056] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30888304 kB) free.
[1587837656] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9526824 kB) free!
[1587837776] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (5772224 kB) free!
[1587837896] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2181984 kB) free!
[1587838013] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 27.3% (35961472 kB) free.
[1587839096] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.1% (22583176 kB) free!