[1588703779] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16087204 kB) free!
[1588704084] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1588704084] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1588704084] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1588704204] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 140741 MB (62% inode=99%):
[1588704204] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1588704204] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1588704379] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28337848 kB) free.
[1588704979] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7937380 kB) free!
[1588705099] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37475312 kB) free.
[1588705248] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 15.5% (3799940 kB) free!
[1588705308] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (1069128 kB) free!
[1588705368] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (185756 kB) free!
[1588705428] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (183120 kB) free!
[1588705701] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13829948 kB) free!
[1588705821] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (10166984 kB) free!
[1588705941] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (6702180 kB) free!
[1588705955] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (208768 kB) free!
[1588706075] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (702512 kB) free!
[1588706195] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (200336 kB) free!
[1588706328] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 75.5% (18505364 kB) free.
[1588706541] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15842700 kB) free!
[1588707395] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.9% (2099552 kB) free!
[1588707741] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (6356604 kB) free!
[1588708341] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (15887456 kB) free!
[1588708941] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26595464 kB) free.
[1588710141] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15326348 kB) free!
[1588710261] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (11562972 kB) free!
[1588710381] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26564380 kB) free.
[1588710981] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15622100 kB) free!
[1588711101] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (11950204 kB) free!
[1588711221] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (8921204 kB) free!
[1588711821] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16572844 kB) free!
[1588713021] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34046456 kB) free.
[1588713131] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 87.0% (42894192 kB) free.
[1588713547] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.4% (44050008 kB) free.
[1588713621] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12908808 kB) free!
[1588713741] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37119676 kB) free.
[1588713975] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 24.8% (6075796 kB) free.
[1588713995] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3371904 kB) free.
[1588714236] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1588714341] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15989932 kB) free!
[1588714444] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 20.6% (5053988 kB) free.
[1588714461] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (12364496 kB) free!
[1588714464] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2853) < 203.212.70.170 [06/May/2020:03:01:42 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1588714521] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(333) < 183.82.105.122 [06/May/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1588714575] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (748996 kB) free!
[1588714581] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37088448 kB) free.
[1588714584] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588714641] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588714695] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (653880 kB) free!
[1588714778] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(456) < 2020/05/06 03:01:32 [error] 22134#0: *86954296 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1588714815] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (228628 kB) free!
[1588714899] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588715044] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (2135276 kB) free!
[1588715164] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.7% (2131164 kB) free!
[1588715181] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15362416 kB) free!
[1588715285] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2125028 kB) free!
[1588715302] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (11952736 kB) free!
[1588715422] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35950028 kB) free.
[1588716022] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7460324 kB) free!
[1588716142] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1608940 kB) free!
[1588716262] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34079960 kB) free.
[1588716862] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (7013920 kB) free!
[1588716982] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2205708 kB) free!
[1588717102] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34220620 kB) free.
[1588717702] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7587340 kB) free!
[1588717815] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 37.0% (9064948 kB) free.
[1588717822] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2936052 kB) free!
[1588717942] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33550516 kB) free.
[1588718415] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (236344 kB) free!
[1588718535] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (237552 kB) free!
[1588718542] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9769332 kB) free!
[1588718655] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1588718655] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (244544 kB) free!
[1588718662] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36400296 kB) free.
[1588719262] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14070384 kB) free!
[1588719382] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (10521216 kB) free!
[1588719503] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38481784 kB) free.
[1588719855] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 38.1% (9343488 kB) free.
[1588719995] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1588720102] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (17997472 kB) free!
[1588720222] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38470700 kB) free.
[1588720822] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20069908 kB) free!
[1588720942] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.5% (16465436 kB) free!
[1588721062] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37310264 kB) free.
[1588721655] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193816 kB) free!
[1588721662] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17760892 kB) free!
[1588721775] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182948 kB) free!
[1588721782] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38432960 kB) free.
[1588721895] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (230940 kB) free!
[1588722382] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19122340 kB) free!
[1588722502] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32506844 kB) free.
[1588723076] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 58.0% (14317488 kB) free.
[1588723085] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 57.1% (13995328 kB) free.
[1588723102] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23228696 kB) free!
[1588723222] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.9% (19632492 kB) free!
[1588723233] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 76.3% (18695200 kB) free.
[1588723342] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39021688 kB) free.
[1588723942] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (18963624 kB) free!
[1588724062] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39336804 kB) free.
[1588724662] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19720704 kB) free!
[1588724783] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40125852 kB) free.
[1588725382] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19426212 kB) free!
[1588725502] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39854100 kB) free.
[1588726102] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19605564 kB) free!
[1588726222] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31639680 kB) free.
[1588726729] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4704692 kB) free!
[1588726789] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 12.6% (3092972 kB) free!
[1588726822] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21394964 kB) free!
[1588726850] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 30.5% (7471224 kB) free.
[1588726942] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (17528884 kB) free!
[1588727062] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1588727062] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (38024600 kB) free.
[1588727150] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4856612 kB) free!
[1588727210] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 19.8% (4857100 kB) free!
[1588727270] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 19.8% (4854996 kB) free!
[1588727330] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 19.8% (4858668 kB) free!
[1588727662] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18197272 kB) free!
[1588727782] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38447576 kB) free.
[1588727876] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (219024 kB) free!
[1588727996] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (213080 kB) free!
[1588728116] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194932 kB) free!
[1588728382] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19085628 kB) free!
[1588728502] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38675116 kB) free.
[1588729102] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18133608 kB) free!
[1588729222] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38783168 kB) free.
[1588729822] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19608112 kB) free!
[1588729943] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38526844 kB) free.
[1588730542] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20221664 kB) free!
[1588730662] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.5% (16414196 kB) free!
[1588730782] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36486516 kB) free.
[1588731382] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16620948 kB) free!
[1588731502] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36834824 kB) free.
[1588732095] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1588732102] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16856960 kB) free!
[1588732222] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36883484 kB) free.
[1588732237] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 353 processes
[1588732357] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 354 processes
[1588732477] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1588732822] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15898288 kB) free!
[1588732942] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37286988 kB) free.
[1588733542] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16635348 kB) free!
[1588733662] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37234592 kB) free.
[1588734262] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16278480 kB) free!
[1588734382] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36833660 kB) free.
[1588734982] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16189564 kB) free!
[1588735103] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36343768 kB) free.
[1588735702] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15899740 kB) free!
[1588735822] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11725332 kB) free!
[1588735942] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32866368 kB) free.
[1588736285] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (2010264 kB) free!
[1588736400] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 18.02, 14.52, 9.70
[1588736405] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (1719528 kB) free!
[1588736498] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(29) < 2020/05/06 09:07:00 [error] 22123#0: *86970602 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1588736520] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 21.80, 16.74, 11.08
[1588736525] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (1459300 kB) free!
[1588736542] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7919256 kB) free!
[1588736618] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588736640] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 18.12, 16.77, 11.77
[1588736662] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33744428 kB) free.
[1588737262] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9316040 kB) free!
[1588737382] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4861936 kB) free!
[1588737502] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29024148 kB) free.
[1588738102] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4181828 kB) free!
[1588738222] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3248296 kB) free!
[1588738342] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34687112 kB) free.
[1588739077] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1588739542] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16593884 kB) free!
[1588739662] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (11941936 kB) free!
[1588739783] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36529408 kB) free.
[1588740033] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182720 kB) free!
[1588740045] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1588740153] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171752 kB) free!
[1588740165] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1588740273] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189352 kB) free!
[1588740285] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1588740382] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11755076 kB) free!
[1588740395] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.4% (2183536 kB) free!
[1588740502] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7360500 kB) free!
[1588740515] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (3189788 kB) free!
[1588740622] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35003840 kB) free.
[1588740635] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (3108464 kB) free!
[1588740885] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1588740984] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(832) < 103.255.146.164 [06/May/2020:10:19:38 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588741105] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588741222] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (11984136 kB) free!
[1588741342] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7694556 kB) free!
[1588741462] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32780820 kB) free.
[1588742062] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8048096 kB) free!
[1588742141] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 53.2% (13042632 kB) free.
[1588742182] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3554728 kB) free!
[1588742302] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31668252 kB) free.
[1588742433] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.30662.189> (10.147.212.158:45330 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1588742553] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1588742902] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6462636 kB) free!
[1588743022] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27041264 kB) free.
[1588743428] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 10.5% (2349940 kB) free!
[1588743548] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (1830304 kB) free!
[1588743622] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7543536 kB) free!
[1588743668] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (1372644 kB) free!
[1588743742] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2928460 kB) free!
[1588743862] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381756 kB) free!
[1588744235] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1606908 kB) free!
[1588744835] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.5% (2350112 kB) free!
[1588745062] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.9% (24964356 kB) free!
[1588745662] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405960 kB) free!
[1588746262] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1588746262] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19739084 kB) free!
[1588746424] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (5517852 kB) free!
[1588746669] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.4% (16245568 kB) free.
[1588746862] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34293116 kB) free.
[1588746941] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.7% (4088084 kB) free!
[1588747054] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1588747061] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.9% (3906240 kB) free!
[1588747174] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1588747181] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 15.3% (3761000 kB) free!
[1588747274] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1588747394] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1588747462] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9668776 kB) free!
[1588747514] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1588747582] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32209832 kB) free.
[1588747624] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4396100 kB) free!
[1588747835] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1588747835] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1142212 kB) free!
[1588748115] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1588748182] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15233836 kB) free!
[1588748235] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1588748302] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (10618116 kB) free!
[1588748382] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2162948 kB) free!
[1588748422] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (6385724 kB) free!
[1588748835] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1588748955] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1588748974] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 334 processes
[1588749022] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16663232 kB) free!
[1588749035] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.1% (1641464 kB) free!
[1588749075] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1588749094] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1588749548] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9652288 kB) free!
[1588749622] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26917292 kB) free.
[1588749635] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (1374632 kB) free!
[1588749668] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9413972 kB) free!
[1588749788] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.5% (9101524 kB) free!
[1588750222] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (3167732 kB) free!
[1588750235] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.7% (2067944 kB) free!
[1588750342] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33461648 kB) free.
[1588750835] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (979208 kB) free!
[1588750942] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7546200 kB) free!
[1588751063] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (3063568 kB) free!
[1588751183] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397664 kB) free!
[1588752035] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.3% (1667052 kB) free!
[1588752188] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4172176 kB) free!
[1588752383] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16364820 kB) free!
[1588752873] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.8% (4847588 kB) free!
[1588752983] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2093392 kB) free!
[1588753235] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1396240 kB) free!
[1588753275] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1588753395] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 332 processes
[1588753515] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1588753869] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (2046604 kB) free!
[1588753989] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (1414392 kB) free!
[1588754109] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (793216 kB) free!
[1588754115] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1588754183] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30169720 kB) free.
[1588754673] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2265292 kB) free!
[1588754715] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 340 processes
[1588754783] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (3847148 kB) free!
[1588754835] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1588754903] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (389260 kB) free!
[1588754943] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6823244 kB) free!
[1588754956] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1588755023] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38106072 kB) free.
[1588755035] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.4% (2510520 kB) free!
[1588755046] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.5% (1710244 kB) free!
[1588755063] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 36.8% (18114016 kB) free.
[1588755556] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1588755556] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 334 processes
[1588755623] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11136264 kB) free!
[1588755646] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (404012 kB) free!
[1588755743] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6615368 kB) free!
[1588755863] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (2516312 kB) free!
[1588756104] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(414) < 103.255.146.164 [06/May/2020:14:32:42 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588756156] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1588756224] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588756332] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9065092 kB) free!
[1588756452] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (8970448 kB) free!
[1588756463] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18704032 kB) free!
[1588756572] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.0% (8852216 kB) free!
[1588756825] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(569) < 103.255.146.164 [06/May/2020:14:47:08 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588756835] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1138008 kB) free!
[1588756945] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588757063] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35738264 kB) free.
[1588757085] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1588757205] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 349 processes
[1588757325] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 351 processes
[1588757435] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.2% (2149576 kB) free!
[1588757545] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(313) < 185.255.8.58 [06/May/2020:15:00:04 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1588757663] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9670776 kB) free!
[1588757665] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588757710] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 73.0% (16395980 kB) free.
[1588757784] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4941240 kB) free!
[1588757904] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (436468 kB) free!
[1588757925] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1588757982] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 61.2% (15010072 kB) free.
[1588758046] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.1% (1649584 kB) free!
[1588758277] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1588758504] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (14004632 kB) free!
[1588758635] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1530112 kB) free!
[1588758865] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6) < 172.105.59.173 [06/May/2020:15:20:19 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1588758986] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588759104] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31925324 kB) free.
[1588759235] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.3% (3463652 kB) free.
[1588759478] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1588759598] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1588759704] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6685220 kB) free!
[1588759718] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1588759824] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2359740 kB) free!
[1588759836] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (3066932 kB) free!
[1588759944] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35888504 kB) free.
[1588759956] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (3241100 kB) free!
[1588760076] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 20.0% (3247744 kB) free!
[1588760544] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11120880 kB) free!
[1588760618] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/06 15:43:47 [error] 22135#0: *87004195 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1588760664] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6362040 kB) free!
[1588760738] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588760772] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4880700 kB) free!
[1588760784] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (1916352 kB) free!
[1588760918] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1588761384] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14154400 kB) free!
[1588761984] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37621372 kB) free.
[1588762476] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (469984 kB) free!
[1588762584] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10914732 kB) free!
[1588762586] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [06/May/2020:16:21:21 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1588762704] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6295840 kB) free!
[1588762706] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588762824] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1844352 kB) free!
[1588762846] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1588762846] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1587744 kB) free!
[1588763076] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (2582256 kB) free!
[1588763138] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/06 16:35:03 [error] 22135#0: *87010079 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1588763259] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588763425] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20464632 kB) free!
[1588763446] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 11.3% (1838652 kB) free!
[1588763494] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1588763614] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1588763734] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1588764276] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (344048 kB) free!
[1588764625] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (11747316 kB) free!
[1588764876] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.2% (2636776 kB) free!
[1588765225] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31851048 kB) free.
[1588765510] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.4% (3915628 kB) free!
[1588765630] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.4% (3679448 kB) free!
[1588765658] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/06 17:09:22 [error] 22135#0: *87013752 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1588765750] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.3% (3442980 kB) free!
[1588765778] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588765825] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (5436932 kB) free!
[1588765945] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (641480 kB) free!
[1588766065] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17598884 kB) free!
[1588766076] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (599400 kB) free!
[1588766350] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1476092 kB) free!
[1588766378] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/06 17:24:51 [error] 22135#0: *87015313 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1588766498] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588766665] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (10251232 kB) free!
[1588766676] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.4% (2830048 kB) free!
[1588766918] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 347 processes
[1588766924] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 378 processes
[1588766982] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.4% (4276944 kB) free!
[1588767038] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 352 processes
[1588767044] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 379 processes
[1588767102] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.7% (4091160 kB) free!
[1588767158] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1588767158] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1588767165] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 377 processes
[1588767223] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 15.6% (3832376 kB) free!
[1588767265] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30176780 kB) free.
[1588767556] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1588767698] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/05/06 17:49:34 [error] 22135#0: *87017748 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1588767704] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 45.3% (11091484 kB) free.
[1588767818] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588767865] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (3210460 kB) free!
[1588767985] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (411516 kB) free!
[1588768105] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (378712 kB) free!
[1588768423] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1926936 kB) free!
[1588768751] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1588768751] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 73.4% (16480648 kB) free.
[1588768846] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (201012 kB) free!
[1588768958] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 348 processes
[1588768964] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 389 processes
[1588769084] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 387 processes
[1588769305] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21490484 kB) free!
[1588769618] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/06 18:22:56 [error] 22134#0: *87020294 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1588769738] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588769905] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (505508 kB) free!
[1588770046] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 13.9% (2264756 kB) free!
[1588770240] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 20.64, 27.47, 23.13
[1588770506] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (38986936 kB) free.
[1588770753] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.29568.107> (10.147.212.158:49078 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1588770841] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 14.25, 15.95, 18.71
[1588770874] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1588770876] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (158804 kB) free!
[1588771095] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.7% (3607308 kB) free!
[1588771106] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11248796 kB) free!
[1588771226] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6639820 kB) free!
[1588771346] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2357704 kB) free!
[1588771431] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 13 users currently logged in
[1588771476] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.8% (4189340 kB) free.
[1588771551] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 13 users currently logged in
[1588771671] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1588771695] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167160 kB) free!
[1588771946] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.6% (24553572 kB) free!
[1588772272] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 17 users currently logged in
[1588772546] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397588 kB) free!
[1588773146] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23099508 kB) free!
[1588773276] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.6% (1721648 kB) free!
[1588773396] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3159084 kB) free!
[1588773517] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.3% (4120432 kB) free.
[1588773746] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (577508 kB) free!
[1588774116] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (884268 kB) free!
[1588774236] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.3% (2317980 kB) free!
[1588774346] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21171632 kB) free!
[1588774356] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3644384 kB) free.
[1588774946] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399352 kB) free!
[1588775272] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 15 users currently logged in
[1588775546] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16909004 kB) free!
[1588775556] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.7% (1894900 kB) free!
[1588775676] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3311772 kB) free.
[1588775872] SERVICE FLAPPING ALERT: sigqscluster1;Current Users;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1588775872] SERVICE ALERT: sigqscluster1;Current Users;OK;HARD;1;USERS OK - 12 users currently logged in
[1588776146] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38925916 kB) free.
[1588776241] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.33, 3.89, 4.48
[1588776746] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13097588 kB) free!
[1588776866] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8615388 kB) free!
[1588776876] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (591196 kB) free!
[1588776986] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (4626492 kB) free!
[1588776996] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.8% (1916932 kB) free!
[1588777116] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3177320 kB) free!
[1588777441] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.07, 6.85, 5.60
[1588777561] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.90, 6.45, 5.60
[1588777587] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.1% (23847868 kB) free!
[1588777681] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1588777681] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.12, 6.16, 5.58
[1588777716] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3476240 kB) free.
[1588777846] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (535928 kB) free!
[1588778187] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (1288360 kB) free!
[1588778281] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.43, 4.46, 4.91
[1588778351] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1588778787] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19860384 kB) free!
[1588778916] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (809624 kB) free!
[1588778951] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4193964 kB) free!
[1588779037] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.3% (2004692 kB) free!
[1588779072] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 73.1% (16408860 kB) free.
[1588779157] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.0% (3418288 kB) free.
[1588779158] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1588779387] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39951460 kB) free.
[1588779987] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16278432 kB) free!
[1588780107] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11862052 kB) free!
[1588780227] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (7821532 kB) free!
[1588780827] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.2% (25309464 kB) free!
[1588781427] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (4082332 kB) free!
[1588782027] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.9% (20904436 kB) free!
[1588782157] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (3087768 kB) free!
[1588782277] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.4% (2669524 kB) free!
[1588782332] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1685 processes
[1588782397] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.3% (3459716 kB) free.
[1588782627] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39057452 kB) free.
[1588782933] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1755 processes
[1588782997] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.3% (2481048 kB) free!
[1588783117] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3262568 kB) free.
[1588783227] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15476052 kB) free!
[1588783347] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11104988 kB) free!
[1588783467] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (6927536 kB) free!
[1588783717] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.6% (1892556 kB) free!
[1588783837] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.8% (2575636 kB) free!
[1588783846] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1588783957] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3291116 kB) free.
[1588784067] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.3% (24179744 kB) free!
[1588784133] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1501 processes
[1588784558] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.1% (1970756 kB) free!
[1588784667] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2355004 kB) free!
[1588784678] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.9% (2585392 kB) free!
[1588784734] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1588784734] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1754 processes
[1588784798] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3291452 kB) free.
[1588785268] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19793848 kB) free!
[1588785398] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.6% (1888588 kB) free!
[1588785518] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.9% (2424968 kB) free!
[1588785638] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (3038744 kB) free!
[1588785868] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38594724 kB) free.
[1588786238] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.2% (3774392 kB) free.
[1588786468] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15709828 kB) free!
[1588786588] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (11609888 kB) free!
[1588786708] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (7633032 kB) free!
[1588787272] SERVICE FLAPPING ALERT: sigqscluster1;Current Users;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1588787308] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.7% (23318744 kB) free!
[1588787438] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (1398056 kB) free!
[1588787558] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.1% (1804048 kB) free!
[1588787678] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.6% (2203712 kB) free!
[1588787908] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3371444 kB) free!
[1588788158] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1588788278] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.9% (3561660 kB) free.
[1588788508] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17877988 kB) free!
[1588788878] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (1609952 kB) free!
[1588788998] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.6% (1891556 kB) free!
[1588789108] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (33984716 kB) free.
[1588789119] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.9% (2255076 kB) free!
[1588789681] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1588789708] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12778252 kB) free!
[1588789719] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.0% (3406756 kB) free.