[1548959489] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9635472 kB) free!
[1548959609] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8064564 kB) free!
[1548959729] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9970404 kB) free.
[1548960930] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1548960963] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1548962163] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548962283] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548962404] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548962729] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9626548 kB) free!
[1548962849] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8024000 kB) free!
[1548962969] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9959744 kB) free.
[1548965969] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9855672 kB) free!
[1548966003] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548966090] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.2% (8491704 kB) free!
[1548966123] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548966210] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9944920 kB) free.
[1548966244] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1548968188] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548968308] SERVICE FLAPPING ALERT: db1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1548968308] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1548969243] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548969363] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548969484] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548969810] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9431304 kB) free!
[1548969930] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7850496 kB) free!
[1548970050] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9928604 kB) free.
[1548971908] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548972028] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548972148] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548972748] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548973050] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9731156 kB) free!
[1548973083] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548973170] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8207308 kB) free!
[1548973203] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548973290] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9911328 kB) free.
[1548973324] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548976290] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9768988 kB) free!
[1548976323] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548976410] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.9% (8321792 kB) free!
[1548976443] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548976530] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9905940 kB) free.
[1548976563] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548978148] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548978268] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548979530] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9844896 kB) free!
[1548979563] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548979650] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.2% (8499216 kB) free!
[1548979683] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548979770] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 20.0% (9839084 kB) free!
[1548979803] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548980370] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9910636 kB) free.
[1548982803] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548982923] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548983043] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548983068] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548983188] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1548983370] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9465404 kB) free!
[1548983490] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7889316 kB) free!
[1548983610] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9906648 kB) free.
[1548983789] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548983909] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548984029] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548984629] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1548986043] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548986163] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548986283] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548986429] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1548986549] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1548986610] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9391712 kB) free!
[1548986730] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7778668 kB) free!
[1548986850] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9884164 kB) free.
[1548987450] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9877444 kB) free!
[1548987571] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 20.0% (9877096 kB) free!
[1548987691] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.9% (9332420 kB) free!
[1548987750] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548987870] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548989283] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548989403] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548989523] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548990091] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1548990091] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9893756 kB) free.
[1548990691] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9878768 kB) free!
[1548990811] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 20.0% (9878304 kB) free!
[1548990931] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.5% (9606180 kB) free!
[1548991924] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1548992044] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548992071] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548992164] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548992191] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1548992764] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548993364] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548993484] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548993604] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548994204] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548994325] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548994445] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548995045] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548995165] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548995285] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1548995885] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548996005] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548996125] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548996726] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1548996846] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548996966] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548996992] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548997112] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548998166] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1548998166] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1548998766] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548998886] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1548999486] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548999606] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548999727] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1548999932] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10105680 kB) free.
[1549000326] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1549000446] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1549000532] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7049804 kB) free!
[1549000566] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1549000652] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10074716 kB) free.
[1549001312] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1549001432] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1549002632] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1549002752] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1549003052] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9485576 kB) free!
[1549003172] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7886384 kB) free!
[1549003292] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10038612 kB) free.
[1549003567] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1549003687] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1549003807] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1549003952] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1549004072] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1549004672] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1549004723] SERVICE ALERT: sservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1549004745] HOST ALERT: sservercluster2;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1549004793] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1549004843] SERVICE ALERT: sservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1549004847] HOST ALERT: sservercluster2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1549005007] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1549005093] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8892568 kB) free!
[1549005124] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1549005127] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1549005213] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7167260 kB) free!
[1549005244] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 325 processes
[1549005247] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1549005333] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.5% (9591176 kB) free!
[1549005364] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 330 processes
[1549005964] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 330 processes
[1549006390] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1549006448] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549006510] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1549006564] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 322 processes
[1549006568] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549006630] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1549006688] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549006910] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 317 processes
[1549007030] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1549007150] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1549007192] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1549007312] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1549008350] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 316 processes
[1549008488] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1549008512] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1549008633] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1549009088] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549009208] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1549009809] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549009832] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1549009929] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549009952] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1549010049] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549010750] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 322 processes
[1549011152] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1549011272] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1549011351] SERVICE FLAPPING ALERT: webservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1549011351] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 326 processes
[1549011437] SERVICE ALERT: webservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1549012449] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1549012473] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1549012533] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1549012593] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1549013049] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549013170] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549013290] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549013793] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1549013913] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1549014430] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1549014944] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47029 MB (20% inode=99%):
[1549015002] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24610 MB (10% inode=99%):
[1549015030] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1549015064] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47027 MB (20% inode=99%):
[1549015113] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1549015150] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1549015184] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46995 MB (20% inode=99%):
[1549015233] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1549015270] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1549015353] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1549016290] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1549016553] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1549016673] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1549016793] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1549018690] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549018810] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549018930] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549019470] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1549020394] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1549020514] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1549020634] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1549020951] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 322 processes
[1549021551] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 322 processes
[1549021671] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 323 processes
[1549021791] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 323 processes
[1549021834] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1549021930] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1549021954] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1549022530] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549022577] SERVICE ALERT: webserver2;Current Load;WARNING;SOFT;1;WARNING - load average: 7.89, 4.65, 2.33
[1549022637] SERVICE ALERT: webserver2;Current Load;WARNING;SOFT;2;WARNING - load average: 6.45, 4.85, 2.55
[1549022650] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1549022697] SERVICE ALERT: webserver2;Current Load;WARNING;SOFT;3;WARNING - load average: 7.60, 5.39, 2.87
[1549022757] SERVICE ALERT: webserver2;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 9.13, 6.11, 3.27
[1549023154] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1549023274] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1549023850] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1549023957] SERVICE ALERT: webserver2;Current Load;WARNING;HARD;4;WARNING - load average: 0.08, 1.47, 3.79
[1549023970] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1549024090] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549024191] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1549024257] SERVICE ALERT: webserver2;Current Load;OK;HARD;1;OK - load average: 0.13, 0.64, 2.79
[1549024791] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1549024911] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 330 processes
[1549025031] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 331 processes
[1549025074] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1549025194] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1549025295] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 355 processes
[1549025314] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1549025415] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 388 processes
[1549025470] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1549025590] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1549025710] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1549026490] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1549027090] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549027114] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1549027210] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549027235] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1549027330] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1549027431] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 323 processes
[1549028031] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 324 processes
[1549028152] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 1 warn out of 328 processes
[1549028272] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 321 processes
[1549029131] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549029251] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549029371] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549030072] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 328 processes
[1549030672] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 328 processes
[1549030908] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 41 users currently logged in
[1549030968] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 41 users currently logged in
[1549031028] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 41 users currently logged in
[1549031088] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 41 users currently logged in
[1549031988] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 40 users currently logged in
[1549032371] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549032492] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549032612] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549033212] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549033812] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549033932] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549034052] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1549036638] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1549037052] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549037172] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1549037292] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1549037710] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1549038034] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1549038154] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1549038274] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1549038874] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1549038910] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1549039030] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1549039150] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1549039474] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1549039594] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1549040194] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1549040292] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549040314] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1549040412] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549040435] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1549040532] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549041701] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.6% (1722520 kB) free!
[1549042072] SERVICE FLAPPING ALERT: webservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1549042234] SERVICE ALERT: db1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1549042354] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1549042474] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1549042671] SERVICE ALERT: db1;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 6534 MB (12% inode=99%):
[1549042672] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 316 processes
[1549042932] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1549043940] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 528 processes
[1549043950] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1549044000] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 529 processes
[1549044132] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1549044252] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1549044372] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1549044438] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1549044972] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1549045638] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1549045758] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in