[1558895720] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558895771] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3479976 kB) free!
[1558895840] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558895859] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8377544 kB) free!
[1558895891] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9917356 kB) free.
[1558895960] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1558895979] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.7% (12644960 kB) free.
[1558896491] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (4781372 kB) free!
[1558896561] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558896611] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (2543120 kB) free!
[1558896681] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1558896732] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9915628 kB) free.
[1558897281] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558897299] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1558897320] SERVICE FLAPPING ALERT: secondary-server;Current Users;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1558897331] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (4007700 kB) free!
[1558897401] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558897451] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9919196 kB) free.
[1558897521] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1558897779] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9823792 kB) free!
[1558897899] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8158568 kB) free!
[1558898019] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.7% (12679960 kB) free.
[1558898051] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.3% (5081748 kB) free!
[1558898121] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558898171] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (3013020 kB) free!
[1558898241] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1558898291] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9915100 kB) free.
[1558898841] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558898891] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4680592 kB) free!
[1558898962] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558899011] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (2530268 kB) free!
[1558899082] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1558899131] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9901668 kB) free.
[1558899682] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558899731] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3946524 kB) free!
[1558899802] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558899819] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8172728 kB) free!
[1558899851] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1558899851] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9904312 kB) free.
[1558899922] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558899940] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.8% (12705832 kB) free.
[1558900451] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.2% (5016688 kB) free!
[1558900522] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558900572] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (2838352 kB) free!
[1558900642] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558900692] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9900528 kB) free.
[1558901243] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558901292] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (4210624 kB) free!
[1558901363] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558901412] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9896084 kB) free.
[1558901483] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1558902012] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5631368 kB) free!
[1558902083] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558902132] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (3705188 kB) free!
[1558902203] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558902252] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9912228 kB) free.
[1558902339] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8387112 kB) free!
[1558902459] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.9% (12745584 kB) free.
[1558902803] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558902852] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4909344 kB) free!
[1558902923] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558902972] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (2503028 kB) free!
[1558903043] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1558903092] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9907452 kB) free.
[1558903299] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 248.95, 128.03, 51.30
[1558903420] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 256.13, 170.77, 76.27
[1558903540] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 94.50, 152.65, 82.14
[1558903575] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 296.06, 203.29, 92.54
[1558903644] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558903692] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3844980 kB) free!
[1558903695] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 77.16, 163.83, 92.53
[1558903732] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 580 processes
[1558903764] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558903812] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9899668 kB) free.
[1558903815] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 10.70, 109.76, 81.36
[1558903852] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 347 processes
[1558903884] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1558904413] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.5% (5161408 kB) free!
[1558904484] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558904533] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (2995892 kB) free!
[1558904604] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1558904653] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9908080 kB) free.
[1558905204] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558905253] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (3796380 kB) free!
[1558905324] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558905340] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.43, 0.86, 12.21
[1558905373] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9895420 kB) free.
[1558905444] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1558905459] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8958880 kB) free!
[1558905579] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.0% (12803644 kB) free.
[1558905615] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.29, 0.58, 11.95
[1558905973] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (4851500 kB) free!
[1558906045] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558906093] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (2588024 kB) free!
[1558906165] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1558906179] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8531864 kB) free!
[1558906213] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9889780 kB) free.
[1558906299] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.0% (12793704 kB) free.
[1558906540] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.47, 0.48, 3.67
[1558906765] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558906813] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (4424112 kB) free!
[1558906815] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.21, 0.18, 3.43
[1558906885] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558906933] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.5% (5645684 kB) free!
[1558907005] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558907053] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9898048 kB) free.
[1558907499] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8491844 kB) free!
[1558907605] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558907619] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.0% (12801856 kB) free.
[1558907653] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3460284 kB) free!
[1558907725] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1558907774] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9892428 kB) free.
[1558908325] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558908373] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4698416 kB) free!
[1558908446] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558908493] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (2350624 kB) free!
[1558908566] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1558908613] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9881916 kB) free.
[1558908980] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 252.16, 105.51, 40.52
[1558909100] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 275.16, 162.77, 69.38
[1558909166] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558909213] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3589696 kB) free!
[1558909217] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 951 processes
[1558909220] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 289.72, 206.41, 96.72
[1558909286] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558909333] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9892212 kB) free.
[1558909337] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 963 processes
[1558909419] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9096704 kB) free!
[1558909457] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 957 processes
[1558909539] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.1% (12853880 kB) free.
[1558909886] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558909933] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4735364 kB) free!
[1558910006] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558910053] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (2410528 kB) free!
[1558910057] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 365 processes
[1558910126] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1558910173] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9892816 kB) free.
[1558910726] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558910739] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8724688 kB) free!
[1558910773] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (4052332 kB) free!
[1558910847] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558910859] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.1% (12869300 kB) free.
[1558910893] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9901308 kB) free.
[1558911446] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558911493] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.4% (5136204 kB) free!
[1558911566] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558911614] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (2835656 kB) free!
[1558911620] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.78, 1.80, 11.66
[1558911686] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1558911734] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9902084 kB) free.
[1558912059] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8664788 kB) free!
[1558912179] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.1% (12885512 kB) free.
[1558912286] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558912334] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3942092 kB) free!
[1558912406] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1558912454] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9901556 kB) free.
[1558912820] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.81, 1.57, 4.31
[1558913006] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558913054] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4909100 kB) free!
[1558913126] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558913174] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9891096 kB) free.
[1558913247] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1558913774] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.1% (5479256 kB) free!
[1558913846] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558913894] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (3509836 kB) free!
[1558913966] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1558913979] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9036528 kB) free!
[1558914014] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9883908 kB) free.
[1558914099] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.2% (12893304 kB) free.
[1558914566] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558914614] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4711832 kB) free!
[1558914686] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558914734] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.3% (9487068 kB) free!
[1558914806] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1558914854] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9894624 kB) free.
[1558915406] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558915455] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3494900 kB) free!
[1558915527] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558915575] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9886880 kB) free.
[1558915899] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8764336 kB) free!
[1558916019] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.2% (12927860 kB) free.
[1558916126] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558916175] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4668664 kB) free!
[1558916246] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1558916295] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9895028 kB) free.
[1558916366] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1558916895] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.1% (5447284 kB) free!
[1558916966] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1558917015] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (3521376 kB) free!
[1558917086] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1558917135] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9889256 kB) free.
[1558917686] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558917735] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (4638792 kB) free!
[1558917806] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558917855] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9891212 kB) free.
[1558917927] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1558918419] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9160332 kB) free!
[1558918455] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.5% (5686044 kB) free!
[1558918526] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558918539] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.3% (12963076 kB) free.
[1558918575] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (3719648 kB) free!
[1558918646] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558918696] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9896128 kB) free.
[1558919246] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558919295] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (4900244 kB) free!
[1558919366] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558919415] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9889816 kB) free.
[1558919486] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558920015] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.6% (5708936 kB) free!
[1558920086] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1558920135] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (3742320 kB) free!
[1558920206] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1558920255] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9894880 kB) free.
[1558920339] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9152488 kB) free!
[1558920459] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.1% (12842828 kB) free.
[1558920807] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558920855] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4488016 kB) free!
[1558920927] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1558920975] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9884356 kB) free.
[1558921047] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1558921575] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.8% (5346188 kB) free!
[1558921647] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558921695] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (3263564 kB) free!
[1558921767] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558921815] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 20.0% (9877940 kB) free!
[1558922140] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 293.98, 161.48, 66.22
[1558922260] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 299.38, 207.38, 94.59
[1558922367] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558922380] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 85.95, 171.52, 96.25
[1558922415] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4367772 kB) free!
[1558922487] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558922607] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1558922859] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9002248 kB) free!
[1558922979] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 26.1% (12869948 kB) free.
[1558923016] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.9% (6872860 kB) free!
[1558923208] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558923328] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1558923928] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558924048] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558924168] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1558924180] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.46, 0.86, 14.18
[1558924768] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558924816] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1558924816] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (2919148 kB) free!
[1558924888] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1558925380] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.40, 0.43, 4.22
[1558925416] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.5% (6178156 kB) free!
[1558925488] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558925609] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558925729] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1558926329] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558926449] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558926569] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1558927169] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558927216] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2935044 kB) free!
[1558927289] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558927409] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558927816] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 20.0% (9861696 kB) free!
[1558928379] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9529676 kB) free!
[1558928420] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 293.26, 118.45, 45.30
[1558928499] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.9% (12747920 kB) free.
[1558928540] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 325.06, 188.49, 79.84
[1558928661] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 49.70, 130.60, 72.00
[1558929099] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9284208 kB) free!
[1558929219] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.7% (12644568 kB) free.
[1558929819] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9189392 kB) free!
[1558929939] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.4% (12535956 kB) free.
[1558930461] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.03, 2.96, 12.61
[1558930539] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9083132 kB) free!
[1558930659] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.2% (12418144 kB) free.
[1558930920] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1558931040] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1558931138] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (4350992 kB) free!
[1558931160] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1558931258] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.7% (4328988 kB) free!
[1558931260] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9158864 kB) free!
[1558931378] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (4298072 kB) free!
[1558931380] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.0% (12334772 kB) free.
[1558931980] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8986236 kB) free!
[1558932016] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4896944 kB) free!
[1558932100] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 25.0% (12299556 kB) free.
[1558932209] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1558932261] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.43, 1.64, 3.40
[1558932700] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9008144 kB) free!
[1558932809] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558932820] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.8% (12217720 kB) free.
[1558932930] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1558933050] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558933420] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8648856 kB) free!
[1558933540] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 24.3% (11955256 kB) free.
[1558933634] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 12 users currently logged in
[1558933754] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 12 users currently logged in
[1558933816] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.8% (9737860 kB) free!
[1558933874] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1558934140] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8470548 kB) free!
[1558934260] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.8% (11722644 kB) free.
[1558934395] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183280 kB) free!
[1558934516] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189180 kB) free!
[1558934637] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189440 kB) free!
[1558934860] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8021232 kB) free!
[1558934980] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.2% (11438620 kB) free.
[1558935450] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1558935450] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558936050] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558936170] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558936180] SERVICE ALERT: sigqscluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9032468 kB) free!
[1558936290] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558936300] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 23.7% (11674132 kB) free.
[1558936400] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1558936520] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1558936640] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1558937417] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (3471168 kB) free!
[1558937474] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 9 users currently logged in
[1558937661] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.15, 19.91, 13.37
[1558937720] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.7% (4339992 kB) free!
[1558937781] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.52, 14.04, 12.01
[1558937840] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.9% (4137624 kB) free!
[1558937901] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1558937901] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.04, 10.05, 10.79
[1558937960] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.2% (3982100 kB) free!
[1558937978] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (1477428 kB) free!
[1558938132] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 47.1% (23187204 kB) free.
[1558938185] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 76.9% (18985400 kB) free.
[1558938227] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.75, 6.49, 5.03
[1558938230] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 88.4% (21661636 kB) free.
[1558938276] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 63.8% (15645612 kB) free.
[1558938347] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.74, 6.31, 5.15
[1558938467] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.10, 5.89, 5.14
[1558938690] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1558938787] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189880 kB) free!
[1558938907] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181480 kB) free!
[1558939028] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189236 kB) free!
[1558939056] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.1% (4179892 kB) free!
[1558939101] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.62, 2.23, 4.54
[1558939176] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.7% (4095988 kB) free!
[1558939296] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.4% (4019052 kB) free!
[1558939861] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 13.4% (2179708 kB) free!
[1558939890] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558939981] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 12.0% (1957132 kB) free!
[1558940010] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558940102] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 10.6% (1726836 kB) free!
[1558940131] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558940702] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (506796 kB) free!
[1558941440] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1558942101] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 80.14, 19.11, 7.60
[1558942181] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 60.96, 35.59, 16.36
[1558942221] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 31.81, 29.95, 13.72
[1558942301] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 12.81, 25.56, 15.02
[1558942341] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.52, 20.54, 12.25
[1558942421] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.54, 18.39, 13.68
[1558942531] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558942896] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2460736 kB) free!
[1558943131] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558943251] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558943371] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558943536] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9402856 kB) free!
[1558943541] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.53, 2.09, 4.65
[1558943656] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 33.1% (16294080 kB) free.
[1558944821] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.48, 1.21, 3.61
[1558945067] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.38, 3.90, 4.57
[1558945953] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 13.9% (1023432 kB) free!
[1558946073] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 12.2% (894852 kB) free!
[1558946193] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 10.4% (762768 kB) free!
[1558946656] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (4174040 kB) free!
[1558946776] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (1729448 kB) free!
[1558946793] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (125432 kB) free!
[1558946896] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9949856 kB) free.
[1558947497] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9159856 kB) free!
[1558947617] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.0% (6908804 kB) free!
[1558947737] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4852848 kB) free!
[1558948093] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 350 processes
[1558948213] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 353 processes
[1558948334] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 353 processes
[1558948376] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (184648 kB) free!
[1558948497] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (165036 kB) free!
[1558948617] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155572 kB) free!
[1558948874] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1558948994] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1558949114] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1558949534] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1558949793] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.9% (6831980 kB) free.
[1558949840] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1558949960] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1558950080] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1558950137] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10041052 kB) free.
[1558950238] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (4218492 kB) free!
[1558950358] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.7% (4094752 kB) free!
[1558950478] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.2% (3980072 kB) free!
[1558950737] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9732028 kB) free!
[1558950857] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7232872 kB) free!
[1558950977] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.3% (5058860 kB) free!
[1558951422] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 4.19, 74.75, 62.88
[1558951542] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.28, 50.31, 55.36
[1558951577] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4743860 kB) free!
[1558951662] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.88, 33.93, 48.75
[1558951678] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2336788 kB) free!
[1558951771] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1558951986] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1558952046] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1558952106] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1558952166] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1558952810] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1558952821] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1558952862] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1558952862] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.42, 1.15, 13.79
[1558952887] SERVICE ALERT: redis-1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1558952941] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558952953] SERVICE ALERT: redis-1;Mariadb status;CRITICAL;SOFT;1;NRPE: Command 'check_mysql' not defined
[1558953007] SERVICE ALERT: redis-1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1558953061] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1558953073] SERVICE ALERT: redis-1;Mariadb status;CRITICAL;SOFT;2;NRPE: Command 'check_mysql' not defined
[1558953084] SERVICE ALERT: redis-1;Sensors Status;UNKNOWN;SOFT;1;SENSORS UNKNOWN - command not found (did you install lmsensors?)
[1558953127] SERVICE ALERT: redis-1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with command name 'ntpd'
[1558953181] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1558953181] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558953193] SERVICE ALERT: redis-1;Mariadb status;CRITICAL;HARD;3;NRPE: Command 'check_mysql' not defined
[1558953204] SERVICE ALERT: redis-1;Sensors Status;UNKNOWN;SOFT;2;SENSORS UNKNOWN - command not found (did you install lmsensors?)
[1558953324] SERVICE ALERT: redis-1;Sensors Status;UNKNOWN;HARD;3;SENSORS UNKNOWN - command not found (did you install lmsensors?)
[1558953589] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 356 processes
[1558953709] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1558954059] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1558954059] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1558954059] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1558954069] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1558954069] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1558954069] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1558954126] SERVICE ALERT: redis-2;Sensors Status;UNKNOWN;SOFT;1;SENSORS UNKNOWN - command not found (did you install lmsensors?)
[1558954134] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.94, 0.88, 4.04
[1558954246] SERVICE ALERT: redis-2;Sensors Status;UNKNOWN;SOFT;2;SENSORS UNKNOWN - command not found (did you install lmsensors?)
[1558954327] SERVICE ALERT: redis-1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 2 processes with command name 'ntpd'
[1558954366] SERVICE ALERT: redis-2;Sensors Status;OK;HARD;1;SENSORS OK
[1558954524] SERVICE ALERT: redis-1;Sensors Status;OK;HARD;1;SENSORS OK
[1558954534] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.9% (9801452 kB) free!
[1558954579] SERVICE ALERT: redis-2;Mariadb status;CRITICAL;SOFT;1;PROCS CRITICAL: 0 processes with command name 'mysqld'
[1558954699] SERVICE ALERT: redis-2;Mariadb status;CRITICAL;SOFT;2;PROCS CRITICAL: 0 processes with command name 'mysqld'
[1558954819] SERVICE ALERT: redis-2;Mariadb status;CRITICAL;HARD;3;PROCS CRITICAL: 0 processes with command name 'mysqld'
[1558954872] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (28.0% change >= 20.0% threshold)
[1558954872] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1558955021] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 60 users currently logged in
[1558955581] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1558955581] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1558956181] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558956301] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558956422] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558957097] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 16.8% (1233032 kB) free!
[1558957217] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 14.8% (1085240 kB) free!
[1558957337] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 12.9% (945824 kB) free!
[1558957534] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4267180 kB) free!
[1558957855] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 14 users currently logged in
[1558957937] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (24.5% change >= 20.0% threshold)
[1558957937] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (411384 kB) free!
[1558959111] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1558959198] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 618 processes
[1558959258] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 611 processes
[1558959422] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558959709] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 363 processes
[1558959829] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 363 processes
[1558959934] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.8% (9753444 kB) free!
[1558959949] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1558960022] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558960142] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558960263] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558960311] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1558960337] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.7% (6894768 kB) free.
[1558960342] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 268.24, 186.17, 85.79
[1558960431] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1558960462] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 268.67, 213.48, 107.97
[1558960534] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10191492 kB) free.
[1558960551] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1558960582] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 117.85, 190.66, 113.48
[1558961134] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8916772 kB) free!
[1558961254] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6825780 kB) free!
[1558961374] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4675780 kB) free!
[1558962382] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.71, 2.01, 17.72
[1558963263] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558963863] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558963983] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558964104] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1558964182] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.42, 1.50, 3.91
[1558964374] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.7% (9698596 kB) free!
[1558964746] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1558964751] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1558964934] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1558965579] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 251.31, 114.10, 43.98
[1558965699] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 266.19, 165.19, 71.20
[1558965819] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 268.08, 199.31, 95.10
[1558965973] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2729584 kB) free!
[1558966504] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1558967104] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1558967224] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558967344] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1558967619] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.31, 0.76, 13.92
[1558967974] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4166592 kB) free!
[1558968819] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.13, 0.19, 3.97
[1558969707] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1558969774] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 20.0% (9863352 kB) free!
[1558969827] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1558969937] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1558970173] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2444864 kB) free!
[1558973374] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4166204 kB) free!
[1558973974] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.9% (6831044 kB) free!
[1558974544] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558974946] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1558975066] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1558975186] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1558975744] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558975864] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558975984] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1558976019] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 20.47, 12.94, 5.74
[1558976139] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 20.51, 15.41, 7.51
[1558976173] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.6% (2831368 kB) free!
[1558976259] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.86, 10.28, 6.62
[1558976859] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1558976859] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.10, 1.56, 3.59
[1558977784] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558978984] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1558979104] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1558979704] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558979824] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558979945] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1558979986] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1558980544] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1558980574] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (3281952 kB) free!
[1558980664] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1558981174] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.2% (6017460 kB) free!
[1558981264] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1558981384] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes