[1579718020] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8476112 kB) free!
[1579718046] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.9% (29540848 kB) free.
[1579718140] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29183800 kB) free.
[1579718368] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180244 kB) free!
[1579718399] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183608 kB) free!
[1579718488] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (204836 kB) free!
[1579718520] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173484 kB) free!
[1579718610] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189680 kB) free!
[1579718639] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173364 kB) free!
[1579718641] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180124 kB) free!
[1579718740] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8430148 kB) free!
[1579718759] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194632 kB) free!
[1579718860] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29126568 kB) free.
[1579718879] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184556 kB) free!
[1579719460] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8047504 kB) free!
[1579719581] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29098456 kB) free.
[1579719905] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1579719920] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2020/01/23 00:26:00 [error] 18357#0: *42699271 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:127407/38820, bytes from/to upstream:38820/127407
[1579720025] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1579720040] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579720180] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7877788 kB) free!
[1579720300] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29053088 kB) free.
[1579720900] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7404552 kB) free!
[1579721020] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29021404 kB) free.
[1579721620] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7216444 kB) free!
[1579721740] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28202892 kB) free.
[1579721880] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.0% (2952988 kB) free!
[1579722340] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6703744 kB) free!
[1579722460] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 20.0% (26349744 kB) free!
[1579722480] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184244 kB) free!
[1579722541] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9138912 kB) free!
[1579722580] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.9% (20957680 kB) free!
[1579722661] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.9% (8333716 kB) free!
[1579722781] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.1% (7422044 kB) free!
[1579723041] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/23 01:18:43 [error] 18357#0: *42711661 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:58/17, bytes from/to upstream:17/58
[1579723161] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579723381] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (2668288 kB) free!
[1579723780] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27966324 kB) free.
[1579724980] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (405540 kB) free!
[1579725100] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27875836 kB) free.
[1579725701] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12894480 kB) free!
[1579725821] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.7% (26001932 kB) free!
[1579725941] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19891364 kB) free!
[1579726608] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (237440 kB) free!
[1579726728] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (241612 kB) free!
[1579726848] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224032 kB) free!
[1579727010] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 43.8% (10723584 kB) free.
[1579727141] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29281288 kB) free.
[1579727448] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.0% (43363716 kB) free.
[1579727741] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (17164388 kB) free!
[1579727861] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26516420 kB) free.
[1579728527] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 779 processes
[1579728561] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(149) < 2020/01/23 02:53:22 [error] 18358#0: *42726639 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579728648] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 788 processes
[1579728681] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(42) < 2020/01/23 03:00:37 [error] 18367#0: *42728205 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:46/0, bytes from/to upstream:0/46
[1579728768] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 330 processes
[1579728801] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579728923] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1425) < 49.50.64.122 [23/Jan/2020:03:02:32 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1579728956] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(30) < 49.50.86.152 [23/Jan/2020:03:00:13 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1579729043] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579729077] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579729661] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13534844 kB) free!
[1579729781] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30665660 kB) free.
[1579732368] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1211 processes
[1579732401] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(216) < 2020/01/23 04:02:33 [error] 18370#0: *42741595 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:58/0, bytes from/to upstream:0/58
[1579732489] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1209 processes
[1579732521] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579732609] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 326 processes
[1579732643] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(7) < 172.105.59.173 [23/Jan/2020:04:03:57 +0530] TCP 502 0 0 120.000 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1579732764] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579734320] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.1% (13288416 kB) free!
[1579735521] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/23 04:54:10 [error] 18355#0: *42752197 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:74/17, bytes from/to upstream:17/74
[1579735641] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579736010] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (139288 kB) free!
[1579736130] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193392 kB) free!
[1579736251] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173432 kB) free!
[1579736810] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 691 processes
[1579736841] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(76) < 2020/01/23 05:13:23 [error] 18357#0: *42755651 connect() failed (111: Connection refused) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579736880] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 47.2% (11644500 kB) free.
[1579736930] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 396 processes
[1579736961] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(107) < 2020/01/23 05:18:41 [error] 18370#0: *42756919 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579736981] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19122632 kB) free!
[1579737081] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579737101] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32516040 kB) free.
[1579737241] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.8% (21767040 kB) free.
[1579737563] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(16) < 203.212.70.170 [23/Jan/2020:05:21:47 +0530] TCP 502 0 0 180.000 "10.147.212.155:3001, 10.147.212.143:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 0.000"
[1579737683] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579738301] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1579739720] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (12400260 kB) free!
[1579740320] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 28.3% (37364980 kB) free.
[1579741436] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1579741556] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 384 processes
[1579741883] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(31) < 172.105.59.173 [23/Jan/2020:06:38:26 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579742003] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579743681] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(112) < 2020/01/23 07:05:23 [error] 18370#0: *42779229 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:154/0, bytes from/to upstream:0/154
[1579743701] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26325272 kB) free!
[1579743801] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579743821] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.2% (22703008 kB) free!
[1579743941] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19173536 kB) free!
[1579744541] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34316468 kB) free.
[1579745141] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (25734804 kB) free!
[1579745261] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32517216 kB) free.
[1579745601] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2020/01/23 07:40:08 [error] 18356#0: *42786310 connect() failed (111: Connection refused) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579745721] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579745756] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 407 processes
[1579745876] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 406 processes
[1579746203] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6) < 203.212.70.170 [23/Jan/2020:07:48:20 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1579746322] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(162) < 2020/01/23 07:46:49 [error] 18359#0: *42787795 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:74/0, bytes from/to upstream:0/74
[1579746324] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579746442] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579746923] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4) < 49.50.64.122 [23/Jan/2020:08:04:38 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579747042] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(269) < 2020/01/23 08:05:16 [error] 18359#0: *42791198 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1579747043] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(98) < 203.212.70.170 [23/Jan/2020:08:06:33 +0530] TCP 502 0 0 180.000 "10.147.212.143:3001, 10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 0.000"
[1579747077] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(138) < 103.16.101.72 [23/Jan/2020:08:03:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1579747162] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579747163] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(1) < 203.212.70.170 [23/Jan/2020:08:07:33 +0530] TCP 502 0 0 240.000 "10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.142:3001, 10.147.212.143:3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 60.000"
[1579747197] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579747625] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1579747745] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1579747763] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579750163] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(26) < 2020/01/23 08:57:39 [error] 18358#0: *42804500 connect() failed (111: Connection refused) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579750283] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579750883] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(206) < 2020/01/23 09:10:06 [error] 18358#0: *42807529 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:62/0, bytes from/to upstream:0/62
[1579751003] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579751963] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 203.212.70.170 [23/Jan/2020:09:24:12 +0530] TCP 502 0 0 63.007 "10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "3.007, 60.000, 0.000"
[1579752084] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579752461] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25542168 kB) free!
[1579752581] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.4% (21676480 kB) free!
[1579752701] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18019652 kB) free!
[1579753901] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1579753901] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (8264616 kB) free!
[1579754687] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.8% (4859068 kB) free!
[1579754807] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.2% (4708748 kB) free!
[1579754927] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.8% (4615628 kB) free!
[1579755101] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.8% (48482336 kB) free.
[1579756081] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.7% (3633396 kB) free!
[1579756201] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.3% (3525788 kB) free!
[1579756283] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(654) < 103.255.146.164 [23/Jan/2020:10:40:54 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1579756301] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26244328 kB) free!
[1579756321] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.7% (3619060 kB) free!
[1579756403] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579756421] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.1% (44919140 kB) free.
[1579757003] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(647) < 103.255.146.164 [23/Jan/2020:10:52:00 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1579757123] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579757327] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2194052 kB) free!
[1579757521] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2401776 kB) free!
[1579758138] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.8% (21758204 kB) free.
[1579758200] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.4% (20825620 kB) free.
[1579758221] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22205884 kB) free!
[1579758235] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.8% (40326696 kB) free.
[1579758266] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.4% (21674592 kB) free.
[1579758299] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 55.3% (72851620 kB) free.
[1579759409] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.1% (4230736 kB) free!
[1579759529] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (969552 kB) free!
[1579759649] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1579759649] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (311312 kB) free!
[1579759867] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.4% (21063568 kB) free.
[1579759881] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.6% (3953204 kB) free!
[1579760001] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.1% (3613796 kB) free!
[1579760003] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/23 11:39:44 [error] 18358#0: *42848936 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net"
[1579760121] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.3% (3216252 kB) free!
[1579760123] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579760721] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1757568 kB) free!
[1579760723] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/23 11:46:12 [error] 18359#0: *42852384 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1579760843] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579761322] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 74.6% (16741180 kB) free.
[1579761444] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/23 12:01:13 [error] 18369#0: *42859814 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1579761564] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579762867] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.6% (2617240 kB) free!
[1579762987] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (1403996 kB) free!
[1579763107] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (201020 kB) free!
[1579763378] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (182348 kB) free!
[1579763499] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (182748 kB) free!
[1579763619] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (183992 kB) free!
[1579764021] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.4% (22050180 kB) free.
[1579764219] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.2% (14334992 kB) free.
[1579764311] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1579765438] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.4% (8559120 kB) free!
[1579765468] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164356 kB) free!
[1579765511] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20262820 kB) free!
[1579765558] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.9% (8838608 kB) free!
[1579765588] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (202628 kB) free!
[1579765632] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (42091740 kB) free.
[1579765678] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.9% (8838560 kB) free!
[1579765708] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170768 kB) free!
[1579765827] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.9% (3674620 kB) free!
[1579765922] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.2% (22353960 kB) free.
[1579765947] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.1% (3472328 kB) free!
[1579765956] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.9% (40373356 kB) free.
[1579765988] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.1% (21975660 kB) free.
[1579766231] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18433004 kB) free!
[1579766351] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40142992 kB) free.
[1579766951] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16457972 kB) free!
[1579767071] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39867016 kB) free.
[1579767199] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167884 kB) free!
[1579767311] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.4% (21821692 kB) free.
[1579767671] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15805324 kB) free!
[1579767791] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39298376 kB) free.
[1579768391] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15118744 kB) free!
[1579768511] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38565636 kB) free.
[1579769111] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14738080 kB) free!
[1579769231] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34195132 kB) free.
[1579769652] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579769772] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579769831] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14783644 kB) free!
[1579769892] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1579769951] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.6% (25814496 kB) free!
[1579770071] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33402240 kB) free.
[1579770241] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169352 kB) free!
[1579770317] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1579770322] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4369172 kB) free!
[1579770361] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163528 kB) free!
[1579770442] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.7% (4194344 kB) free!
[1579770467] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.3% (21892624 kB) free.
[1579770492] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1579770497] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 74.3% (16691816 kB) free.
[1579770648] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9563104 kB) free!
[1579770672] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9930760 kB) free!
[1579770768] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.0% (9374248 kB) free!
[1579770792] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32830092 kB) free.
[1579770888] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.7% (9205584 kB) free!
[1579770917] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.4% (4300256 kB) free!
[1579771037] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.8% (4153028 kB) free!
[1579771067] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169792 kB) free!
[1579771157] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.3% (4014100 kB) free!
[1579771188] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179308 kB) free!
[1579771308] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181660 kB) free!
[1579771392] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9328528 kB) free!
[1579771512] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32976836 kB) free.
[1579772112] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8495088 kB) free!
[1579772232] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33480564 kB) free.
[1579772357] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1921148 kB) free!
[1579772680] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1579772800] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1579772832] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8776356 kB) free!
[1579772876] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579772920] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1579772926] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.1% (21502420 kB) free.
[1579772952] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.1% (23845700 kB) free!
[1579772986] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 56.3% (74246332 kB) free.
[1579772996] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1579773024] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.0% (21810268 kB) free.
[1579773175] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177452 kB) free!
[1579773288] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4699528 kB) free!
[1579773295] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186872 kB) free!
[1579773336] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.1% (21829252 kB) free.
[1579773374] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.9% (45301304 kB) free.
[1579773520] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1579773762] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9349372 kB) free!
[1579773882] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.0% (9341616 kB) free!
[1579774002] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.9% (9336836 kB) free!
[1579774092] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579774196] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1579774213] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579774228] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193380 kB) free!
[1579774316] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1579774330] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.1% (4442640 kB) free!
[1579774333] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1579774348] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164308 kB) free!
[1579774436] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1579774450] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.7% (4331904 kB) free!
[1579774468] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174768 kB) free!
[1579774570] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.1% (4186204 kB) free!
[1579775636] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1579775929] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179936 kB) free!
[1579776050] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178148 kB) free!
[1579776170] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1579776170] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177064 kB) free!
[1579776370] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166836 kB) free!
[1579777190] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23759092 kB) free!
[1579777310] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.6% (19215268 kB) free!
[1579777430] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (15062192 kB) free!
[1579777436] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1579777541] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192776 kB) free!
[1579777557] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1579777632] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.8% (21409820 kB) free.
[1579777661] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192040 kB) free!
[1579777675] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 85.7% (21001432 kB) free.
[1579777677] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1579777720] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1579777724] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.1% (22339384 kB) free.
[1579777774] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1579777774] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.0% (21818608 kB) free.
[1579777840] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1579778277] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1579778277] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1579778630] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28415732 kB) free.
[1579778877] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1579778997] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1579779117] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1579779578] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168968 kB) free!
[1579779698] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172516 kB) free!
[1579779818] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180672 kB) free!
[1579779830] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11256676 kB) free!
[1579779950] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37320040 kB) free.
[1579780002] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3633908 kB) free!
[1579780550] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12723212 kB) free!
[1579780602] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1579780638] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.8% (3900212 kB) free!
[1579780670] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8510564 kB) free!
[1579780722] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1579780758] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.7% (3380484 kB) free!
[1579780790] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31999080 kB) free.
[1579780878] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.7% (2631732 kB) free!
[1579781390] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (2435540 kB) free!
[1579781478] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (1225712 kB) free!
[1579781510] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (408056 kB) free!
[1579781517] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1579781630] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38198976 kB) free.
[1579781897] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9836300 kB) free!
[1579781916] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.7% (21647616 kB) free.
[1579781922] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1579781952] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.9% (40351032 kB) free.
[1579781981] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.6% (22203044 kB) free.
[1579782018] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 91.4% (45021528 kB) free.
[1579782042] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1579782230] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6810228 kB) free!
[1579782350] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2290008 kB) free!
[1579782470] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38414104 kB) free.
[1579783070] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8033904 kB) free!
[1579783190] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3664576 kB) free!
[1579783310] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394744 kB) free!
[1579783523] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(10) < 139.59.21.208 [23/Jan/2020:18:15:18 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1579783638] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1579783638] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 43.6% (57516016 kB) free.
[1579783643] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579783719] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194856 kB) free!
[1579783839] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167060 kB) free!
[1579783959] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178812 kB) free!
[1579784243] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(11) < 183.82.105.122 [23/Jan/2020:18:20:58 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1579784279] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196664 kB) free!
[1579784364] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579784389] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166552 kB) free!
[1579784399] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193884 kB) free!
[1579784440] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1579784509] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176092 kB) free!
[1579784519] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178460 kB) free!
[1579784560] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1579784629] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177580 kB) free!
[1579784963] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(18) < 94.237.65.17 [23/Jan/2020:18:33:10 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1579785083] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579785117] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1579785120] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 82.7% (20276320 kB) free.
[1579785237] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1579785357] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1579785683] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(86) < 183.82.105.122 [23/Jan/2020:18:49:18 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1579785803] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579786648] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26211076 kB) free!
[1579786727] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195928 kB) free!
[1579786733] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.0% (22053880 kB) free.
[1579786768] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.0% (22464664 kB) free!
[1579786847] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.7% (21397788 kB) free.
[1579786847] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183212 kB) free!
[1579786877] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.7% (22727900 kB) free.
[1579786888] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 39.6% (52191984 kB) free.
[1579787003] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(165) < 92.42.104.166 [23/Jan/2020:19:09:37 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1579787123] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579787158] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1579787365] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(58) < 2020/01/23 19:16:52 [error] 18362#0: *43059584 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579787486] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579787723] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6) < 182.18.144.246 [23/Jan/2020:19:24:15 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579787843] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(5) < 203.212.70.170 [23/Jan/2020:19:26:07 +0530] TCP 502 0 0 180.001 "10.147.212.142:3001, 10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.001, 0.000"
[1579787964] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579789763] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1352) < 103.255.146.164 [23/Jan/2020:19:51:05 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1579789883] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579791133] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579791688] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26132336 kB) free!
[1579791808] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.3% (50455600 kB) free.
[1579792157] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9843080 kB) free!
[1579792277] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (9843096 kB) free!
[1579792398] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 20.0% (9843688 kB) free!
[1579793608] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1579794733] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579794853] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579794973] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1579795141] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1579795283] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(298) < 49.50.86.152 [23/Jan/2020:21:27:21 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579795285] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/23 21:28:12 [error] 18363#0: *43091382 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579795404] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579795406] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579795847] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1579796173] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579797358] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1579801573] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579801693] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579801813] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1579801813] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1579802603] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(230) < 182.18.144.246 [23/Jan/2020:23:29:10 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579802724] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found