[1561919681] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8228092 kB) free!
[1561919701] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561919802] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12488252 kB) free.
[1561919821] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561920003] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188076 kB) free!
[1561920123] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (201476 kB) free!
[1561920243] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182132 kB) free!
[1561920401] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8248464 kB) free!
[1561920421] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561920521] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12493620 kB) free.
[1561920541] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561921121] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8134192 kB) free!
[1561921141] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561921241] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12475808 kB) free.
[1561921261] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561921365] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.5% (2356228 kB) free!
[1561921485] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.5% (2363280 kB) free!
[1561921605] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.6% (2368200 kB) free!
[1561921836] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561921841] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7978912 kB) free!
[1561921861] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561921961] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12483828 kB) free.
[1561921981] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561922266] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 492 processes
[1561922326] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 492 processes
[1561922561] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7830628 kB) free!
[1561922681] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12486136 kB) free.
[1561923181] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561923225] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 17.3% (4231348 kB) free!
[1561923281] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7695152 kB) free!
[1561923285] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 17.3% (4228560 kB) free!
[1561923301] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561923346] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 17.3% (4234976 kB) free!
[1561923401] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12477192 kB) free.
[1561923406] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 17.3% (4235776 kB) free!
[1561923902] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561924022] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561924602] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9160108 kB) free!
[1561924622] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561924722] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12485556 kB) free.
[1561924742] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561925322] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9125108 kB) free!
[1561925342] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561925442] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12478388 kB) free.
[1561925462] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561925661] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 191.18, 59.70, 21.17
[1561925781] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 256.96, 128.36, 51.00
[1561925854] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561925901] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 265.88, 174.32, 77.22
[1561926042] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9070352 kB) free!
[1561926062] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561926162] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12483492 kB) free.
[1561926182] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561926762] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (8998248 kB) free!
[1561926783] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561926882] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12486720 kB) free.
[1561926903] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561927005] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (233692 kB) free!
[1561927482] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8885012 kB) free!
[1561927503] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1561927602] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12479932 kB) free.
[1561927623] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561927701] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.60, 1.13, 15.18
[1561928202] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8811824 kB) free!
[1561928223] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561928322] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12484116 kB) free.
[1561928343] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561928901] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.15, 0.34, 4.45
[1561928923] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8786208 kB) free!
[1561928943] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561929043] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12475288 kB) free.
[1561929063] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561929643] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8738048 kB) free!
[1561929664] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561929763] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12485064 kB) free.
[1561929784] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561930101] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 82.69, 51.92, 23.76
[1561930221] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 83.36, 62.32, 30.98
[1561930341] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1561930341] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 20.60, 48.93, 30.16
[1561930363] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8741784 kB) free!
[1561930384] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561930443] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.5% (2844896 kB) free!
[1561930465] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 35.6% (17522876 kB) free.
[1561930504] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1561930538] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.2% (23496060 kB) free.
[1561930636] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 73.8% (11988012 kB) free.
[1561930941] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.71, 7.06, 16.05
[1561931104] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561931224] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561931824] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561931945] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561932141] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.72, 0.74, 4.83
[1561932544] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561932664] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561933264] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561933384] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561933984] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561934104] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561934665] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8470796 kB) free!
[1561934785] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12791864 kB) free.
[1561935304] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561935385] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8214296 kB) free!
[1561935425] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561935505] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12799936 kB) free.
[1561936024] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561936105] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8020776 kB) free!
[1561936144] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561936225] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12793292 kB) free.
[1561936744] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561936825] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7992468 kB) free!
[1561936864] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561936945] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12799100 kB) free.
[1561937464] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561937584] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561938145] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9452824 kB) free!
[1561938184] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561938266] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12804740 kB) free.
[1561938305] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561938904] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561939024] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561939093] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4656088 kB) free!
[1561939144] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561939213] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.1% (4185992 kB) free!
[1561939238] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 207.44, 109.58, 44.28
[1561939242] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.3% (20917112 kB) free.
[1561939358] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 212.33, 143.83, 64.74
[1561939478] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 213.51, 166.94, 82.80
[1561939744] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561940065] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9701764 kB) free!
[1561940185] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8157268 kB) free!
[1561940305] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12268072 kB) free.
[1561940344] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561940464] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561940584] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561941185] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1561941185] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561941279] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.24, 0.71, 12.15
[1561941454] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 50.83, 40.87, 22.05
[1561941505] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9190244 kB) free!
[1561941574] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 50.80, 44.14, 25.53
[1561941625] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12735048 kB) free.
[1561941694] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 141.77, 71.15, 37.03
[1561941785] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561941905] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561942025] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561942252] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (4540184 kB) free!
[1561942372] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 29.3% (7174200 kB) free.
[1561942479] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.16, 0.19, 3.48
[1561942941] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561943225] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561943356] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.0% (3082936 kB) free!
[1561943476] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.9% (3067656 kB) free!
[1561943596] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.7% (3039860 kB) free!
[1561944094] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.68, 2.89, 18.58
[1561944425] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561944546] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561945145] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561945265] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561945825] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8354720 kB) free!
[1561945895] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.32, 2.44, 4.82
[1561945945] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12578476 kB) free.
[1561946465] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561946545] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8089032 kB) free!
[1561946585] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561946665] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12095952 kB) free.
[1561947279] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 116.95, 88.44, 42.79
[1561947399] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 43.05, 76.65, 44.49
[1561947519] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1561947519] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 5.88, 51.31, 39.11
[1561947786] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561947906] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561948026] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561948719] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.04, 1.01, 10.82
[1561949826] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561949919] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.15, 0.19, 3.10
[1561950196] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (1538548 kB) free!
[1561950865] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9132204 kB) free!
[1561950985] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12565908 kB) free.
[1561951585] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9019204 kB) free!
[1561951627] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561951705] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12572676 kB) free.
[1561951747] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561952305] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8971128 kB) free!
[1561952347] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561952425] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12562644 kB) free.
[1561952467] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561953025] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9599116 kB) free!
[1561953067] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561953145] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.1% (8442148 kB) free!
[1561953172] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4908280 kB) free!
[1561953187] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561953265] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12569832 kB) free.
[1561953292] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (4858520 kB) free!
[1561953307] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561953412] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (4810792 kB) free!
[1561953865] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9455460 kB) free!
[1561953907] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561953985] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.9% (8333180 kB) free!
[1561954027] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561954106] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12566616 kB) free.
[1561954148] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561954705] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8725360 kB) free!
[1561954747] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561954825] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12557596 kB) free.
[1561954867] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561955425] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8569396 kB) free!
[1561955467] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561955545] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12564412 kB) free.
[1561955587] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561956145] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8583088 kB) free!
[1561956187] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561956265] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12566520 kB) free.
[1561956307] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561956865] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8334804 kB) free!
[1561956985] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12561732 kB) free.
[1561957341] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.56, 11.31, 5.49
[1561957461] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.53, 7.74, 4.88
[1561957507] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561957585] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8362468 kB) free!
[1561957627] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561957705] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12567992 kB) free.
[1561958227] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561958305] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8060816 kB) free!
[1561958347] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561958426] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12553448 kB) free.
[1561958927] SERVICE ALERT: sig2-vm1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 312 processes
[1561958947] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561959047] SERVICE ALERT: sig2-vm1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 312 processes
[1561959067] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561959167] SERVICE ALERT: sig2-vm1;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1561959200] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (161792 kB) free!
[1561959320] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (151392 kB) free!
[1561959345] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188412 kB) free!
[1561959415] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (205144 kB) free!
[1561959440] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (140828 kB) free!
[1561959465] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (227972 kB) free!
[1561959585] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179276 kB) free!
[1561959625] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9194220 kB) free!
[1561959649] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185252 kB) free!
[1561959667] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561959745] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12562520 kB) free.
[1561959769] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175000 kB) free!
[1561959788] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561959889] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185164 kB) free!
[1561960345] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9184860 kB) free!
[1561960387] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561960465] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12561056 kB) free.
[1561960507] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561960719] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561960907] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (1482404 kB) free!
[1561961027] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (178300 kB) free!
[1561961065] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9068692 kB) free!
[1561961107] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561961147] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (175552 kB) free!
[1561961185] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12550360 kB) free.
[1561961227] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561961785] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8896380 kB) free!
[1561961827] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561961905] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12559372 kB) free.
[1561961947] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1561961947] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561962339] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.7% (23128076 kB) free.
[1561962457] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.3% (22384880 kB) free.
[1561962552] SERVICE ALERT: sig2-vm1;Memory;OK;HARD;1;OK - 94.9% (6977924 kB) free.
[1561962578] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4738864 kB) free!
[1561962698] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (3556924 kB) free!
[1561962757] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 79.2% (12872140 kB) free.
[1561962818] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (2862916 kB) free!
[1561963105] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8761552 kB) free!
[1561963148] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561963226] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12543484 kB) free.
[1561963268] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561963825] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8348016 kB) free!
[1561963945] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12551020 kB) free.
[1561964150] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (4398260 kB) free!
[1561964230] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1561964270] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (1404872 kB) free!
[1561964350] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1561964390] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193660 kB) free!
[1561964468] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561964470] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1561964545] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7763700 kB) free!
[1561964588] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561964665] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12550268 kB) free.
[1561965188] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561965308] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1561965865] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8845604 kB) free!
[1561965908] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561965985] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12537328 kB) free.
[1561966028] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1561966461] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.03, 5.54, 5.11
[1561966582] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.52, 5.53, 5.16
[1561966585] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8575348 kB) free!
[1561966629] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561966702] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.23, 5.45, 5.18
[1561966705] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12546812 kB) free.
[1561966749] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561966791] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.49, 7.31, 6.50
[1561966911] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.95, 6.49, 6.30
[1561967031] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.71, 6.76, 6.40
[1561967302] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 5.32, 4.48, 4.67
[1561967305] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9321320 kB) free!
[1561967349] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561967425] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7779732 kB) free!
[1561967469] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561967546] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12540740 kB) free.
[1561968069] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561968145] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8718408 kB) free!
[1561968189] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561968265] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12546560 kB) free.
[1561968309] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561968406] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 27.4% (6717456 kB) free.
[1561968502] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.45, 5.45, 5.20
[1561968622] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.01, 5.65, 5.30
[1561968742] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.42, 5.87, 5.42
[1561968909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561969030] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561969150] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561969369] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 14.3% (2320204 kB) free!
[1561969489] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 12.7% (2071968 kB) free!
[1561969609] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 11.2% (1815996 kB) free!
[1561969739] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (487384 kB) free!
[1561969750] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1561969773] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 74.2% (18301312 kB) free.
[1561969858] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.4% (6797100 kB) free.
[1561970065] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9395584 kB) free!
[1561970103] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.6% (39250736 kB) free.
[1561970185] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7660932 kB) free!
[1561970189] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.2% (22341988 kB) free.
[1561970209] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (861452 kB) free!
[1561970269] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.7% (14409652 kB) free.
[1561970305] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11718516 kB) free.
[1561970350] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561970379] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188900 kB) free!
[1561970470] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561970499] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185328 kB) free!
[1561970536] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1561970536] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 78.2% (12707976 kB) free.
[1561970590] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561970619] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 47.0% (11607980 kB) free.
[1561970876] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (157568 kB) free!
[1561970996] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (181596 kB) free!
[1561971116] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (164456 kB) free!
[1561971142] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1561971142] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.24, 4.40, 4.74
[1561971226] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (1295580 kB) free!
[1561971346] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193980 kB) free!
[1561971466] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1561971466] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191792 kB) free!
[1561971505] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8514372 kB) free!
[1561971625] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12502968 kB) free.
[1561972267] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 15.4% (1131960 kB) free!
[1561972342] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.96, 5.50, 5.05
[1561972387] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 11.2% (826408 kB) free!
[1561972462] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.98, 5.33, 5.04
[1561972507] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (557344 kB) free!
[1561972582] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.91, 5.18, 5.02
[1561972990] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1561972990] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561973182] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 5.06, 4.98, 4.97
[1561973425] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7753896 kB) free!
[1561973545] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12510332 kB) free.
[1561973707] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.9% (6835184 kB) free.
[1561973782] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.55, 5.51, 5.23
[1561973902] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 67.39, 25.23, 12.19
[1561974023] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 86.17, 46.33, 21.50
[1561974095] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 73.51, 71.27, 40.29
[1561974215] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 12.40, 48.66, 35.77
[1561974335] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.40, 33.63, 31.83
[1561974426] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 0.01, 0.07, 0.08
[1561974488] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 242 processes
[1561974608] SERVICE ALERT: sig4-vm2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1561974631] SERVICE ALERT: sig4-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 122827 MB (95% inode=99%):
[1561974745] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8896920 kB) free!
[1561974754] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 243 processes
[1561974790] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561974800] SERVICE ALERT: sig4-vm2;Memory;OK;HARD;1;OK - 82.2% (6006676 kB) free.
[1561974811] SERVICE ALERT: sig4-vm2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1561974865] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12511140 kB) free.
[1561974910] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1561974935] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.68, 6.97, 18.04
[1561975223] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.30, 8.17, 17.55
[1561975346] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.4% (2821408 kB) free!
[1561975465] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8594152 kB) free!
[1561975466] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.7% (2559596 kB) free!
[1561975510] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561975585] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12505656 kB) free.
[1561975586] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 14.2% (2304360 kB) free!
[1561975631] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561975666] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (4494512 kB) free!
[1561975787] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4397656 kB) free!
[1561975907] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.4% (3534060 kB) free!
[1561976185] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8297292 kB) free!
[1561976187] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (1325776 kB) free!
[1561976305] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12510228 kB) free.
[1561976735] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.98, 2.79, 4.92
[1561976830] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561976905] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7862424 kB) free!
[1561976950] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561977025] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.4% (12495504 kB) free.
[1561977107] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168864 kB) free!
[1561977550] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561977670] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561978049] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (164988 kB) free!
[1561978169] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (180728 kB) free!
[1561978223] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 5.08, 4.21, 4.97
[1561978270] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1561978290] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (814220 kB) free!
[1561978390] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561978511] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561979108] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.1% (7941012 kB) free!
[1561979111] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561979228] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.1% (7939528 kB) free!
[1561979348] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.1% (7934564 kB) free!
[1561979423] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.79, 5.06, 5.04
[1561979424] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.2% (2635812 kB) free!
[1561979425] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8756152 kB) free!
[1561979543] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.80, 4.67, 4.91
[1561979544] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 15.7% (2545228 kB) free!
[1561979545] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12472376 kB) free.
[1561979664] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.1% (2456188 kB) free!
[1561979711] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1561979735] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 310.56, 214.40, 99.09
[1561979831] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561979855] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 312.41, 246.84, 124.95
[1561979919] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 98.51, 52.74, 21.65
[1561979951] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561979975] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1561979975] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 312.28, 268.51, 147.67
[1561980040] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 101.65, 69.06, 31.40
[1561980160] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 85.40, 76.20, 38.71
[1561980398] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.2% (2491968 kB) free!
[1561980518] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (2404688 kB) free!
[1561980548] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3803508 kB) free!
[1561980638] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2337524 kB) free!
[1561980907] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 18.0% (1321920 kB) free!
[1561981027] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 16.4% (1207204 kB) free!
[1561981066] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1561981147] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 14.8% (1087336 kB) free!
[1561981345] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8040136 kB) free!
[1561981360] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.42, 1.78, 10.94
[1561981431] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.69, 4.34, 4.60
[1561981465] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12469916 kB) free.
[1561981747] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1561981747] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (403700 kB) free!
[1561982031] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.94, 7.70, 6.20
[1561982064] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1397636 kB) free!
[1561982152] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.17, 6.97, 6.12
[1561982272] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.21, 6.00, 5.87
[1561982352] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1561982375] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.76, 3.41, 14.25
[1561982560] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.13, 0.42, 3.31
[1561982665] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8823136 kB) free!
[1561982785] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12458256 kB) free.
[1561982952] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1561983072] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561983192] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561984072] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.41, 4.00, 4.47
[1561984175] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.53, 3.14, 4.96
[1561984465] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 84.4% (13711328 kB) free.
[1561984748] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.9% (6829720 kB) free.
[1561984943] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 5.68, 64.52, 54.58
[1561985063] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.96, 43.73, 48.18
[1561985184] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.39, 29.70, 42.50
[1561985185] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8813832 kB) free!
[1561985207] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 580 processes
[1561985305] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12467104 kB) free.
[1561985327] SERVICE ALERT: sservercluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 579 processes
[1561985447] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 322 processes
[1561985560] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.66, 18.61, 18.48
[1561985680] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.36, 12.55, 16.28
[1561985800] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1561985800] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.41, 8.51, 14.35
[1561985905] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8738920 kB) free!
[1561986025] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12452772 kB) free.
[1561986192] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561986385] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.04, 5.97, 15.61
[1561986625] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8663888 kB) free!
[1561986745] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12456736 kB) free.
[1561987000] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.44, 0.55, 4.20
[1561987345] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8277276 kB) free!
[1561987465] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12455896 kB) free.
[1561987587] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1561987992] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561988065] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8304624 kB) free!
[1561988112] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1561988186] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12452860 kB) free.
[1561988712] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1561988785] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9141532 kB) free!
[1561988832] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561988905] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7689708 kB) free!
[1561988952] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1561989025] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12444644 kB) free.
[1561989552] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1561989625] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8445040 kB) free!
[1561989672] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561989745] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12450320 kB) free.
[1561989793] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561990345] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9093260 kB) free!
[1561990392] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561990465] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12455640 kB) free.
[1561990512] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1561991065] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9070064 kB) free!
[1561991112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561991185] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12444800 kB) free.
[1561991232] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561991785] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9085908 kB) free!
[1561991832] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1561991906] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12455456 kB) free.
[1561991952] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561992385] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.36, 2.93, 4.45
[1561992505] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8933340 kB) free!
[1561992552] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561992625] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12442076 kB) free.
[1561992673] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1561993225] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8944444 kB) free!
[1561993272] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561993345] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12450540 kB) free.
[1561993392] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561993495] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1561993585] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.07, 5.33, 5.12
[1561993615] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1561993705] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.73, 5.12, 5.07
[1561993748] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1561993775] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 134.54, 105.52, 53.66
[1561993825] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.23, 5.12, 5.07
[1561993895] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 134.97, 115.35, 63.56
[1561993945] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8880772 kB) free!
[1561993992] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561994015] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 32.44, 87.49, 60.02
[1561994065] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12451708 kB) free.
[1561994112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1561994232] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561994714] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47041 MB (20% inode=99%):
[1561994835] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47027 MB (20% inode=99%):
[1561994955] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47023 MB (20% inode=99%):
[1561995215] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.84, 4.80, 19.22
[1561995555] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47767 MB (21% inode=99%):
[1561996225] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 5.06, 4.00, 4.30
[1561996633] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1561996755] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46606 MB (20% inode=99%):
[1561996875] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46541 MB (20% inode=99%):
[1561996995] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46540 MB (20% inode=99%):
[1561997065] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9001636 kB) free!
[1561997186] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12438964 kB) free.
[1561997425] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.65, 5.40, 5.06
[1561997545] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.00, 5.29, 5.06
[1561997615] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.89, 2.96, 4.41
[1561997665] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.30, 5.32, 5.10
[1561997785] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8814044 kB) free!
[1561997800] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1561997833] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561997905] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12444080 kB) free.
[1561997953] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561998553] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1561998673] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1561998793] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1561999393] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1561999465] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.07, 3.44, 4.57
[1561999994] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562000114] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562000714] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562000729] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 65.5% (10638984 kB) free.
[1562000750] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 67.3% (10931632 kB) free.
[1562000776] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.5% (21445452 kB) free.
[1562000834] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562000843] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 58.2% (14369336 kB) free.
[1562000858] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.9% (22523704 kB) free.
[1562000880] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.7% (39278576 kB) free.
[1562000930] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 75.9% (12340292 kB) free.
[1562001215] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 106.82, 29.55, 11.95
[1562001335] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 201.90, 91.44, 36.76
[1562001356] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (169844 kB) free!
[1562001434] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562001448] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (260748 kB) free!
[1562001455] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 38.28, 69.65, 35.77
[1562001476] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (153120 kB) free!
[1562001554] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1562001568] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (279652 kB) free!
[1562001596] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (181196 kB) free!
[1562001689] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (212860 kB) free!
[1562002055] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.49, 11.58, 19.98
[1562002105] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8919020 kB) free!
[1562002154] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562002225] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12419976 kB) free.
[1562002274] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562002600] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 75.74, 48.46, 21.30
[1562002720] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 76.20, 57.67, 27.97
[1562002826] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8878500 kB) free!
[1562002841] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 12.49, 40.42, 25.34
[1562002875] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562002946] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12429552 kB) free.
[1562002995] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562003065] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 115.21, 205.48, 118.91
[1562003185] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 18.61, 138.68, 104.93
[1562003305] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.74, 93.43, 92.46
[1562003441] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.34, 5.68, 13.41
[1562003546] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8939988 kB) free!
[1562003595] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562003666] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12430988 kB) free.
[1562003715] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562004266] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8952540 kB) free!
[1562004315] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562004386] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12423268 kB) free.
[1562004435] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562004455] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.88, 2.83, 4.05
[1562004641] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.30, 0.39, 3.90
[1562004986] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8739496 kB) free!
[1562005035] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562005106] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.00, 5.13, 17.37
[1562005106] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.2% (12430320 kB) free.
[1562005156] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562005706] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8724104 kB) free!
[1562005755] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes