[1552329782] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 71.4% (11608368 kB) free.
[1552329944] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9109384 kB) free!
[1552330064] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.2% (6510364 kB) free!
[1552330184] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3890388 kB) free!
[1552330684] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552331284] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552331404] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552331524] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552331984] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.8% (11212368 kB) free.
[1552332257] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9579420 kB) free!
[1552332377] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.8% (7810644 kB) free!
[1552332497] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.1% (15339972 kB) free.
[1552333097] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9774396 kB) free!
[1552333218] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.5% (8142168 kB) free!
[1552333325] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1552333338] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.1% (15347868 kB) free.
[1552333925] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552334045] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552334165] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552334765] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552334885] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552335005] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1552335605] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552335725] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552335738] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9083712 kB) free!
[1552335846] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1552335858] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.1% (15341112 kB) free.
[1552336184] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (4373724 kB) free!
[1552336304] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (1899032 kB) free!
[1552336424] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.0% (11321984 kB) free.
[1552336445] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552336565] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552336685] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552337024] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (2759368 kB) free!
[1552337145] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.0% (11349648 kB) free.
[1552338258] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9718408 kB) free!
[1552338344] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.3% (8527736 kB) free!
[1552338378] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8010380 kB) free!
[1552338464] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.8% (5834920 kB) free!
[1552338485] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1552338499] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.1% (15341248 kB) free.
[1552338584] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (3080960 kB) free!
[1552339085] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552339184] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.1% (11408384 kB) free.
[1552339205] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552339326] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1552339925] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552340045] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552340165] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552340384] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8990720 kB) free!
[1552340504] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.9% (6369220 kB) free!
[1552340624] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3657164 kB) free!
[1552340765] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552340885] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552340898] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9226164 kB) free!
[1552341005] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1552341018] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.1% (15329588 kB) free.
[1552341224] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.2% (11431020 kB) free.
[1552341605] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552341726] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552341846] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552342424] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.3% (8999896 kB) free!
[1552342545] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.2% (6495104 kB) free!
[1552342665] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3836832 kB) free!
[1552343265] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.3% (11483208 kB) free.
[1552343646] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1552344246] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552344366] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552344465] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9078388 kB) free!
[1552344486] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1552344585] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.0% (6427588 kB) free!
[1552344705] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3683584 kB) free!
[1552345087] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552345207] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552345219] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9092368 kB) free!
[1552345327] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1552345339] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.1% (15327652 kB) free.
[1552345905] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.7% (6259616 kB) free!
[1552345927] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552346047] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552346167] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552346505] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.5% (10107852 kB) free.
[1552347967] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552348306] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3819388 kB) free!
[1552348426] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.0% (11849400 kB) free.
[1552348567] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552348687] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552348807] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552349026] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4313208 kB) free!
[1552349146] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.1% (11887244 kB) free.
[1552349407] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552349527] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552349539] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9034632 kB) free!
[1552349647] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1552349659] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.1% (15320776 kB) free.
[1552349746] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (4881944 kB) free!
[1552349866] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.2% (11918008 kB) free.
[1552350247] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552350368] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552350466] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 10.5% (5171952 kB) free!
[1552350488] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552350586] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.2% (11941912 kB) free.
[1552351088] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552351186] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 11.0% (5417320 kB) free!
[1552351208] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552351306] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (2898756 kB) free!
[1552351328] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552351426] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.5% (11083752 kB) free.
[1552352059] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9552364 kB) free!
[1552352179] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.7% (7760904 kB) free!
[1552352299] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.1% (15315652 kB) free.
[1552352899] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9788544 kB) free!
[1552353019] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8204424 kB) free!
[1552353128] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1552353139] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.1% (15310704 kB) free.
[1552353729] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552353826] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (3197824 kB) free!
[1552353849] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552353946] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.2% (11924856 kB) free.
[1552353969] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1552354546] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (3401144 kB) free!
[1552354569] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552354666] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.2% (11941020 kB) free.
[1552354689] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552354809] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1552355409] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552355530] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552355540] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9262336 kB) free!
[1552355650] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1552355660] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.1% (15323872 kB) free.
[1552356250] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552356370] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552356490] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552358266] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.0% (8887056 kB) free!
[1552358290] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1552358387] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 12.9% (6352756 kB) free!
[1552358507] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3756648 kB) free!
[1552358751] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1552358871] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1552358890] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552358991] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1552359011] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552359107] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.9% (11784628 kB) free.
[1552359131] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1552359731] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552359851] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552359860] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8986756 kB) free!
[1552359971] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1552359980] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 31.1% (15308892 kB) free.
[1552360307] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.3% (8504528 kB) free!
[1552360427] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.9% (5860476 kB) free!
[1552360547] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3171912 kB) free!
[1552360571] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1552360691] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552360811] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552361147] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.1% (11375036 kB) free.
[1552361412] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1552361532] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1552361652] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1552361747] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9180840 kB) free!
[1552361867] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.4% (6601840 kB) free!
[1552361987] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3935992 kB) free!
[1552361991] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1552362587] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.2% (11929984 kB) free.
[1552363788] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8106252 kB) free!
[1552363908] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.0% (5421360 kB) free!
[1552364028] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (2745268 kB) free!
[1552364180] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7414188 kB) free!
[1552364300] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.3% (14457188 kB) free.
[1552364628] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.8% (11215768 kB) free.
[1552365228] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9070648 kB) free!
[1552365348] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.0% (6389228 kB) free!
[1552365468] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3608344 kB) free!
[1552365849] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.50, 5.48, 5.08
[1552365969] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.78, 5.21, 5.03
[1552366089] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.43, 4.99, 4.97
[1552366191] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1552366311] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1552366431] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1552366668] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.4% (10531992 kB) free.
[1552366700] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6957340 kB) free!
[1552366820] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.0% (14292192 kB) free.
[1552367042] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.5% (3173736 kB) free!
[1552367110] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 22360 MB (10% inode=99%):
[1552367162] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.0% (3095076 kB) free!
[1552367268] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.0% (7866956 kB) free!
[1552367282] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.5% (3006056 kB) free!
[1552367388] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.3% (5082264 kB) free!
[1552367508] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (2233396 kB) free!
[1552367710] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 34072 MB (16% inode=99%):
[1552368708] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 20.0% (9832116 kB) free!
[1552369213] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4912128 kB) free!
[1552369220] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.5% (7164920 kB) free!
[1552369333] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (4839952 kB) free!
[1552369340] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.3% (14452028 kB) free.
[1552369453] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.4% (4778396 kB) free!
[1552369682] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1592656 kB) free!
[1552370289] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.20, 5.58, 5.20
[1552370362] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;1;WARNING - 19.9% (1460336 kB) free!
[1552370409] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.93, 5.71, 5.29
[1552370482] SERVICE ALERT: sig4-vm1;Memory;WARNING;SOFT;2;WARNING - 19.5% (1431428 kB) free!
[1552370529] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.44, 5.59, 5.30
[1552370602] SERVICE ALERT: sig4-vm1;Memory;WARNING;HARD;3;WARNING - 19.1% (1404892 kB) free!
[1552370710] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 22369 MB (10% inode=99%):
[1552371108] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (4130716 kB) free!
[1552371140] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6858000 kB) free!
[1552371160] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1552371252] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1552371260] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.9% (14236660 kB) free.
[1552371280] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1552371400] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1552371882] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1552371942] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1552372002] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1552372062] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1552372452] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1552372510] SERVICE FLAPPING ALERT: db1;HDD_Home status;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1552372510] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 29908 MB (14% inode=99%):
[1552372963] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;1;WARNING - 19.2% (1415276 kB) free!
[1552373002] SERVICE ALERT: sig4-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (737904 kB) free!
[1552373053] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2292676 kB) free!
[1552373083] SERVICE ALERT: sig3-vm1;Memory;WARNING;SOFT;2;WARNING - 19.1% (1402372 kB) free!
[1552373203] SERVICE ALERT: sig3-vm1;Memory;WARNING;HARD;3;WARNING - 18.9% (1388996 kB) free!
[1552373652] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 320 processes
[1552373660] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6576824 kB) free!
[1552373772] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 323 processes
[1552373780] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.3% (13947016 kB) free.
[1552373893] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 320 processes
[1552374129] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.87, 5.02, 5.00
[1552374708] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.6% (10156164 kB) free.
[1552374729] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.14, 5.19, 5.01
[1552374849] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.73, 5.39, 5.11
[1552374914] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3037920 kB) free!
[1552374969] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.87, 5.54, 5.19
[1552375034] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.4% (2995188 kB) free!
[1552375154] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.2% (2955376 kB) free!
[1552376180] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6958932 kB) free!
[1552376300] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.9% (14227724 kB) free.
[1552376508] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9228104 kB) free!
[1552376629] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.8% (6807144 kB) free!
[1552376711] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 18938 MB (9% inode=99%):
[1552376749] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4277128 kB) free!
[1552376769] SERVICE FLAPPING ALERT: sigqscluster2;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1552376769] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.54, 4.34, 4.77
[1552376893] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 319 processes
[1552377493] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1552377613] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 324 processes
[1552377734] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1552377734] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 320 processes
[1552378062] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1552378100] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7108276 kB) free!
[1552378220] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.4% (14478052 kB) free.
[1552378603] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (662836 kB) free!
[1552378662] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 59 users currently logged in
[1552378754] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (1547664 kB) free!
[1552379005] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 19.3% (1422708 kB) free!
[1552379031] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1552379111] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 32579 MB (15% inode=99%):
[1552379125] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 19.1% (1408324 kB) free!
[1552379200] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1552379245] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 18.9% (1393072 kB) free!
[1552380020] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (6986732 kB) free!
[1552380140] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.9% (14259096 kB) free.
[1552380369] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.67, 6.14, 5.40
[1552380489] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.83, 5.96, 5.42
[1552380611] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.98, 6.00, 5.50
[1552380734] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1552381244] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4913832 kB) free!
[1552381364] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 20.1% (4920632 kB) free.
[1552381934] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552381940] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6561308 kB) free!
[1552381964] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4913776 kB) free!
[1552382055] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552382060] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.1% (13840164 kB) free.
[1552382084] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4908092 kB) free!
[1552382149] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.3% (6042924 kB) free!
[1552382175] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552382204] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 20.0% (4904012 kB) free!
[1552382411] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 4.60, 4.51, 4.84
[1552382749] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (226896 kB) free!
[1552384004] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 24.7% (6059716 kB) free.
[1552384461] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6857940 kB) free!
[1552384581] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.6% (14111444 kB) free.
[1552384645] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (674972 kB) free!
[1552385149] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.3% (6554484 kB) free!
[1552385175] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1552385411] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.87, 5.31, 5.05
[1552385531] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.25, 5.29, 5.07
[1552385651] SERVICE ALERT: sigqscluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.89, 5.51, 5.17
[1552385749] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.6% (10148528 kB) free.
[1552386349] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (2294548 kB) free!
[1552386376] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1552386381] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6358556 kB) free!
[1552386469] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (223736 kB) free!
[1552386496] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552386501] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.5% (14024144 kB) free.
[1552386589] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.6% (10654792 kB) free.
[1552386616] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1552386851] SERVICE ALERT: sigqscluster2;CPU Load;OK;HARD;1;OK - load average: 3.96, 4.25, 4.65
[1552387190] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (239492 kB) free!
[1552387310] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.7% (10193756 kB) free.
[1552388301] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6425972 kB) free!
[1552388421] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.2% (13895208 kB) free.
[1552388653] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 24.4% (6010556 kB) free.
[1552388711] SERVICE FLAPPING ALERT: db1;HDD_Home status;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1552389016] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1552389710] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (2054208 kB) free!
[1552389830] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (237960 kB) free!
[1552389950] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.8% (10266816 kB) free.
[1552389992] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.55, 8.53, 5.29
[1552390112] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.36, 5.77, 4.67
[1552390217] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1552390222] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6357848 kB) free!
[1552390337] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552390342] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.9% (13737552 kB) free.
[1552390457] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552390550] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9539140 kB) free!
[1552390670] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 14.0% (6889500 kB) free!
[1552390791] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4229500 kB) free!
[1552391494] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;1;WARNING - 20.0% (1474488 kB) free!
[1552391614] SERVICE ALERT: sig2-vm1;Memory;WARNING;SOFT;2;WARNING - 19.9% (1463692 kB) free!
[1552391734] SERVICE ALERT: sig2-vm1;Memory;WARNING;HARD;3;WARNING - 19.8% (1453284 kB) free!
[1552392103] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 321 processes
[1552392142] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.0% (6420592 kB) free!
[1552392262] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.4% (13975184 kB) free.
[1552392857] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1552393303] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 322 processes
[1552393453] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4923928 kB) free!
[1552393457] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552393462] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6264128 kB) free!
[1552393573] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.8% (4890728 kB) free!
[1552393578] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552393582] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.7% (13635784 kB) free.
[1552393693] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.7% (4856772 kB) free!
[1552393698] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552394355] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 91.3% (14839728 kB) free.
[1552394712] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 18312 MB (8% inode=99%):
[1552394991] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.6% (8649404 kB) free!
[1552395383] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.7% (6760352 kB) free!
[1552395498] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1552395503] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.9% (14232088 kB) free.
[1552395592] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.3% (10515548 kB) free.
[1552396098] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1552396103] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.8% (5835908 kB) free!
[1552396218] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1552396223] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.5% (13547172 kB) free.
[1552396338] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1552396604] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4867976 kB) free!
[1552396724] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (4803508 kB) free!
[1552396792] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (235852 kB) free!
[1552396844] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (4740036 kB) free!
[1552396912] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.9% (10296668 kB) free.
[1552397112] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24798 MB (11% inode=99%):
[1552397423] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.2% (6001544 kB) free!
[1552397512] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (226772 kB) free!
[1552397543] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.4% (13496572 kB) free.
[1552397633] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.5% (11107984 kB) free.
[1552397651] SERVICE FLAPPING ALERT: sigqscluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1552398138] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1552398232] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (235788 kB) free!
[1552398312] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 17464 MB (8% inode=99%):
[1552398334] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (678556 kB) free!
[1552398352] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (232652 kB) free!
[1552398472] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.2% (10953268 kB) free.
[1552398738] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552398743] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.0% (5912236 kB) free!
[1552398858] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552398863] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.4% (13503564 kB) free.
[1552398978] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552399072] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (683868 kB) free!
[1552399192] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (213220 kB) free!
[1552399312] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.5% (11078596 kB) free.
[1552399512] SERVICE FLAPPING ALERT: db1;HDD_Home status;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1552399512] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 27918 MB (13% inode=99%):
[1552399913] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (1512760 kB) free!
[1552400033] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (213860 kB) free!
[1552400112] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 18830 MB (9% inode=99%):
[1552400153] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (230288 kB) free!
[1552400663] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6442220 kB) free!
[1552400712] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 24839 MB (11% inode=99%):
[1552400753] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.8% (5307372 kB) free!
[1552400783] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.7% (13626864 kB) free.
[1552401379] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1552401912] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 20035 MB (9% inode=99%):
[1552401953] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (249752 kB) free!
[1552401979] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552401983] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6793192 kB) free!
[1552402099] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552402103] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.3% (13955828 kB) free.
[1552402219] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552402703] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.2% (6026392 kB) free!
[1552402823] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.1% (13356616 kB) free.
[1552403112] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26876 MB (12% inode=99%):
[1552403154] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.7% (6731724 kB) free!
[1552403713] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 19004 MB (9% inode=99%):
[1552403754] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233176 kB) free!
[1552404019] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1552404313] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26751 MB (12% inode=99%):
[1552404354] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.9% (6865572 kB) free!
[1552404954] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.8% (10761184 kB) free.
[1552405219] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1552405340] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552405460] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552405513] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 19724 MB (9% inode=99%):
[1552405554] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (240272 kB) free!
[1552405674] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.1% (9919484 kB) free.
[1552405823] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6642440 kB) free!
[1552405943] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.8% (13685820 kB) free.
[1552407743] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7188572 kB) free!
[1552407860] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1552407863] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.5% (14035208 kB) free.
[1552408074] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (247504 kB) free!
[1552408194] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.2% (10931576 kB) free.
[1552408513] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 26970 MB (12% inode=99%):
[1552408795] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (3869560 kB) free!
[1552408915] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (1304072 kB) free!
[1552409035] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (236272 kB) free!
[1552409060] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552409063] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6304456 kB) free!
[1552409180] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552409183] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.0% (13320580 kB) free.
[1552409301] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552409986] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 94.04, 59.99, 26.91
[1552410032] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5543 MB (10% inode=97%):
[1552410044] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2259952 kB) free!
[1552410106] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 94.56, 71.44, 35.12
[1552410226] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 94.42, 79.05, 42.31
[1552410542] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 56.62, 16.24, 5.96
[1552410632] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 5790 MB (11% inode=97%):
[1552410644] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.3% (2520748 kB) free!
[1552410662] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 89.14, 42.01, 16.65
[1552410782] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 93.59, 59.31, 26.07
[1552411235] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 358 processes
[1552411244] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2434776 kB) free!
[1552411426] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.24, 2.79, 12.91
[1552411583] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7312384 kB) free!
[1552411701] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552411703] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.5% (14028188 kB) free.
[1552411982] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.20, 2.63, 12.67
[1552412301] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552412303] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9284868 kB) free!
[1552412421] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552412423] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8063100 kB) free!
[1552412541] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1552412543] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.3% (6567084 kB) free!
[1552412626] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.38, 1.53, 4.64
[1552413143] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9939412 kB) free.
[1552413182] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.60, 0.40, 3.69
[1552413741] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1552413913] SERVICE ALERT: db1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 21409 MB (10% inode=99%):
[1552414341] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1552414343] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6335532 kB) free!
[1552414435] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.6% (9169296 kB) free!
[1552414461] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1552414463] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.5% (14030528 kB) free.
[1552414513] SERVICE ALERT: db1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 25482 MB (12% inode=99%):
[1552415061] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1552415063] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9255144 kB) free!
[1552415181] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1552415183] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8241272 kB) free!
[1552415301] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1552415304] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.4% (7072764 kB) free!