[1555266684] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2436260 kB) free!
[1555266688] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.09, 0.26, 3.79
[1555266779] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44096 MB (19% inode=99%):
[1555266899] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56390 MB (25% inode=99%):
[1555268699] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44099 MB (19% inode=99%):
[1555268819] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56387 MB (25% inode=99%):
[1555269419] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45564 MB (20% inode=99%):
[1555269539] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45564 MB (20% inode=99%):
[1555269659] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43522 MB (19% inode=99%):
[1555270259] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48195 MB (21% inode=99%):
[1555271459] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46659 MB (20% inode=99%):
[1555271579] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44610 MB (19% inode=99%):
[1555271699] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56386 MB (25% inode=99%):
[1555272300] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44094 MB (19% inode=99%):
[1555272420] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44094 MB (19% inode=99%):
[1555272534] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2452920 kB) free!
[1555272540] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44094 MB (19% inode=99%):
[1555273140] SERVICE FLAPPING ALERT: db2;HDD_Home status;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1555273140] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48189 MB (21% inode=99%):
[1555273735] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.7% (2622976 kB) free!
[1555274340] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44076 MB (19% inode=99%):
[1555274460] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56365 MB (25% inode=99%):
[1555275463] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1555276064] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1555276184] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1555276304] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1555276304] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1555277919] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4890776 kB) free!
[1555277935] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2459948 kB) free!
[1555278039] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (4884848 kB) free!
[1555278159] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.9% (4879824 kB) free!
[1555278660] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44075 MB (19% inode=99%):
[1555278780] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56348 MB (25% inode=99%):
[1555279633] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 410.47, 167.27, 62.73
[1555279753] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 452.05, 263.62, 110.71
[1555279873] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 457.75, 328.15, 152.90
[1555279893] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.0% (4400264 kB) free!
[1555280013] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 28.1% (6896792 kB) free.
[1555280580] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44074 MB (19% inode=99%):
[1555280700] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44074 MB (19% inode=99%):
[1555280820] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56363 MB (25% inode=99%):
[1555282273] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.30, 0.56, 14.32
[1555283473] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.56, 0.67, 4.45
[1555285020] SERVICE FLAPPING ALERT: db2;HDD_Home status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1555286220] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43173 MB (19% inode=99%):
[1555286340] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56356 MB (25% inode=99%):
[1555287704] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1555288140] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43805 MB (19% inode=99%):
[1555288260] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56186 MB (25% inode=99%):
[1555288359] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2061736 kB) free!
[1555289875] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 473.67, 385.56, 198.65
[1555289995] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 340.51, 384.43, 222.04
[1555290061] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46451 MB (20% inode=99%):
[1555290115] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 47.52, 257.85, 195.35
[1555290159] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.9% (2915272 kB) free!
[1555290181] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44403 MB (19% inode=99%):
[1555290301] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55918 MB (24% inode=99%):
[1555290901] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43884 MB (19% inode=99%):
[1555291021] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43884 MB (19% inode=99%):
[1555291141] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43884 MB (19% inode=99%):
[1555291741] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47979 MB (21% inode=99%):
[1555292515] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.62, 1.76, 16.29
[1555292941] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44835 MB (19% inode=99%):
[1555293061] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44843 MB (19% inode=99%):
[1555293181] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44034 MB (19% inode=99%):
[1555293781] SERVICE FLAPPING ALERT: db2;HDD_Home status;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1555293781] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48006 MB (21% inode=99%):
[1555294315] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.86, 1.83, 3.88
[1555294359] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (984116 kB) free!
[1555294981] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44902 MB (19% inode=99%):
[1555295102] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44391 MB (19% inode=99%):
[1555295222] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56166 MB (24% inode=99%):
[1555297022] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44421 MB (19% inode=99%):
[1555297142] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56195 MB (25% inode=99%):
[1555297742] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46274 MB (20% inode=99%):
[1555297862] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46274 MB (20% inode=99%):
[1555297982] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44739 MB (19% inode=99%):
[1555298582] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47520 MB (21% inode=99%):
[1555299104] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1555299704] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1555299825] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1555299945] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 319 processes
[1555300982] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44636 MB (19% inode=99%):
[1555301102] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56185 MB (25% inode=99%):
[1555302535] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.8% (2899836 kB) free!
[1555302591] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 14 users currently logged in
[1555302688] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 31.96, 202.39, 139.86
[1555302808] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.49, 135.52, 122.93
[1555302903] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46327 MB (20% inode=99%):
[1555302928] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.74, 90.73, 108.05
[1555303023] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45399 MB (20% inode=99%):
[1555303143] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42325 MB (18% inode=99%):
[1555303413] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4807492 kB) free!
[1555303534] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (4781100 kB) free!
[1555303654] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (4757516 kB) free!
[1555303680] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555303743] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47860 MB (21% inode=99%):
[1555303800] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555303920] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555304335] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2443532 kB) free!
[1555304343] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43764 MB (19% inode=99%):
[1555304463] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55920 MB (24% inode=99%):
[1555304728] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.50, 0.65, 15.86
[1555305663] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44272 MB (19% inode=99%):
[1555305784] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55787 MB (24% inode=99%):
[1555305900] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46807 MB (20% inode=99%):
[1555305928] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.22, 0.31, 4.60
[1555306020] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46807 MB (20% inode=99%):
[1555306140] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46581 MB (20% inode=99%):
[1555306491] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 391.52, 144.33, 53.08
[1555306611] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 455.31, 250.59, 103.08
[1555306731] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 463.57, 321.46, 146.94
[1555306983] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43755 MB (19% inode=99%):
[1555307103] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55780 MB (24% inode=99%):
[1555308303] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46308 MB (20% inode=99%):
[1555308423] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46308 MB (20% inode=99%):
[1555308543] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44260 MB (19% inode=99%):
[1555309054] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2182608 kB) free!
[1555309131] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.88, 1.38, 17.19
[1555309143] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51676 MB (22% inode=99%):
[1555310343] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43482 MB (19% inode=99%):
[1555310463] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43482 MB (19% inode=99%):
[1555310474] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 430.96, 227.65, 92.15
[1555310579] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555310583] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43482 MB (19% inode=99%):
[1555310594] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 440.83, 298.73, 134.63
[1555310699] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555310714] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 324.57, 319.64, 163.20
[1555310819] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555310854] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.3% (2759612 kB) free!
[1555310903] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.7% (2877504 kB) free!
[1555310931] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.51, 1.46, 3.72
[1555311023] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.4% (2669096 kB) free!
[1555311143] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.2% (2468580 kB) free!
[1555311183] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51663 MB (22% inode=99%):
[1555311454] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1555311454] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2444356 kB) free!
[1555311743] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (1455784 kB) free!
[1555312383] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46031 MB (20% inode=99%):
[1555312504] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46031 MB (20% inode=99%):
[1555312624] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46031 MB (20% inode=99%):
[1555313114] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.17, 0.92, 13.03
[1555313224] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51636 MB (22% inode=99%):
[1555313544] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 85.8% (13953080 kB) free.
[1555314314] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.92, 0.76, 4.06
[1555314424] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43442 MB (19% inode=99%):
[1555314435] SERVICE ALERT: sig4-vm2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 281 processes
[1555314454] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (4376468 kB) free!
[1555314544] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43444 MB (19% inode=99%):
[1555314555] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 276 processes
[1555314563] SERVICE ALERT: sig4-vm2;Memory;WARNING;SOFT;1;WARNING - 10.6% (776616 kB) free!
[1555314636] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (147272 kB) free!
[1555314664] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54683 MB (24% inode=99%):
[1555314683] SERVICE ALERT: sig4-vm2;Memory;WARNING;SOFT;2;WARNING - 10.3% (756396 kB) free!
[1555314756] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (128928 kB) free!
[1555314803] SERVICE ALERT: sig4-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (735176 kB) free!
[1555314876] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (121884 kB) free!
[1555314996] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (116336 kB) free!
[1555315075] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (119124 kB) free!
[1555315116] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (115716 kB) free!
[1555315116] SERVICE ALERT: sig4-vm1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (118308 kB) free!
[1555315195] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.6% (119884 kB) free!
[1555315202] SERVICE ALERT: sig5-vm1;Memory;WARNING;SOFT;1;WARNING - 17.9% (1318000 kB) free!
[1555315236] SERVICE ALERT: sig2-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (125388 kB) free!
[1555315237] SERVICE ALERT: sig4-vm1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (129724 kB) free!
[1555315312] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.5% (1436780 kB) free!
[1555315315] SERVICE ALERT: sig3-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (126332 kB) free!
[1555315322] SERVICE ALERT: sig5-vm1;Memory;WARNING;SOFT;2;WARNING - 17.6% (1294928 kB) free!
[1555315357] SERVICE ALERT: sig4-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (120192 kB) free!
[1555315432] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 19.2% (1412592 kB) free!
[1555315442] SERVICE ALERT: sig5-vm1;Memory;WARNING;HARD;3;WARNING - 17.3% (1269780 kB) free!
[1555315552] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 18.8% (1385956 kB) free!
[1555315864] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43713 MB (19% inode=99%):
[1555315984] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39608 MB (17% inode=99%):
[1555316104] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55217 MB (24% inode=99%):
[1555316476] SERVICE ALERT: sig6-vm1;Memory;WARNING;SOFT;1;WARNING - 19.6% (1439856 kB) free!
[1555316596] SERVICE ALERT: sig6-vm1;Memory;WARNING;SOFT;2;WARNING - 19.4% (1426772 kB) free!
[1555316716] SERVICE ALERT: sig6-vm1;Memory;WARNING;HARD;3;WARNING - 19.1% (1405772 kB) free!
[1555317434] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 19.7% (1448496 kB) free!
[1555317554] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 19.5% (1432236 kB) free!
[1555317674] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 19.3% (1416256 kB) free!
[1555317904] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43183 MB (19% inode=99%):
[1555318025] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43183 MB (19% inode=99%):
[1555318131] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 347.46, 100.57, 35.28
[1555318145] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43185 MB (19% inode=99%):
[1555318252] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 53.69, 73.69, 33.68
[1555318372] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 8.72, 49.79, 29.77
[1555318745] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47293 MB (21% inode=99%):
[1555318972] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1555318972] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.03, 7.52, 16.06
[1555319152] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (716124 kB) free!
[1555319345] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45566 MB (20% inode=99%):
[1555319465] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45566 MB (20% inode=99%):
[1555319585] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43518 MB (19% inode=99%):
[1555319642] SERVICE ALERT: sig5-vm1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (702108 kB) free!
[1555320037] SERVICE ALERT: sig2-vm1;Memory;OK;HARD;1;OK - 96.0% (7063048 kB) free.
[1555320075] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 96.0% (7057364 kB) free.
[1555320116] SERVICE ALERT: sig3-vm1;Memory;OK;HARD;1;OK - 95.8% (7048856 kB) free.
[1555320157] SERVICE ALERT: sig4-vm1;Memory;OK;HARD;1;OK - 95.9% (7053372 kB) free.
[1555320185] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49548 MB (22% inode=99%):
[1555320203] SERVICE ALERT: sig4-vm2;Memory;OK;HARD;1;OK - 95.9% (7049596 kB) free.
[1555320242] SERVICE ALERT: sig5-vm1;Memory;OK;HARD;1;OK - 95.7% (7036880 kB) free.
[1555320276] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 95.5% (7026412 kB) free.
[1555320317] SERVICE ALERT: sig6-vm1;Memory;OK;HARD;1;OK - 95.6% (7028752 kB) free.
[1555320352] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 95.5% (7024592 kB) free.
[1555320772] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.27, 1.33, 3.39
[1555321385] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43396 MB (19% inode=99%):
[1555321505] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43396 MB (19% inode=99%):
[1555321625] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43396 MB (19% inode=99%):
[1555321654] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2315448 kB) free!
[1555322225] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53251 MB (23% inode=99%):
[1555322484] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 26.7% (6588632 kB) free.
[1555322545] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 11.2% (1817844 kB) free!
[1555322665] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 10.3% (1669600 kB) free!
[1555322785] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (1534876 kB) free!
[1555322825] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47105 MB (20% inode=99%):
[1555322945] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47105 MB (20% inode=99%):
[1555323050] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 340 processes
[1555323065] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46595 MB (20% inode=99%):
[1555323454] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.4% (2555136 kB) free!
[1555323945] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1555324054] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 26.6% (6518616 kB) free.
[1555324266] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51176 MB (22% inode=99%):
[1555324545] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555324665] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1555324785] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1555324866] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47080 MB (20% inode=99%):
[1555324986] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47080 MB (20% inode=99%):
[1555325106] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47080 MB (20% inode=99%):
[1555325185] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 84.1% (13665268 kB) free.
[1555325219] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555325339] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555325459] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555325759] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1555325940] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24217 MB (10% inode=99%):
[1555326306] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50750 MB (22% inode=99%):
[1555326359] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 326 processes
[1555326479] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 328 processes
[1555326599] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 328 processes
[1555326906] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47166 MB (20% inode=99%):
[1555327026] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47166 MB (20% inode=99%):
[1555327146] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43051 MB (19% inode=99%):
[1555327284] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4918660 kB) free!
[1555327404] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (4871816 kB) free!
[1555327524] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (4828820 kB) free!
[1555328124] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 24.2% (5962544 kB) free.
[1555328346] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51142 MB (22% inode=99%):
[1555328559] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.7% (2868664 kB) free!
[1555328946] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47046 MB (20% inode=99%):
[1555329050] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 339 processes
[1555329066] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42948 MB (19% inode=99%):
[1555329186] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42948 MB (19% inode=99%):
[1555329385] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 31.8% (5173540 kB) free.
[1555329540] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46315 MB (20% inode=99%):
[1555329786] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55224 MB (24% inode=99%):
[1555330386] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45457 MB (20% inode=99%):
[1555330506] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45457 MB (20% inode=99%):
[1555330627] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45457 MB (20% inode=99%):
[1555330821] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9860612 kB) free!
[1555330941] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 33.3% (16416836 kB) free.
[1555330972] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1555331124] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.3% (4518012 kB) free!
[1555331244] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (4467748 kB) free!
[1555331254] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (4311156 kB) free!
[1555331364] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1555331364] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.9% (4429144 kB) free!
[1555331375] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.3% (4234696 kB) free!
[1555331495] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.9% (4147060 kB) free!
[1555331827] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50072 MB (22% inode=99%):
[1555332385] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.6% (3019328 kB) free!
[1555332427] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45976 MB (20% inode=99%):
[1555332505] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.0% (2928288 kB) free!
[1555332547] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 37766 MB (16% inode=99%):
[1555332599] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 329 processes
[1555332625] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.5% (2851848 kB) free!
[1555332667] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 37766 MB (16% inode=99%):
[1555333267] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52524 MB (23% inode=99%):
[1555333867] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46380 MB (20% inode=99%):
[1555333895] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2074832 kB) free!
[1555333987] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46380 MB (20% inode=99%):
[1555334107] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44331 MB (19% inode=99%):
[1555334425] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (1483056 kB) free!
[1555334459] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555334520] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555334991] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1555335307] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48258 MB (21% inode=99%):
[1555335564] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2477760 kB) free!
[1555335907] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44162 MB (19% inode=99%):
[1555336028] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44162 MB (19% inode=99%):
[1555336148] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44162 MB (19% inode=99%):
[1555337948] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53268 MB (23% inode=99%):
[1555338548] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45652 MB (20% inode=99%):
[1555338668] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45652 MB (20% inode=99%):
[1555338741] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9707912 kB) free!
[1555338759] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2437368 kB) free!
[1555338788] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41554 MB (18% inode=99%):
[1555338861] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 34.1% (16817576 kB) free.
[1555339388] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54252 MB (24% inode=99%):
[1555339988] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46572 MB (20% inode=99%):
[1555340108] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44523 MB (19% inode=99%):
[1555340228] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44523 MB (19% inode=99%):
[1555340828] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50662 MB (22% inode=99%):
[1555341261] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9334152 kB) free!
[1555341381] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 33.8% (16648268 kB) free.
[1555341428] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45807 MB (20% inode=99%):
[1555341548] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44783 MB (19% inode=99%):
[1555341668] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44783 MB (19% inode=99%):
[1555341928] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.32, 6.17, 6.26
[1555341982] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9584816 kB) free!
[1555342029] SERVICE ALERT: sservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45570 MB (20% inode=99%):
[1555342048] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.55, 4.29, 5.55
[1555342102] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 33.8% (16635904 kB) free.
[1555342149] SERVICE ALERT: sservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45570 MB (20% inode=99%):
[1555342168] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.47, 3.01, 4.93
[1555342268] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50715 MB (22% inode=99%):
[1555342269] SERVICE ALERT: sservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45999 MB (20% inode=99%):
[1555342702] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9634752 kB) free!
[1555342822] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 33.5% (16503972 kB) free.
[1555342869] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45083 MB (20% inode=99%):
[1555342989] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43035 MB (19% inode=99%):
[1555343109] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41849 MB (18% inode=99%):
[1555343364] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1555343495] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1555345509] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54667 MB (24% inode=99%):
[1555346109] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46602 MB (20% inode=99%):
[1555346229] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46219 MB (20% inode=99%):
[1555346349] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44683 MB (19% inode=99%):
[1555347585] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1555348149] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47561 MB (21% inode=99%):
[1555348185] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555348305] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 319 processes
[1555348425] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1555349350] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43456 MB (19% inode=99%):
[1555349470] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43456 MB (19% inode=99%):
[1555349590] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43463 MB (19% inode=99%):
[1555350790] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54853 MB (24% inode=99%):
[1555351390] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46660 MB (20% inode=99%):
[1555351510] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46661 MB (20% inode=99%):
[1555351630] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45125 MB (20% inode=99%):