[1546286836] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.27, 46.35, 39.27
[1546286956] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.66, 31.22, 34.59
[1546287076] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.47, 21.03, 30.44
[1546287676] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.58, 3.33, 16.23
[1546288876] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.46, 0.61, 4.89
[1546294022] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 449 processes
[1546294082] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 448 processes
[1546294276] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 627.67, 157.53, 53.22
[1546294396] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 101.90, 121.61, 53.56
[1546294516] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14.19, 81.56, 47.13
[1546295716] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.40, 1.95, 13.35
[1546296916] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.87, 0.79, 4.20
[1546302316] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 9.80, 74.98, 45.84
[1546302436] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.78, 50.36, 40.35
[1546302556] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.90, 33.95, 35.56
[1546303157] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.80, 5.22, 18.98
[1546304957] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.85, 0.73, 3.34
[1546309757] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 41.02, 101.42, 50.57
[1546309877] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 5.90, 68.01, 44.50
[1546309997] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.39, 45.71, 39.17
[1546311197] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.78, 1.56, 11.27
[1546312397] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.76, 0.73, 3.60
[1546314032] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546314089] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9820528 kB) free!
[1546314152] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546314209] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.1% (8420252 kB) free!
[1546314272] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546314329] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10260372 kB) free.
[1546314797] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.76, 10.13, 11.49
[1546314872] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546314917] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.62, 6.97, 10.16
[1546315037] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.46, 6.54, 9.70
[1546316237] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.53, 0.72, 3.12
[1546320872] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1546320992] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1546321112] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1546321712] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1546321712] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1546322837] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 0.62, 28.41, 33.58
[1546322957] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.87, 21.29, 30.36
[1546323077] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.52, 14.63, 26.83
[1546323678] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.41, 2.40, 14.31
[1546324878] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.61, 0.74, 4.42
[1546331478] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.54, 49.37, 43.43
[1546331598] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.92, 33.37, 38.29
[1546331718] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.51, 22.51, 33.72
[1546332318] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.47, 3.41, 17.89
[1546333112] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1546334118] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.64, 0.72, 3.17
[1546334912] SERVICE ALERT: qservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1546337129] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8990832 kB) free!
[1546337249] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.5% (7649176 kB) free!
[1546337369] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10179388 kB) free.
[1546344318] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1546351518] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.45, 40.69, 30.97
[1546351638] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.52, 27.34, 27.26
[1546351758] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.35, 18.41, 24.00
[1546352358] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.40, 2.97, 12.87
[1546353558] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.65, 0.73, 4.04
[1546358958] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 578.61, 170.62, 59.18
[1546359078] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 78.78, 114.43, 52.10
[1546359198] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1546359198] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 11.33, 76.82, 45.87
[1546360398] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.57, 1.92, 13.00
[1546361598] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.55, 0.60, 3.98