[1589740240] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16285304 kB) free!
[1589740840] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (447272 kB) free!
[1589742040] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33372936 kB) free.
[1589742640] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9060220 kB) free!
[1589742760] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (5091884 kB) free!
[1589742881] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1141028 kB) free!
[1589743190] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (1248660 kB) free!
[1589743310] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (1058976 kB) free!
[1589743430] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (333612 kB) free!
[1589743481] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17384952 kB) free!
[1589743884] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (192744 kB) free!
[1589744081] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402336 kB) free!
[1589744591] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 32857 MB (14.62% inode=100%):
[1589744651] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 32857 MB (14.62% inode=100%):
[1589744681] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13452204 kB) free!
[1589744711] SERVICE ALERT: webserver2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 82407 MB (36.66% inode=100%):
[1589744784] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 79.3% (19424652 kB) free.
[1589745281] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31944212 kB) free.
[1589745748] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (238616 kB) free!
[1589745868] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (216644 kB) free!
[1589745881] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8652796 kB) free!
[1589745944] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (403812 kB) free!
[1589745988] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237464 kB) free!
[1589746001] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4818400 kB) free!
[1589746064] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (381032 kB) free!
[1589746121] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1193344 kB) free!
[1589746184] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (378944 kB) free!
[1589746721] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15516548 kB) free!
[1589747321] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32995256 kB) free.
[1589747788] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 12.9% (6369052 kB) free!
[1589747921] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9088332 kB) free!
[1589748041] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5380304 kB) free!
[1589748161] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (1775752 kB) free!
[1589748761] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16730524 kB) free!
[1589749361] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34360892 kB) free.
[1589749430] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1589749701] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 86.1% (42411376 kB) free.
[1589749961] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10789624 kB) free!
[1589750082] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (7125608 kB) free!
[1589750144] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.0% (42860996 kB) free.
[1589750202] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3460148 kB) free!
[1589750630] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (2578952 kB) free!
[1589750802] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16748032 kB) free!
[1589751062] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2265) < 203.212.70.170 [18/May/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1589751183] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(77) < 203.212.70.170 [18/May/2020:03:02:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1589751265] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.3% (3024932 kB) free!
[1589751303] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589751344] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(457) < 2020/05/18 03:01:32 [error] 22135#0: *89025986 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1589751344] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(353) < 103.16.101.72 [18/May/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1589751388] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 89.8% (44262188 kB) free.
[1589751402] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30643708 kB) free.
[1589751464] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589751465] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589752002] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (3062356 kB) free!
[1589752122] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (391484 kB) free!
[1589752242] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403252 kB) free!
[1589752842] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13476692 kB) free!
[1589753061] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 30.4% (7439992 kB) free.
[1589753442] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29219640 kB) free.
[1589753661] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190108 kB) free!
[1589753781] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193960 kB) free!
[1589753901] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188840 kB) free!
[1589754042] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5270936 kB) free!
[1589754162] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (1439156 kB) free!
[1589754282] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35865572 kB) free.
[1589754866] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1937184 kB) free!
[1589754882] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11436744 kB) free!
[1589755002] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7322900 kB) free!
[1589755122] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3377272 kB) free!
[1589755723] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17585840 kB) free!
[1589756323] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34586480 kB) free.
[1589756923] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12160168 kB) free!
[1589757043] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8587576 kB) free!
[1589757163] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (5155780 kB) free!
[1589757584] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.2% (13475212 kB) free!
[1589757763] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (20583932 kB) free!
[1589757830] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3297864 kB) free.
[1589758101] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 49.1% (12043280 kB) free.
[1589758363] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31881188 kB) free.
[1589758701] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198136 kB) free!
[1589758821] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177812 kB) free!
[1589758942] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1589758942] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178456 kB) free!
[1589758963] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12683364 kB) free!
[1589759083] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38250816 kB) free.
[1589759441] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 65.0% (15942996 kB) free.
[1589759666] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 48.2% (11802412 kB) free.
[1589759683] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19329712 kB) free!
[1589759803] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.7% (15391992 kB) free!
[1589759878] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 47.9% (11817960 kB) free.
[1589759923] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (11424220 kB) free!
[1589760523] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.9% (21001756 kB) free!
[1589761123] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34560696 kB) free.
[1589761723] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13585880 kB) free!
[1589761784] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (13007768 kB) free!
[1589761843] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (40004336 kB) free.
[1589762444] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18039496 kB) free!
[1589762564] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.8% (14177292 kB) free!
[1589762684] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (10566040 kB) free!
[1589763284] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.4% (23000988 kB) free!
[1589763884] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32557364 kB) free.
[1589764078] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183656 kB) free!
[1589764198] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (209248 kB) free!
[1589764318] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172944 kB) free!
[1589764484] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13448688 kB) free!
[1589764522] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1589764604] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39072760 kB) free.
[1589764642] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 341 processes
[1589764762] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 342 processes
[1589764784] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 70.9% (93455556 kB) free.
[1589765204] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18631684 kB) free!
[1589765324] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.3% (14832684 kB) free!
[1589765444] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39235612 kB) free.
[1589766044] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19761068 kB) free!
[1589766164] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.0% (15867196 kB) free!
[1589766284] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (11867292 kB) free!
[1589766884] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20450264 kB) free!
[1589767484] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29148584 kB) free.
[1589768084] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22910092 kB) free!
[1589768205] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.0% (18404240 kB) free!
[1589768325] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (13951004 kB) free!
[1589768866] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1589768986] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1589769464] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/18 07:58:15 [error] 22131#0: *89052922 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1589769525] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (2488188 kB) free!
[1589769584] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589770125] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20224800 kB) free!
[1589770342] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589770725] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29389904 kB) free.
[1589771325] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8471304 kB) free!
[1589771445] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4242360 kB) free!
[1589771565] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402636 kB) free!
[1589772141] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.55, 6.68, 5.29
[1589772261] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.04, 6.03, 5.22
[1589772381] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.65, 5.23, 5.03
[1589772765] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21191176 kB) free!
[1589772830] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3039120 kB) free!
[1589772866] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4595368 kB) free!
[1589772950] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (3001092 kB) free!
[1589772986] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (4362680 kB) free!
[1589773070] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (3083512 kB) free!
[1589773106] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.9% (4150928 kB) free!
[1589773184] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(37) < 2020/05/18 09:09:31 [error] 22131#0: *89060496 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1589773304] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/05/18 09:10:36 [error] 22129#0: *89060670 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:236/17, bytes from/to upstream:17/236
[1589773365] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385172 kB) free!
[1589773424] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589773965] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13367084 kB) free!
[1589774024] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/18 09:15:40 [error] 22135#0: *89061101 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1589774144] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589774306] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2130052 kB) free!
[1589774565] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27014792 kB) free.
[1589774962] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1589775166] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (1931316 kB) free!
[1589775286] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (398668 kB) free!
[1589775406] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16549976 kB) free!
[1589776006] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (6649816 kB) free!
[1589776606] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.0% (25038688 kB) free!
[1589776842] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199344 kB) free!
[1589776962] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196360 kB) free!
[1589776978] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1589777082] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172068 kB) free!
[1589777098] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1589777206] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401196 kB) free!
[1589777218] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1589777806] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19828800 kB) free!
[1589778406] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408464 kB) free!
[1589779073] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.8% (3105908 kB) free!
[1589779193] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 10.9% (2453452 kB) free!
[1589779313] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1771036 kB) free!
[1589779670] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (1459428 kB) free!
[1589779818] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 766 processes
[1589779938] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 762 processes
[1589780058] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 762 processes
[1589780206] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16596836 kB) free!
[1589780633] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 42.1% (10330844 kB) free.
[1589780703] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(443) < 103.255.146.164 [18/May/2020:11:08:34 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589780806] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35292248 kB) free.
[1589780823] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589780870] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.3% (1842248 kB) free!
[1589781406] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (4758160 kB) free!
[1589781423] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 203.212.70.170 [18/May/2020:11:18:36 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1589781470] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1589781470] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1616076 kB) free!
[1589781526] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (391680 kB) free!
[1589781543] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589781646] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393596 kB) free!
[1589782143] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(259) < 103.255.146.164 [18/May/2020:11:31:21 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589782186] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1589782264] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589782306] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1589783362] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1589783385] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 17.3% (6044340 kB) free!
[1589783463] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(282) < 103.255.146.164 [18/May/2020:11:58:54 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589783506] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.8% (31800900 kB) free.
[1589783513] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.6% (15860408 kB) free.
[1589783581] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 16.4% (4046000 kB) free!
[1589783583] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589783633] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 14.0% (3433772 kB) free!
[1589783702] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 15.0% (3710660 kB) free!
[1589783753] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 12.7% (3105296 kB) free!
[1589783822] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 13.3% (3280480 kB) free!
[1589783870] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3263140 kB) free.
[1589783873] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 11.2% (2753972 kB) free!
[1589784046] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36817840 kB) free.
[1589784282] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 24.9% (6100128 kB) free.
[1589784422] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (1224472 kB) free!
[1589784470] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3172228 kB) free!
[1589784473] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (867356 kB) free!
[1589784590] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (3112068 kB) free!
[1589784646] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6671212 kB) free!
[1589784710] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (3049948 kB) free!
[1589784766] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1564064 kB) free!
[1589784886] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (411300 kB) free!
[1589785022] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.9% (22927392 kB) free.
[1589785384] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(279) < 103.255.146.164 [18/May/2020:12:31:22 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589785486] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19216192 kB) free!
[1589785504] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589786082] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.7% (4338724 kB) free!
[1589786086] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382776 kB) free!
[1589786144] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9455140 kB) free!
[1589786202] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.9% (4145648 kB) free!
[1589786264] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9191852 kB) free!
[1589786301] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8598192 kB) free!
[1589786322] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (3886300 kB) free!
[1589786340] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 357 processes
[1589786362] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1589786384] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.1% (8923928 kB) free!
[1589786421] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.9% (8320008 kB) free!
[1589786460] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 359 processes
[1589786541] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 16.3% (8053904 kB) free!
[1589786580] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 357 processes
[1589786686] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17596680 kB) free!
[1589786704] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(234) < 103.255.146.164 [18/May/2020:12:49:30 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589786824] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589786984] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4373224 kB) free!
[1589787110] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (735236 kB) free!
[1589787141] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (767164 kB) free!
[1589787286] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (442020 kB) free!
[1589787522] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1589787522] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1542948 kB) free!
[1589787584] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 20.9% (10324112 kB) free.
[1589787706] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1589787710] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.1% (2619472 kB) free!
[1589787741] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 25.4% (12499732 kB) free.
[1589787826] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1589787886] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15171504 kB) free!
[1589788025] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(68) < 103.255.146.164 [18/May/2020:13:14:25 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589788145] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589788184] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9302516 kB) free!
[1589788304] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9099588 kB) free!
[1589788424] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1589788424] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.1% (8924612 kB) free!
[1589788486] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383700 kB) free!
[1589789345] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(268) < 103.255.146.164 [18/May/2020:13:32:10 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589789465] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589789510] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (260380 kB) free!
[1589789541] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9784600 kB) free!
[1589789661] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (9624800 kB) free!
[1589789781] SERVICE FLAPPING ALERT: qservercluster2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1589789781] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (9463296 kB) free!
[1589790381] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.61, 2.94, 4.72
[1589790710] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.1% (1804468 kB) free!
[1589790982] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.12, 6.05, 5.60
[1589791079] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.9% (4867008 kB) free!
[1589791102] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.17, 6.43, 5.77
[1589791199] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.8% (4847724 kB) free!
[1589791222] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.78, 5.40, 5.48
[1589791319] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.7% (4825820 kB) free!
[1589791424] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4673752 kB) free!
[1589791865] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(47) < 103.255.146.164 [18/May/2020:14:12:02 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589791913] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.6% (3942664 kB) free!
[1589791985] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589792033] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.5% (3699044 kB) free!
[1589792153] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.3% (3426300 kB) free!
[1589792362] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1589792510] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (774672 kB) free!
[1589792753] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2183668 kB) free!
[1589793110] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.2% (2143700 kB) free!
[1589793185] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(261) < 103.255.146.164 [18/May/2020:14:41:42 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589793305] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589793381] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4889592 kB) free!
[1589793710] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (287580 kB) free!
[1589793906] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(286) < 103.255.146.164 [18/May/2020:14:51:59 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589794026] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589794311] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.5% (2197568 kB) free!
[1589794380] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1589794553] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1589794553] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.0% (15949556 kB) free.
[1589794626] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(105) < 185.255.8.58 [18/May/2020:15:00:10 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1589794746] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589794911] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (558652 kB) free!
[1589795346] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(283) < 103.255.146.164 [18/May/2020:15:13:09 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589795362] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1589795466] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589795482] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1589795511] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.2% (2465208 kB) free!
[1589795602] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 341 processes
[1589795744] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/18 15:23:31 [error] 22135#0: *89116055 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1589795864] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589796067] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(26) < 103.255.146.164 [18/May/2020:15:25:24 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589796187] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(20) < 103.255.146.164 [18/May/2020:15:31:17 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589796202] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1589796202] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1589796307] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589797064] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(5) < 2020/05/18 15:38:33 [error] 22131#0: *89072873 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:5268043/1568668, bytes from/to upstream:1568668/5268043
[1589797185] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589797305] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(9) < webmachine error: path="/api/whoami"
[1589797311] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (680188 kB) free!
[1589797374] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.1% (1640764 kB) free!
[1589797426] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589797426] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1589797507] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(779) < 172.105.59.173 [18/May/2020:15:54:59 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1589797546] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1589797628] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(31) < 49.50.86.152 [18/May/2020:15:56:02 +0530] TCP 502 0 0 60.000 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1589797748] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589797911] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (2806252 kB) free!
[1589798025] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(14) < webmachine error: path="/api/whoami"
[1589798086] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.0% (23789816 kB) free!
[1589798145] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;(15) < webmachine error: path="/api/whoami"
[1589798266] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589798322] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1589798458] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < webmachine error: path="/api/whoami"
[1589798579] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589798602] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 370 processes
[1589798686] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (425140 kB) free!
[1589799022] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 32.34, 27.49, 18.42
[1589799203] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 344 processes
[1589799287] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.9% (26260388 kB) free!
[1589799323] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1589799443] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1589799548] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 172.105.59.173 [18/May/2020:16:24:54 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1589799622] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1589799622] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 11.82, 15.16, 16.38
[1589799668] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589799887] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399508 kB) free!
[1589800269] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(568) < 103.255.146.164 [18/May/2020:16:33:33 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589800311] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (663464 kB) free!
[1589800378] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(14) < webmachine error: path="/api/whoami"
[1589800389] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589800424] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1589800487] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.4% (25636936 kB) free!
[1589800498] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589800643] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1589800665] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(12) < webmachine error: path="/api/whoami"
[1589800785] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;(12) < webmachine error: path="/api/whoami"
[1589800905] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589801087] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (459164 kB) free!
[1589801098] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < webmachine error: path="/api/whoami"
[1589801153] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4284656 kB) free!
[1589801219] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589801243] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1589801273] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4077848 kB) free!
[1589801363] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1589801393] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.2% (3861904 kB) free!
[1589801506] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(4) < webmachine error: path="/api/whoami"
[1589801574] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166436 kB) free!
[1589801589] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(26) < 203.212.70.170 [18/May/2020:17:00:32 +0530] TCP 502 0 0 180.003 "10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.001, 60.001, 60.001, 0.000"
[1589801626] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589801687] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23051380 kB) free!
[1589801709] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589801781] SERVICE FLAPPING ALERT: qservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1589801818] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(7) < webmachine error: path="/api/whoami"
[1589801938] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589802226] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < webmachine error: path="/api/whoami"
[1589802287] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (404516 kB) free!
[1589802346] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;2;(5) < webmachine error: path="/api/whoami"
[1589802466] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589802538] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(7) < webmachine error: path="/api/whoami"
[1589802563] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 349 processes
[1589802593] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (1036940 kB) free!
[1589802658] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589802684] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1589802711] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.5% (2029088 kB) free!
[1589802887] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21214236 kB) free!
[1589803066] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(2) < webmachine error: path="/api/whoami"
[1589803186] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589803259] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(3) < webmachine error: path="/api/whoami"
[1589803311] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.2% (4739504 kB) free.
[1589803379] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589803487] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (379788 kB) free!
[1589803786] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < webmachine error: path="/api/whoami"
[1589803907] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589803912] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (2887508 kB) free!
[1589803979] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < webmachine error: path="/api/whoami"
[1589804032] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.8% (5006448 kB) free.
[1589804087] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.3% (14897312 kB) free!
[1589804099] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589804632] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.7% (2558432 kB) free!
[1589804687] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405648 kB) free!
[1589804752] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.9% (4855972 kB) free.
[1589805083] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1589805203] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 356 processes
[1589805287] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14777112 kB) free!
[1589805309] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(536) < 103.255.146.164 [18/May/2020:18:01:36 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1589805323] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 350 processes
[1589805352] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.3% (2643044 kB) free!
[1589805429] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589805472] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3697660 kB) free.
[1589805593] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.2% (15979332 kB) free.
[1589805887] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394256 kB) free!
[1589806487] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13536208 kB) free!
[1589806672] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (1056900 kB) free!
[1589806792] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.8% (2563784 kB) free!
[1589806907] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(4) < webmachine error: path="/api/whoami"
[1589806912] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3687056 kB) free.
[1589807028] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589807087] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382844 kB) free!
[1589807099] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(10) < webmachine error: path="/api/whoami"
[1589807123] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 349 processes
[1589807219] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589807512] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (225092 kB) free!
[1589807632] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (1598916 kB) free!
[1589807752] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (3035964 kB) free!
[1589808323] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1589808353] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.3% (3786352 kB) free.
[1589808923] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 345 processes
[1589808953] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (788136 kB) free!
[1589809073] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.8% (2237980 kB) free!
[1589809193] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3682604 kB) free.
[1589809546] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 4 warn out of 337 processes
[1589809619] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(7) < webmachine error: path="/api/whoami"
[1589809666] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 332 processes
[1589809739] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589809786] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1589810393] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (1467576 kB) free!
[1589810513] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3542236 kB) free.
[1589810580] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 362 processes
[1589810700] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 357 processes
[1589810820] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 352 processes
[1589811022] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589812313] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (390696 kB) free!
[1589812433] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.7% (1904548 kB) free!
[1589812553] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3404120 kB) free.
[1589813753] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (499476 kB) free!
[1589813820] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1589813873] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.6% (1716584 kB) free!
[1589813994] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 20.0% (3257712 kB) free!
[1589814593] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.5% (3931048 kB) free!
[1589814594] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3546652 kB) free.
[1589814713] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.7% (3740372 kB) free!
[1589814833] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.9% (3575148 kB) free!
[1589815222] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.83, 4.07, 4.60
[1589815487] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (30937812 kB) free.
[1589815794] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (287176 kB) free!
[1589815914] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (1363316 kB) free!
[1589816034] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.1% (15957016 kB) free.
[1589816034] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.4% (2670112 kB) free!
[1589816087] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (405536 kB) free!
[1589816207] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (429304 kB) free!
[1589816327] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (415284 kB) free!
[1589816634] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3292888 kB) free.
[1589817834] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3212268 kB) free!
[1589817954] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (3194120 kB) free!
[1589818074] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (369484 kB) free!
[1589818727] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24635896 kB) free!
[1589819274] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (3118208 kB) free!
[1589819327] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381596 kB) free!
[1589819927] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26435316 kB) free.
[1589820323] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589820527] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (408188 kB) free!
[1589820647] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (408968 kB) free!
[1589820767] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (427240 kB) free!
[1589821674] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.8% (4523744 kB) free.
[1589821967] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34382608 kB) free.
[1589822386] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1589822506] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1589822567] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8082252 kB) free!
[1589822687] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (3695512 kB) free!
[1589822807] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383416 kB) free!
[1589822874] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.1% (2783704 kB) free!
[1589822994] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3343484 kB) free.
[1589823407] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (18402908 kB) free!
[1589823594] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.1% (1797548 kB) free!
[1589823714] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.7% (2226480 kB) free!
[1589823834] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.4% (2658500 kB) free!
[1589823923] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 337 processes
[1589824007] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393076 kB) free!
[1589824434] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.0% (4233160 kB) free.
[1589825034] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.0% (2921368 kB) free!
[1589825154] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3261412 kB) free.
[1589825207] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34700924 kB) free.
[1589825807] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8064716 kB) free!
[1589825927] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (3932748 kB) free!
[1589826047] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389260 kB) free!
[1589826354] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.7% (2061148 kB) free!
[1589826474] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.7% (2232384 kB) free!
[1589826595] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.5% (2522428 kB) free!