[1560364211] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.3% (9525484 kB) free!
[1560364270] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.0% (6837240 kB) free.
[1560364331] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.9% (6826088 kB) free!
[1560364697] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 88.9% (14448684 kB) free.
[1560364757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1560364848] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.25, 2.08, 14.14
[1560364877] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1560364955] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.1% (2766512 kB) free!
[1560365126] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.41, 0.55, 3.53
[1560365133] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.75, 3.09, 4.19
[1560365478] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1560365531] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10097344 kB) free.
[1560365598] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1560365718] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560366132] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (219032 kB) free!
[1560366252] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10168348 kB) free.
[1560366756] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (982268 kB) free!
[1560366852] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (575972 kB) free!
[1560366918] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1560366972] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9913692 kB) free.
[1560367356] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.0% (1947616 kB) free!
[1560367518] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560367572] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (222124 kB) free!
[1560367639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1560367692] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10430272 kB) free.
[1560367759] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1560367848] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.75, 1.89, 3.63
[1560368292] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (840972 kB) free!
[1560368412] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10100760 kB) free.
[1560368556] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (903000 kB) free!
[1560369013] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (310332 kB) free!
[1560369133] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.3% (9532332 kB) free!
[1560369156] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.8% (2570048 kB) free!
[1560369253] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.9% (7327540 kB) free!
[1560369756] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1131436 kB) free!
[1560369910] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 380 processes
[1560370356] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.3% (2809512 kB) free!
[1560370453] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10140412 kB) free.
[1560370511] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 388 processes
[1560371053] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.5% (5180740 kB) free!
[1560371173] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (3371416 kB) free!
[1560371293] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.2% (1062500 kB) free!
[1560371360] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1560371556] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156364 kB) free!
[1560371711] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 404 processes
[1560371960] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1560372080] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1560372311] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 390 processes
[1560372493] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.4% (7596488 kB) free!
[1560372680] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1560372756] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.6% (4160860 kB) free.
[1560372800] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1560372920] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560373511] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 385 processes
[1560373693] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10205888 kB) free.
[1560374111] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 381 processes
[1560374293] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.0% (5402424 kB) free!
[1560374413] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (3735312 kB) free!
[1560374533] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (1370292 kB) free!
[1560374556] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.2% (3108376 kB) free!
[1560374676] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.6% (4149236 kB) free.
[1560374720] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1560375133] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.0% (6427504 kB) free!
[1560375320] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1560375440] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1560375560] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1560376161] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1560376281] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1560376334] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10198912 kB) free.
[1560376881] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1560376934] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4215048 kB) free!
[1560377001] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560377054] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (1872692 kB) free!
[1560377121] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560377174] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10140272 kB) free.
[1560377676] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.2% (2624852 kB) free!
[1560377711] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 384 processes
[1560377774] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (2078196 kB) free!
[1560377796] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.5% (3976492 kB) free.
[1560377894] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10133692 kB) free.
[1560378396] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.1% (3096264 kB) free!
[1560378494] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (4429936 kB) free!
[1560378516] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.7% (3517308 kB) free.
[1560378614] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (2342676 kB) free!
[1560378735] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (219800 kB) free!
[1560378911] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 404 processes
[1560379521] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1560379935] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.7% (6256420 kB) free!
[1560380121] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560380241] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1560380361] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1560380535] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10136824 kB) free.
[1560380711] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 373 processes
[1560380961] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1560381081] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560381135] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (1606552 kB) free!
[1560381202] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560381255] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10117848 kB) free.
[1560381855] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.6% (5245308 kB) free!
[1560381879] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 337 processes
[1560381976] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (3495432 kB) free!
[1560382096] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (1170552 kB) free!
[1560382402] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1560382512] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 406 processes
[1560382566] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 490.40, 358.19, 168.69
[1560382686] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 85.41, 256.75, 155.32
[1560382696] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.9% (6353920 kB) free!
[1560382806] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 11.83, 171.92, 136.53
[1560383002] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1560383079] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 338 processes
[1560383122] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1560383316] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.3% (2000716 kB) free!
[1560383437] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.8% (4027960 kB) free.
[1560383679] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 339 processes
[1560383722] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560383842] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1560384036] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1560384443] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1560384496] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (225220 kB) free!
[1560384563] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1560384607] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.23, 0.71, 19.87
[1560384648] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 151.77, 131.32, 73.12
[1560384768] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 22.71, 88.70, 64.56
[1560384879] SERVICE FLAPPING ALERT: redis-1;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1560384879] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 338 processes
[1560384889] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 5.03, 60.14, 57.02
[1560385163] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560385207] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 79.31, 55.98, 36.15
[1560385283] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560385403] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560385479] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 336 processes
[1560385646] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (1803440 kB) free!
[1560385836] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.9% (2424996 kB) free!
[1560385956] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (426956 kB) free!
[1560386076] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (170168 kB) free!
[1560386079] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 1 warn out of 345 processes
[1560386089] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.36, 3.29, 17.35
[1560386112] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 370 processes
[1560386246] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (158560 kB) free!
[1560386296] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.4% (5139308 kB) free!
[1560386407] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1560386407] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.09, 1.79, 11.31
[1560386476] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191236 kB) free!
[1560386596] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168232 kB) free!
[1560386604] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1560386656] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187076 kB) free!
[1560386712] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 385 processes
[1560386716] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195024 kB) free!
[1560386777] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173680 kB) free!
[1560386778] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182464 kB) free!
[1560386897] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187264 kB) free!
[1560386898] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178520 kB) free!
[1560387018] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186852 kB) free!
[1560387204] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1560387276] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.2% (1823420 kB) free!
[1560387324] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1560387607] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.22, 0.32, 3.43
[1560387876] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 54.1% (8779468 kB) free.
[1560387889] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.93, 1.96, 4.41
[1560387913] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 399 processes
[1560387924] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1560388044] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1560388096] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10145108 kB) free.
[1560388479] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 336 processes
[1560388513] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 395 processes
[1560388644] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560388696] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (214476 kB) free!
[1560388765] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1560388816] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.2% (9471644 kB) free!
[1560388936] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.0% (7372740 kB) free!
[1560389364] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560389484] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1560389679] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 338 processes
[1560390084] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560390137] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (217804 kB) free!
[1560390204] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1560390279] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 337 processes
[1560390804] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1560390879] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 338 processes
[1560390913] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 398 processes
[1560390924] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560391045] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1560391480] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 337 processes
[1560391513] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 386 processes
[1560391937] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.8% (5310124 kB) free!
[1560392680] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 338 processes
[1560393737] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10089380 kB) free.
[1560393880] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 337 processes
[1560393913] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 379 processes
[1560394338] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (815520 kB) free!
[1560394458] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9962984 kB) free.
[1560394480] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 338 processes
[1560394498] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.8% (3211176 kB) free!
[1560394513] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 395 processes
[1560394618] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.6% (3189540 kB) free!
[1560394738] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.5% (3171892 kB) free!
[1560395058] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (432064 kB) free!
[1560395080] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 338 processes
[1560395178] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10077620 kB) free.
[1560395680] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 339 processes
[1560395778] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6721304 kB) free!
[1560395898] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 18.7% (9216592 kB) free!
[1560396018] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.4% (7078816 kB) free!
[1560396313] SERVICE ALERT: redis-2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 397 processes
[1560396880] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 339 processes
[1560396913] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 416 processes
[1560397218] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11078708 kB) free.
[1560397481] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 338 processes
[1560397738] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1614684 kB) free!
[1560397818] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (217472 kB) free!
[1560397938] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11303464 kB) free.
[1560398081] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 336 processes
[1560398407] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1560398538] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (238372 kB) free!
[1560398658] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.2% (10463796 kB) free.
[1560398681] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 338 processes
[1560398727] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 104.68, 219.00, 127.35
[1560398733] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.77, 9.83, 7.68
[1560398847] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 14.79, 146.81, 112.02
[1560398853] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.88, 9.69, 7.90
[1560398967] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.45, 98.42, 98.50
[1560398974] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 10.10, 10.08, 8.27
[1560399018] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (2549492 kB) free!
[1560399289] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 349.93, 180.71, 73.56
[1560399409] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 359.37, 240.33, 108.40
[1560399445] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1560399529] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 361.26, 280.39, 139.07
[1560399618] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1939004 kB) free!
[1560399670] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.8% (354136 kB) free!
[1560399790] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (176832 kB) free!
[1560399859] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (679368 kB) free!
[1560399910] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (118768 kB) free!
[1560399979] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 20.0% (9851196 kB) free!
[1560400074] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 8 users currently logged in
[1560400099] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.6% (8190036 kB) free!
[1560400194] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 9 users currently logged in
[1560400314] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1560400510] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.0% (6840684 kB) free.
[1560400699] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10065480 kB) free.
[1560400767] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.65, 1.18, 15.06
[1560401299] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (221900 kB) free!
[1560401419] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (10022856 kB) free.
[1560401929] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.01, 2.23, 16.80
[1560401967] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.69, 0.67, 4.67
[1560402020] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (221436 kB) free!
[1560402140] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (10005172 kB) free.
[1560402740] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (230464 kB) free!
[1560402860] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.2% (9475624 kB) free!
[1560402911] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46875 MB (20% inode=99%):
[1560402980] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.6% (6702720 kB) free!
[1560403031] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45464 MB (20% inode=99%):
[1560403152] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44908 MB (19% inode=99%):
[1560403729] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.47, 2.65, 4.59
[1560404138] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9631020 kB) free!
[1560404181] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (220192 kB) free!
[1560404245] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1560404258] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8894856 kB) free!
[1560404378] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (8901056 kB) free!
[1560404846] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1560404966] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560405086] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1560406581] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.4% (6125824 kB) free!
[1560406729] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 158.17, 97.20, 43.40
[1560406767] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.95, 6.02, 6.98
[1560406849] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 160.27, 118.16, 57.61
[1560406887] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.30, 4.45, 6.29
[1560406969] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1560406969] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 160.51, 132.18, 70.09
[1560407007] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1560407007] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.90, 3.28, 5.64
[1560407111] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (345928 kB) free!
[1560407231] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (124488 kB) free!
[1560407351] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (122124 kB) free!
[1560407607] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.99, 1.09, 3.28
[1560407781] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9941900 kB) free.
[1560407952] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1560407952] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 91.5% (6731800 kB) free.
[1560408277] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.9% (3074284 kB) free!
[1560408314] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560408327] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183432 kB) free!
[1560408381] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (229084 kB) free!
[1560408387] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189500 kB) free!
[1560408397] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 20.0% (3249488 kB) free!
[1560408447] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (189800 kB) free!
[1560408501] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10331788 kB) free.
[1560408507] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (186636 kB) free!
[1560408517] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.7% (3520884 kB) free.
[1560408769] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.85, 3.40, 15.08
[1560409102] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (249136 kB) free!
[1560409222] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.9% (9806204 kB) free!
[1560409342] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.0% (7391892 kB) free!
[1560409707] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 22.9% (5623852 kB) free.
[1560409718] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (918852 kB) free!
[1560409838] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (1456648 kB) free!
[1560409886] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1560409958] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.2% (1986320 kB) free!
[1560410082] SERVICE FLAPPING ALERT: redis-1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560410487] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1560410543] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (222444 kB) free!
[1560410569] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.52, 2.20, 4.10
[1560410607] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1560410727] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1560410727] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1560411158] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.2% (3759080 kB) free.
[1560411334] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1560411578] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4187188 kB) free!
[1560411758] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.9% (2421084 kB) free!
[1560411878] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.7% (3519716 kB) free.
[1560412343] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.4% (5135800 kB) free!
[1560413952] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (144072 kB) free!
[1560413952] SERVICE ALERT: smscapp1;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 24570 MB (10% inode=99%):
[1560414072] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (123108 kB) free!
[1560414143] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (229388 kB) free!
[1560414192] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (122140 kB) free!
[1560414607] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 208.35, 94.82, 36.91
[1560414727] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 220.48, 137.04, 59.40
[1560414848] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 223.17, 165.71, 79.32
[1560414878] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (217040 kB) free!
[1560414998] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (1565292 kB) free!
[1560415097] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (2534316 kB) free!
[1560415115] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 353 processes
[1560415118] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (1568532 kB) free!
[1560415235] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 353 processes
[1560415355] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 353 processes
[1560415697] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1849544 kB) free!
[1560415719] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1560415719] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.6% (4153656 kB) free.
[1560415943] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.1% (5955876 kB) free!
[1560416319] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (350144 kB) free!
[1560416439] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 11.8% (1914232 kB) free!
[1560416555] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1560416559] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.8% (3703832 kB) free.
[1560416648] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.80, 1.83, 13.57
[1560417143] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (221444 kB) free!
[1560417159] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (1109808 kB) free!
[1560417279] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.9% (2259512 kB) free!
[1560417399] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.9% (4197132 kB) free.
[1560417848] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.34, 0.57, 4.23
[1560417999] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.0% (2111400 kB) free!
[1560418119] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.2% (4088616 kB) free.
[1560418343] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.2% (6003428 kB) free!
[1560418392] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.9% (6834608 kB) free.
[1560418719] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (907488 kB) free!
[1560418839] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.8% (3378252 kB) free.
[1560419007] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560419135] SERVICE ALERT: sigqscluster1;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1560419307] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 13.2% (3225500 kB) free!
[1560419367] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 20.0% (4912760 kB) free!
[1560419427] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 19.1% (4680668 kB) free!
[1560419439] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.7% (2543468 kB) free!
[1560419487] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 17.2% (4209616 kB) free!
[1560419544] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237288 kB) free!
[1560419559] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.1% (2780304 kB) free!
[1560419607] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 56.30, 140.04, 82.81
[1560419679] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (168816 kB) free!
[1560419727] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 8.11, 93.90, 72.84
[1560419735] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1560419847] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.92, 63.07, 64.09
[1560419855] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1560419975] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 12 users currently logged in
[1560420280] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 27.6% (4482212 kB) free.
[1560420446] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (2395740 kB) free!
[1560420755] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 351 processes
[1560420769] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1560420875] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1560420880] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 11.7% (1898824 kB) free!
[1560421000] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.8% (4182924 kB) free.
[1560421046] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (201608 kB) free!
[1560421105] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 33.8% (8277428 kB) free.
[1560421219] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 34.9% (8616384 kB) free.
[1560421344] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.6% (8671340 kB) free!
[1560421600] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (1340976 kB) free!
[1560421647] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.50, 3.56, 13.50
[1560421705] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3171548 kB) free!
[1560421720] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (639592 kB) free!
[1560421819] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190792 kB) free!
[1560421825] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 20.4% (3317756 kB) free.
[1560421840] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.2% (2311380 kB) free!
[1560421939] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190056 kB) free!
[1560421944] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (215740 kB) free!
[1560422059] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (213700 kB) free!
[1560422127] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560422305] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179128 kB) free!
[1560422338] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.1% (2459892 kB) free!
[1560422352] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 130498 MB (58% inode=99%):
[1560422425] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 15.7% (2550556 kB) free!
[1560422426] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173952 kB) free!
[1560422440] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.2% (4093644 kB) free.
[1560422545] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 14.7% (2396704 kB) free!
[1560422546] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1560422546] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194740 kB) free!
[1560422665] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 13.7% (2230460 kB) free!
[1560422847] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.14, 1.29, 4.71
[1560422938] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1178504 kB) free!
[1560423145] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.0% (6922284 kB) free!
[1560423248] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.23, 6.49, 8.26
[1560423265] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1558916 kB) free!
[1560423358] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1560423368] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.38, 4.81, 7.42
[1560423387] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 37.3% (9139696 kB) free.
[1560423478] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1560423488] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1560423488] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.08, 3.60, 6.67
[1560423599] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1560423745] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233128 kB) free!
[1560424089] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.55, 1.43, 4.06
[1560424240] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (181424 kB) free!
[1560424360] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 10.8% (1751112 kB) free!
[1560424392] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 13.0% (954204 kB) free!
[1560424480] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.2% (3446856 kB) free.
[1560424512] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (692284 kB) free!
[1560424632] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (436712 kB) free!
[1560425080] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (175440 kB) free!
[1560425200] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (969468 kB) free!
[1560425321] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.5% (3977852 kB) free.
[1560425529] SERVICE ALERT: redis-2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1560425532] HOST ALERT: redis-2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.202)
[1560425537] SERVICE ALERT: redis-1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1560425540] HOST ALERT: redis-1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.201)
[1560425545] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.7% (7756272 kB) free!
[1560425552] SERVICE ALERT: redis-1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'redis-1' was out of bounds)
[1560425587] SERVICE ALERT: redis-2;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'redis-2' was out of bounds)
[1560425590] SERVICE ALERT: redis-2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'redis-2' was out of bounds)
[1560425618] SERVICE ALERT: redis-1;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'redis-1' was out of bounds)
[1560425652] SERVICE ALERT: redis-2;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'redis-2' was out of bounds)
[1560425655] SERVICE ALERT: redis-2;HDD_VNS1 status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_VNS1 status' on host 'redis-2' was out of bounds)
[1560425660] SERVICE ALERT: redis-1;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'redis-1' was out of bounds)
[1560425687] SERVICE ALERT: redis-1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'redis-1' was out of bounds)
[1560425735] SERVICE ALERT: redis-2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'redis-2' was out of bounds)
[1560425853] SERVICE ALERT: redis-1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'redis-1' was out of bounds)
[1560425920] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'redis-1' was out of bounds)
[1560425939] SERVICE ALERT: redis-1;HDD_VNS1 status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_VNS1 status' on host 'redis-1' was out of bounds)
[1560425939] SERVICE ALERT: redis-1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'redis-1' was out of bounds)
[1560425974] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 0.36, 1.60, 4.07
[1560425995] SERVICE ALERT: redis-2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'redis-2' was out of bounds)
[1560426050] SERVICE ALERT: redis-2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'redis-2' was out of bounds)
[1560426058] SERVICE ALERT: redis-2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'redis-2' was out of bounds)
[1560426088] SERVICE ALERT: redis-1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'redis-1' was out of bounds)
[1560426093] HOST ALERT: redis-2;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.202)
[1560426105] HOST ALERT: redis-1;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.201)
[1560426145] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (1379228 kB) free!
[1560426745] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6078352 kB) free!
[1560427346] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233228 kB) free!
[1560428065] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 10.2% (1662108 kB) free!
[1560428727] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 351 processes
[1560428978] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.5% (7168936 kB) free!
[1560429040] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 58.8% (9534732 kB) free.
[1560429044] HOST ALERT: redis-1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1560429059] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 815991 MB (76% inode=99%):
[1560429059] SERVICE ALERT: redis-1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 45065 MB (88% inode=99%):
[1560429097] HOST ALERT: redis-2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1560429115] SERVICE ALERT: redis-2;Total Processes;OK;HARD;1;PROCS OK: 347 processes
[1560429170] SERVICE ALERT: redis-2;CPU Load;OK;HARD;1;OK - load average: 0.92, 0.40, 0.15
[1560429178] SERVICE ALERT: redis-2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1560429207] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 636 processes
[1560429208] SERVICE ALERT: redis-1;CPU Load;OK;HARD;1;OK - load average: 0.98, 0.55, 0.23
[1560429249] SERVICE ALERT: redis-2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 270916 MB (68% inode=99%):
[1560429257] SERVICE ALERT: redis-1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1560429266] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1560429266] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 88.8% (14432420 kB) free.
[1560429267] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 633 processes
[1560429272] SERVICE ALERT: redis-1;Total Processes;OK;HARD;1;PROCS OK: 353 processes
[1560429307] SERVICE ALERT: redis-2;Sensors Status;OK;HARD;1;SENSORS OK
[1560429310] SERVICE ALERT: redis-2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 45677 MB (89% inode=99%):
[1560429338] SERVICE ALERT: redis-1;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1560429375] SERVICE ALERT: redis-2;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 1069171 MB (99% inode=99%):
[1560429407] SERVICE ALERT: redis-1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 379865 MB (95% inode=99%):
[1560429432] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 90.5% (6657564 kB) free.
[1560429456] SERVICE ALERT: redis-2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1560429489] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 43.08, 153.44, 85.01
[1560429573] SERVICE ALERT: redis-1;Sensors Status;OK;HARD;1;SENSORS OK
[1560429578] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3742428 kB) free!
[1560429609] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 7.44, 103.29, 74.93
[1560429640] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (342212 kB) free!
[1560429729] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.23, 69.55, 66.00
[1560429760] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.6% (3834456 kB) free.
[1560430775] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.46, 8.23, 6.65
[1560430799] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1560430895] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.53, 8.53, 6.93
[1560430929] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.28, 2.22, 18.80
[1560430960] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (162220 kB) free!
[1560431015] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 13.61, 10.29, 7.77
[1560431080] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (336908 kB) free!
[1560431200] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.9% (4044388 kB) free.
[1560431569] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 125.24, 38.38, 14.58
[1560431690] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 180.23, 88.07, 35.67
[1560431810] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 186.94, 121.14, 54.14
[1560431991] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.5% (15035564 kB) free.
[1560432047] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.5% (23161276 kB) free.
[1560432102] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.0% (14956192 kB) free.
[1560432143] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 84.3% (41561868 kB) free.
[1560432239] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.9% (23023752 kB) free.
[1560432281] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.3% (23270904 kB) free.
[1560432327] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.5% (7640692 kB) free!
[1560432729] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.29, 1.57, 3.78
[1560432849] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191720 kB) free!
[1560432885] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193892 kB) free!
[1560432927] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1560432947] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10909004 kB) free.
[1560432969] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194032 kB) free!
[1560433000] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (300620 kB) free!
[1560433005] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176440 kB) free!
[1560433090] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191940 kB) free!
[1560433121] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (1286784 kB) free!
[1560433125] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164824 kB) free!
[1560433241] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (196644 kB) free!
[1560433257] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193384 kB) free!
[1560433377] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181300 kB) free!
[1560433497] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185956 kB) free!
[1560433547] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3936304 kB) free!
[1560433668] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (1361712 kB) free!
[1560433788] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.8% (7274140 kB) free!
[1560433839] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.9% (22773160 kB) free.
[1560433905] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (171992 kB) free!
[1560433963] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 73.9% (18114232 kB) free.
[1560433979] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.3% (21296104 kB) free.
[1560434025] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (174776 kB) free!
[1560434075] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 351 processes
[1560434145] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (182264 kB) free!
[1560434195] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1560434210] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.49, 3.40, 13.15
[1560434388] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (225708 kB) free!
[1560434399] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1560434448] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172604 kB) free!
[1560434519] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1560434566] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174572 kB) free!
[1560434568] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191036 kB) free!
[1560434585] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198964 kB) free!
[1560434639] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1560434686] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182996 kB) free!
[1560434688] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171944 kB) free!
[1560434705] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166724 kB) free!
[1560434812] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1560434812] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177136 kB) free!
[1560434825] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1560434825] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188700 kB) free!
[1560435114] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1560435188] SERVICE ALERT: redis-2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1560435218] HOST ALERT: redis-2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1560435218] SERVICE ALERT: redis-1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1560435221] HOST ALERT: redis-1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.201)
[1560435252] SERVICE ALERT: redis-2;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'redis-2' was out of bounds)
[1560435260] SERVICE ALERT: redis-1;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'redis-1' was out of bounds)
[1560435275] SERVICE ALERT: redis-1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'redis-1' was out of bounds)
[1560435310] SERVICE ALERT: redis-2;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'redis-2' was out of bounds)
[1560435310] SERVICE ALERT: redis-2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'redis-2' was out of bounds)
[1560435313] SERVICE ALERT: redis-2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'redis-2' was out of bounds)
[1560435341] SERVICE ALERT: redis-1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'redis-1' was out of bounds)
[1560435341] SERVICE ALERT: redis-1;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'redis-1' was out of bounds)
[1560435378] SERVICE ALERT: redis-2;HDD_VNS1 status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_VNS1 status' on host 'redis-2' was out of bounds)
[1560435410] SERVICE ALERT: redis-1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'redis-1' was out of bounds)
[1560435458] SERVICE ALERT: redis-2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'redis-2' was out of bounds)
[1560435575] SERVICE ALERT: redis-1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'redis-1' was out of bounds)
[1560435588] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9914296 kB) free.
[1560435641] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.0% (3729436 kB) free.
[1560435645] HOST ALERT: redis-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1560435660] HOST ALERT: redis-2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1560435695] SERVICE FLAPPING ALERT: redis-1;Sensors Status;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1560435695] SERVICE ALERT: redis-1;Sensors Status;OK;HARD;1;SENSORS OK
[1560435729] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.64, 5.82, 5.01
[1560435815] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 0.14, 0.96, 2.92
[1560435849] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.28, 3.97, 4.43
[1560435959] HOST ALERT: redis-2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.202)
[1560435971] HOST ALERT: redis-1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1560436188] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.2% (5026944 kB) free!
[1560436244] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'redis-1' was out of bounds)
[1560436262] SERVICE ALERT: redis-1;HDD_VNS1 status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_VNS1 status' on host 'redis-1' was out of bounds)
[1560436262] SERVICE ALERT: redis-1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'redis-1' was out of bounds)
[1560436298] SERVICE ALERT: redis-1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'redis-1' was out of bounds)
[1560436308] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (4579748 kB) free!
[1560436318] SERVICE ALERT: redis-2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'redis-2' was out of bounds)
[1560436348] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (2837532 kB) free!
[1560436373] SERVICE ALERT: redis-2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'redis-2' was out of bounds)
[1560436428] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (3538228 kB) free!
[1560436468] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (1301628 kB) free!
[1560436526] HOST ALERT: redis-2;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.202)
[1560436538] HOST ALERT: redis-1;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.201)
[1560436589] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1560436589] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (856372 kB) free!
[1560436610] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.03, 2.10, 3.67
[1560436632] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1560437947] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.6% (23099936 kB) free.
[1560438003] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47070 MB (20% inode=99%):
[1560438016] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.3% (22628388 kB) free.
[1560438076] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.7% (40283920 kB) free.
[1560438111] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.8% (22985156 kB) free.
[1560438123] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46839 MB (20% inode=99%):
[1560438143] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.2% (14988184 kB) free.
[1560438215] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.23, 9.25, 7.04
[1560438228] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10221132 kB) free.
[1560438243] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46839 MB (20% inode=99%):
[1560438335] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.22, 8.36, 6.97
[1560438455] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.14, 7.39, 6.77
[1560438828] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (2511476 kB) free!
[1560438948] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (883384 kB) free!
[1560439068] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (220800 kB) free!
[1560440066] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1560440268] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.2% (7474984 kB) free!
[1560440667] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2988128 kB) free!
[1560440787] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.4% (2661560 kB) free!
[1560440847] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 53.85, 13.79, 4.88
[1560440868] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (234480 kB) free!
[1560440907] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.5% (2361180 kB) free!
[1560440967] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 112.00, 49.29, 18.97
[1560441087] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 120.05, 73.15, 31.39
[1560441507] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (850268 kB) free!
[1560441839] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1560442287] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.64, 3.00, 13.93
[1560442633] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (649512 kB) free!
[1560442668] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.8% (5826888 kB) free!
[1560442753] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (420948 kB) free!
[1560442873] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (186224 kB) free!
[1560443049] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1560443210] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 472.83, 368.65, 183.01
[1560443330] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 473.32, 403.30, 218.22
[1560443451] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 126.44, 315.80, 210.26
[1560443487] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.41, 0.53, 4.22
[1560443546] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.6% (3185480 kB) free!
[1560443640] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1560443666] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.3% (2982336 kB) free!
[1560443760] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1560443786] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.2% (2789028 kB) free!
[1560443855] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.64, 3.83, 4.64
[1560443869] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (232048 kB) free!
[1560443880] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1560444114] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1560444234] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1560444354] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1560444628] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1560444642] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1560444642] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.5% (14864872 kB) free.
[1560444673] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.1% (6844192 kB) free.
[1560444711] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 87.4% (14204848 kB) free.
[1560445069] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.1% (5940384 kB) free!
[1560445155] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1560445227] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 381 processes
[1560445347] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 375 processes
[1560445595] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 352 processes
[1560445669] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (225292 kB) free!
[1560445716] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 355 processes
[1560445797] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (168668 kB) free!
[1560445830] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190148 kB) free!
[1560445836] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 354 processes
[1560445851] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1560445851] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.75, 2.38, 17.74
[1560445915] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1560445917] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (157560 kB) free!
[1560445950] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196284 kB) free!
[1560446037] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (163864 kB) free!
[1560446070] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186228 kB) free!
[1560446727] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1560446870] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.3% (7031776 kB) free!
[1560447036] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1560447085] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1560447167] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560447221] SERVICE FLAPPING ALERT: redis-1;Sensors Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560447243] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 152887 MB (68% inode=99%):
[1560447270] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 79.4% (19462536 kB) free.
[1560447327] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1560447447] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1560447470] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (224376 kB) free!
[1560447567] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1560447651] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.73, 1.70, 4.19
[1560448671] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.0% (7389272 kB) free!
[1560448947] SERVICE ALERT: sigqscluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 380 processes
[1560449067] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 381 processes
[1560449093] SERVICE ALERT: redis-2;CPU Load;OK;HARD;1;OK - load average: 0.95, 0.36, 0.13
[1560449097] HOST ALERT: redis-2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1560449108] SERVICE ALERT: redis-2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1560449138] SERVICE ALERT: redis-1;CPU Load;OK;HARD;1;OK - load average: 0.99, 0.40, 0.15
[1560449142] HOST ALERT: redis-1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1560449172] SERVICE ALERT: redis-2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 261404 MB (65% inode=99%):
[1560449180] SERVICE ALERT: redis-1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1560449195] SERVICE ALERT: redis-1;Total Processes;OK;HARD;1;PROCS OK: 386 processes
[1560449230] SERVICE ALERT: redis-2;Sensors Status;OK;HARD;1;SENSORS OK
[1560449234] SERVICE ALERT: redis-2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 45677 MB (89% inode=99%):
[1560449262] SERVICE ALERT: redis-1;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1560449298] SERVICE ALERT: redis-2;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 1069171 MB (99% inode=99%):
[1560449330] SERVICE ALERT: redis-1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 378569 MB (95% inode=99%):
[1560449378] SERVICE ALERT: redis-2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1560449565] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.9% (4040756 kB) free.
[1560449582] SERVICE ALERT: redis-1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 45060 MB (88% inode=99%):
[1560449582] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 538349 MB (50% inode=99%):
[1560449618] SERVICE ALERT: redis-1;Sensors Status;OK;HARD;1;SENSORS OK
[1560449638] SERVICE ALERT: redis-2;Total Processes;OK;HARD;1;PROCS OK: 334 processes
[1560449871] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233336 kB) free!
[1560450164] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.6% (2528948 kB) free!
[1560450284] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.3% (2484028 kB) free!
[1560450404] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.1% (2452368 kB) free!