[1590085819] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (1137008 kB) free!
[1590085939] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.1% (2453932 kB) free!
[1590086001] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (2569800 kB) free!
[1590086059] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.9% (3556988 kB) free.
[1590086121] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (401780 kB) free!
[1590086241] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (423224 kB) free!
[1590086841] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27736300 kB) free.
[1590087086] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1590087259] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (516272 kB) free!
[1590087379] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.5% (1703536 kB) free!
[1590087436] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(809) < 172.105.59.173 [22/May/2020:00:21:54 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1590087441] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (1126040 kB) free!
[1590087499] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1590087499] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.7% (2877708 kB) free!
[1590087556] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590087561] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (415800 kB) free!
[1590087681] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (424116 kB) free!
[1590087686] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 358 processes
[1590087806] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 349 processes
[1590087926] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 353 processes
[1590088156] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(558) < 172.105.59.173 [22/May/2020:00:38:55 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1590088276] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(62) < 203.212.70.170 [22/May/2020:00:40:45 +0530] TCP 502 0 0 120.001 "10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1590088281] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20449296 kB) free!
[1590088396] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590088426] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(607) < 182.18.144.246 [22/May/2020:00:43:20 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1590088526] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1590088546] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(43) < 172.105.57.57 [22/May/2020:00:45:36 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1590088667] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;HARD;3;(536) < 182.18.144.246 [22/May/2020:00:46:01 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1590088673] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1694) < 2020/05/22 00:47:09 [error] 22121#0: *89878139 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1590088793] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590088881] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (409796 kB) free!
[1590088996] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(108) < 185.255.8.58 [22/May/2020:00:50:59 +0530] TCP 502 0 0 0.001 "10.147.212.143:13001, 10.147.212.154:13001, 10.147.212.142:13001, 10.147.212.155:13001" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 0.001, 0.000, 0.000"
[1590089117] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590089267] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590089299] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1276472 kB) free!
[1590089311] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590089482] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27968212 kB) free.
[1590089624] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (170244 kB) free!
[1590089902] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (3048820 kB) free!
[1590090010] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1703 processes
[1590090082] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (383720 kB) free!
[1590090202] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (392348 kB) free!
[1590090262] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 30774 MB (13.69% inode=100%):
[1590090322] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (387552 kB) free!
[1590090322] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 25134 MB (11.18% inode=100%):
[1590090382] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;3;DISK WARNING - free space: /home 23680 MB (10.53% inode=100%):
[1590090442] SERVICE ALERT: webserver2;HDD_Home status;CRITICAL;HARD;4;DISK CRITICAL - free space: /home 18038 MB (8.02% inode=100%):
[1590090610] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1698 processes
[1590090922] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (17062036 kB) free!
[1590091042] SERVICE ALERT: webserver2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 68981 MB (30.69% inode=100%):
[1590091102] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (833900 kB) free!
[1590091110] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (239904 kB) free!
[1590091124] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 87.3% (21398996 kB) free.
[1590091231] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (237696 kB) free!
[1590091304] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.0% (21083128 kB) free!
[1590091351] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233900 kB) free!
[1590091425] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (10415228 kB) free!
[1590091522] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399832 kB) free!
[1590091545] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401248 kB) free!
[1590091702] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.4% (2665928 kB) free!
[1590092122] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17211140 kB) free!
[1590092411] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1703 processes
[1590092722] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (414352 kB) free!
[1590093151] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.4% (8568784 kB) free!
[1590093322] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20147412 kB) free!
[1590093922] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (418364 kB) free!
[1590094522] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.5% (21726568 kB) free!
[1590095122] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408644 kB) free!
[1590095311] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.08, 6.07, 5.52
[1590095432] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.64, 5.87, 5.51
[1590095436] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (4549524 kB) free!
[1590095552] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.07, 5.21, 5.32
[1590095635] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.70, 5.31, 5.14
[1590095722] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23164128 kB) free!
[1590095755] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.26, 5.01, 5.05
[1590095875] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.49, 4.64, 4.92
[1590096011] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1700 processes
[1590096036] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (237648 kB) free!
[1590096141] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 90.6% (44648920 kB) free.
[1590096152] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.25, 3.81, 4.57
[1590096322] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402244 kB) free!
[1590096502] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (1622092 kB) free!
[1590096592] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.0% (44377336 kB) free.
[1590096752] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.0% (44353060 kB) free.
[1590096916] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2328) < 203.212.70.170 [22/May/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1590096922] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.0% (21135564 kB) free!
[1590097036] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590097068] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(408) < 103.16.101.72 [22/May/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1590097102] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.6% (2209908 kB) free!
[1590097188] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590097193] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(536) < 2020/05/22 03:01:32 [error] 22135#0: *89895117 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1590097211] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1705 processes
[1590097313] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590097522] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386520 kB) free!
[1590098123] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16245144 kB) free!
[1590098723] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (413372 kB) free!
[1590099012] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1696 processes
[1590099323] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.9% (23600644 kB) free!
[1590099502] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3432260 kB) free.
[1590099545] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590099923] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407816 kB) free!
[1590100145] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(60) < 2020/05/22 03:59:00 [error] 14372#0: *48055488 access forbidden by rule while initializing session, client: 103.209.99.7, server: 0.0.0.0:13000
[1590100265] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(12) < 2020/05/22 04:01:03 [error] 14367#0: *48055571 access forbidden by rule while initializing session, client: 103.209.99.7, server: 0.0.0.0:13000
[1590100352] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.8% (4855888 kB) free!
[1590100385] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(12) < 2020/05/22 04:03:03 [error] 14376#0: *48055664 access forbidden by rule while initializing session, client: 103.209.99.7, server: 0.0.0.0:13000
[1590100523] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.0% (23682660 kB) free!
[1590100952] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190464 kB) free!
[1590101123] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (415072 kB) free!
[1590101413] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1701 processes
[1590101723] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27463856 kB) free.
[1590102323] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4379688 kB) free!
[1590102443] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (402336 kB) free!
[1590102563] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398544 kB) free!
[1590102613] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1691 processes
[1590103763] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (31013080 kB) free.
[1590104363] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8908224 kB) free!
[1590104413] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1701 processes
[1590104483] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (5071172 kB) free!
[1590104603] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (1274876 kB) free!
[1590105204] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17110644 kB) free!
[1590105614] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1695 processes
[1590105643] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 73.3% (17958372 kB) free.
[1590105804] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (443580 kB) free!
[1590106236] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 52.0% (12737724 kB) free.
[1590106404] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17763140 kB) free!
[1590106441] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.3% (3764272 kB) free!
[1590107004] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30445904 kB) free.
[1590107041] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194252 kB) free!
[1590107604] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8378596 kB) free!
[1590107724] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (4165124 kB) free!
[1590107844] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.2% (21302192 kB) free!
[1590107875] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.98, 5.32, 5.13
[1590107995] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.02, 5.39, 5.17
[1590108015] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1702 processes
[1590108115] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.07, 5.81, 5.37
[1590108444] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (13225600 kB) free!
[1590108615] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1699 processes
[1590108715] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.28, 4.68, 4.99
[1590109044] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29121936 kB) free.
[1590109102] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1590109315] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.76, 4.78, 5.12
[1590109435] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.01, 4.90, 5.12
[1590109556] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.27, 4.23, 4.85
[1590109644] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8725840 kB) free!
[1590109764] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (4973392 kB) free!
[1590109884] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34201600 kB) free.
[1590110024] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 15.0% (3670604 kB) free!
[1590110084] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (2037236 kB) free!
[1590110144] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 1.7% (425664 kB) free!
[1590110204] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (179436 kB) free!
[1590110436] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199968 kB) free!
[1590110484] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14635156 kB) free!
[1590110505] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 19.2% (4709300 kB) free!
[1590110556] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168364 kB) free!
[1590110604] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (10803440 kB) free!
[1590110676] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181332 kB) free!
[1590110724] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (6963944 kB) free!
[1590110805] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 9.9% (2417928 kB) free!
[1590111324] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19259048 kB) free!
[1590111345] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 74.0% (97557492 kB) free.
[1590111655] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.14643.556> (10.147.212.158:43598 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590111775] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590111925] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33599380 kB) free.
[1590112525] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11350928 kB) free!
[1590112555] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.59, 5.63, 5.03
[1590112645] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7380556 kB) free!
[1590112675] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.36, 4.93, 4.84
[1590112765] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3493104 kB) free!
[1590113275] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.12, 5.30, 5.13
[1590113365] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.0% (14539712 kB) free!
[1590113395] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.87, 5.28, 5.15
[1590113515] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.01, 5.25, 5.16
[1590113965] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (9423056 kB) free!
[1590115215] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1701 processes
[1590115528] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1590115648] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 343 processes
[1590115765] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.7% (23293000 kB) free!
[1590115768] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1590115815] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1689 processes
[1590115915] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1590115915] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.07, 4.73, 4.94
[1590116365] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1402944 kB) free!
[1590116965] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.6% (21858376 kB) free!
[1590117565] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385828 kB) free!
[1590117715] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.31, 5.71, 5.19
[1590117775] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 12.2% (1988776 kB) free!
[1590117835] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.64, 5.21, 5.07
[1590117955] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.89, 5.17, 5.08
[1590118165] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16534104 kB) free!
[1590118352] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 13.40, 6.56, 5.01
[1590118472] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 20.43, 10.86, 6.74
[1590118593] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 13.22, 10.94, 7.24
[1590118765] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391968 kB) free!
[1590119365] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13364600 kB) free!
[1590119966] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393064 kB) free!
[1590120393] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 25.51, 27.40, 19.77
[1590122368] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1590122902] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (2952500 kB) free!
[1590122958] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1590123022] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.3% (3623992 kB) free.
[1590123043] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189048 kB) free!
[1590123078] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1590123163] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183720 kB) free!
[1590123198] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1590123283] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187744 kB) free!
[1590123393] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 26.22, 20.91, 19.91
[1590123436] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(91) < 103.255.146.164 [22/May/2020:10:25:09 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590123556] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590123622] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3176260 kB) free!
[1590123742] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (3116428 kB) free!
[1590123799] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1590123862] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.6% (2695732 kB) free!
[1590123978] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 747 processes
[1590123993] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1590123993] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 34.47, 32.12, 26.03
[1590124098] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 747 processes
[1590124208] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.7% (3311188 kB) free!
[1590124218] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 748 processes
[1590124329] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.2% (2745260 kB) free!
[1590124449] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2135952 kB) free!
[1590125415] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1780 processes
[1590125523] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 53.9% (13213920 kB) free.
[1590125677] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.3% (2779748 kB) free!
[1590126277] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1636856 kB) free!
[1590127316] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 17.8% (6236232 kB) free!
[1590127436] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 16.1% (5631604 kB) free!
[1590127556] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 14.6% (5102252 kB) free!
[1590127756] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(268) < 103.255.146.164 [22/May/2020:11:31:45 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590127766] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.3% (13571748 kB) free!
[1590127876] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590128063] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.3% (3785148 kB) free.
[1590128105] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 14.2% (3515956 kB) free!
[1590128155] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1590128156] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (2374552 kB) free!
[1590128226] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 12.4% (3070860 kB) free!
[1590128346] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 10.6% (2623988 kB) free!
[1590128366] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (444836 kB) free!
[1590128541] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9628508 kB) free!
[1590128661] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (9365800 kB) free!
[1590128781] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (9056156 kB) free!
[1590128946] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (366584 kB) free!
[1590128967] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35033260 kB) free.
[1590128992] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.3% (8526664 kB) free!
[1590128994] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1590129112] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.7% (8253540 kB) free!
[1590129114] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1590129123] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.7% (4102620 kB) free!
[1590129232] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.9% (7853888 kB) free!
[1590129244] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.4% (3783224 kB) free!
[1590129249] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.6% (16077428 kB) free.
[1590129357] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 91.0% (31865460 kB) free.
[1590129364] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.5% (3565204 kB) free!
[1590129567] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (2130580 kB) free!
[1590129687] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (388008 kB) free!
[1590129807] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (439356 kB) free!
[1590129863] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.7% (2877460 kB) free!
[1590129983] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.9% (2582696 kB) free!
[1590130043] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.1% (4980808 kB) free!
[1590130103] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.7% (3848244 kB) free.
[1590130146] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.2% (22998904 kB) free.
[1590130407] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.0% (23699380 kB) free!
[1590130432] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4788708 kB) free!
[1590130564] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1992748 kB) free!
[1590130643] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (264600 kB) free!
[1590131007] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398240 kB) free!
[1590131181] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (238228 kB) free!
[1590131418] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1590131418] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 657 processes
[1590131903] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.8% (2568140 kB) free!
[1590132018] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 738 processes
[1590132023] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.1% (2447124 kB) free!
[1590132138] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 741 processes
[1590132143] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.7% (2225552 kB) free!
[1590132259] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 741 processes
[1590133277] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(386) < 103.255.146.164 [22/May/2020:13:01:38 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590133313] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/05/22 13:10:43 [error] 22127#0: *89990542 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590133375] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.24008.662> (10.147.212.147:58842 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590133397] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590133407] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32016272 kB) free.
[1590133433] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/05/22 13:12:57 [error] 22135#0: *89991142 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590133496] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590133553] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590134007] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3718680 kB) free!
[1590134127] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393824 kB) free!
[1590134153] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/22 13:20:27 [error] 22135#0: *89993072 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590134247] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399388 kB) free!
[1590134273] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590134847] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21581200 kB) free!
[1590135295] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.32415.649> (10.147.212.158:55728 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590135393] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590135416] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590135448] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391760 kB) free!
[1590136048] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33423700 kB) free.
[1590136314] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 342 processes
[1590136434] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1590136648] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4924336 kB) free!
[1590136673] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/22 14:00:48 [error] 22131#0: *90002100 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590136768] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (750800 kB) free!
[1590136793] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590136888] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (441400 kB) free!
[1590136975] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (403452 kB) free!
[1590137488] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.9% (20981252 kB) free!
[1590137597] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(905) < 103.255.146.164 [22/May/2020:14:14:15 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590137599] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1590137616] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1590137634] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1590137717] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590137719] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 349 processes
[1590137736] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1590137754] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1590137755] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 6.94, 5.37, 4.98
[1590137839] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1590137856] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1590137993] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/22 14:29:27 [error] 22135#0: *90008855 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590138088] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382964 kB) free!
[1590138113] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590138143] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (1489264 kB) free!
[1590138439] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1590138559] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1590138679] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1590138688] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28285288 kB) free.
[1590138743] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1590138743] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (2534184 kB) free!
[1590139056] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1590139288] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4646456 kB) free!
[1590139343] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (973604 kB) free!
[1590139408] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (915284 kB) free!
[1590139528] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (436492 kB) free!
[1590139555] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.08, 6.27, 5.38
[1590139675] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.02, 5.70, 5.30
[1590139795] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 2.15, 4.50, 4.91
[1590139913] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/22 14:59:23 [error] 22135#0: *90015444 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590139943] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.9% (2415096 kB) free!
[1590140033] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/05/22 15:02:48 [error] 22135#0: *90016167 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590140118] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(463) < 103.255.146.164 [22/May/2020:15:00:07 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590140153] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590140238] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590140543] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (471992 kB) free!
[1590141143] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.9% (1930320 kB) free!
[1590141353] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/22 15:23:51 [error] 22135#0: *90020834 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590141438] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(623) < 103.255.146.164 [22/May/2020:15:17:40 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590141473] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590141558] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590141595] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.72, 6.29, 5.05
[1590141715] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.86, 6.26, 5.19
[1590141835] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.15, 6.13, 5.29
[1590142529] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (15929660 kB) free!
[1590142943] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (630560 kB) free!
[1590143129] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (413604 kB) free!
[1590143273] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/22 15:53:13 [error] 22135#0: *90026397 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/dashboard HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/dashboard", host: "altsmsc.tesync.net", referrer: "http://wap.tesync.net/tesync/AnalysisSMS.jsp"
[1590143393] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590143543] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (3182092 kB) free!
[1590143729] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (22847916 kB) free!
[1590143968] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1590144329] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408944 kB) free!
[1590145168] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1590145288] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1590145343] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1316664 kB) free!
[1590145758] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(457) < 103.255.146.164 [22/May/2020:16:31:29 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590145878] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590146394] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/05/22 16:47:43 [error] 22135#0: *90036828 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590146489] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1590146514] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/05/22 16:50:58 [error] 22134#0: *90037554 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590146544] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.4% (2175384 kB) free!
[1590146609] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1590146634] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590146649] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.0% (2909512 kB) free!
[1590146729] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1590146729] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19587352 kB) free!
[1590146769] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.4% (2569396 kB) free!
[1590146889] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2219436 kB) free!
[1590147144] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1602576 kB) free!
[1590147234] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2020/05/22 16:57:09 [error] 22135#0: *90038850 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1590147329] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1590147329] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406492 kB) free!
[1590147354] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590147678] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(657) < 103.255.146.164 [22/May/2020:17:01:23 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1590147798] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590147929] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (25649492 kB) free!
[1590148344] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.5% (2027408 kB) free!
[1590148529] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (420376 kB) free!
[1590148554] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/22 17:17:59 [error] 22135#0: *90042440 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590148674] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590149544] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3337192 kB) free.
[1590149874] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/22 17:40:19 [error] 22135#0: *90046438 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590149994] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590150594] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/22 17:57:41 [error] 22125#0: *90049685 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590150714] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590150744] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (1056484 kB) free!
[1590150864] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (190520 kB) free!
[1590150929] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 351 processes
[1590150985] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.6% (1883348 kB) free!
[1590151049] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 348 processes
[1590151090] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.8% (16118776 kB) free.
[1590151146] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 16.7% (4122688 kB) free!
[1590151170] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1590151170] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1590151266] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 15.2% (3739676 kB) free!
[1590151314] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/22 18:08:32 [error] 22134#0: *90051946 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/dashboard HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/dashboard", host: "altsmsc.tesync.net"
[1590151386] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 13.6% (3356060 kB) free!
[1590151434] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/05/22 18:13:25 [error] 22135#0: *90052949 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590151554] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590151585] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (526488 kB) free!
[1590151770] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 351 processes
[1590151986] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.2% (22999448 kB) free.
[1590152129] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (32223912 kB) free.
[1590152155] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/05/22 18:23:17 [error] 22135#0: *90054780 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1590152185] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.2% (4913600 kB) free.
[1590152275] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590152729] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (399328 kB) free!
[1590152786] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (173188 kB) free!
[1590152849] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (428724 kB) free!
[1590152906] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (1517184 kB) free!
[1590152969] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (480076 kB) free!
[1590152970] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1590153026] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.1% (4722096 kB) free.
[1590153569] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32856940 kB) free.
[1590153570] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1590153626] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (493776 kB) free!
[1590153746] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.6% (2045596 kB) free!
[1590153866] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.0% (3744064 kB) free.
[1590154169] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (385612 kB) free!
[1590154289] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (395636 kB) free!
[1590154409] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (472392 kB) free!
[1590155009] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33812356 kB) free.
[1590155066] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.7% (1909772 kB) free!
[1590155187] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.6% (4160084 kB) free.
[1590155609] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (389932 kB) free!
[1590155729] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (443048 kB) free!
[1590155764] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 20.1% (4936564 kB) free.
[1590155786] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (1217832 kB) free!
[1590155849] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (418716 kB) free!
[1590155906] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.7% (2877808 kB) free!
[1590156026] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.5% (3989140 kB) free.
[1590156364] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 14.2% (3489824 kB) free!
[1590156449] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36635732 kB) free.
[1590156484] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 13.1% (3207616 kB) free!
[1590156604] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 12.0% (2947832 kB) free!
[1590156626] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (918264 kB) free!
[1590156746] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.6% (2534640 kB) free!
[1590156866] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.3% (3955964 kB) free.
[1590157049] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (1795932 kB) free!
[1590157170] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (404468 kB) free!
[1590157204] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1664020 kB) free!
[1590157290] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384464 kB) free!
[1590157352] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.4% (4751296 kB) free!
[1590157466] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (366196 kB) free!
[1590157586] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.5% (2035692 kB) free!
[1590157706] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.2% (3764224 kB) free.
[1590158290] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.3% (3661472 kB) free!
[1590158306] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (179616 kB) free!
[1590158410] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.9% (3344264 kB) free!
[1590158426] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.1% (1806824 kB) free!
[1590158530] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.1% (2949148 kB) free!
[1590158546] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.6% (3992616 kB) free.
[1590158552] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1771052 kB) free!
[1590158754] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 346 processes
[1590158874] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1590158875] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/22 20:16:43 [error] 22135#0: *90073978 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1590158994] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1590158995] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590159130] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (1070824 kB) free!
[1590159147] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (304268 kB) free!
[1590159267] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.1% (2136892 kB) free!
[1590159387] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.1% (4079484 kB) free.
[1590159690] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20286968 kB) free!
[1590159987] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.6% (1887664 kB) free!
[1590160107] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.1% (3916420 kB) free.
[1590160290] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388376 kB) free!
[1590160707] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (166944 kB) free!
[1590160828] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.2% (1982616 kB) free!
[1590160948] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.5% (3664872 kB) free.
[1590161531] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1590161531] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.0% (16175220 kB) free.
[1590162148] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.8% (1920888 kB) free!
[1590162235] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.44, 4.51, 4.85
[1590162268] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3343808 kB) free.
[1590162690] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.0% (14536764 kB) free!
[1590163195] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/22 21:24:29 [error] 22135#0: *90086126 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/uuid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/uuid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/uuid"
[1590163290] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (405560 kB) free!
[1590163315] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1590163436] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.60, 5.24, 5.04
[1590163468] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (1496132 kB) free!
[1590163556] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.92, 4.81, 4.91
[1590163588] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.0% (2760772 kB) free!
[1590163708] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3513416 kB) free.
[1590163890] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.2% (25352288 kB) free!
[1590164156] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.61, 5.55, 5.22
[1590164185] SERVICE ALERT: sigqscluster1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1590164276] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.30, 5.24, 5.15
[1590164307] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 1844 processes
[1590164396] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.07, 5.61, 5.28
[1590164490] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (378512 kB) free!
[1590164768] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 20.0% (4906020 kB) free!
[1590164888] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 20.0% (4905760 kB) free!
[1590164908] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (970984 kB) free!
[1590164970] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1590165008] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 20.0% (4906284 kB) free!
[1590165028] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.1% (2462304 kB) free!
[1590165090] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30539056 kB) free.
[1590165148] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3293508 kB) free.
[1590165551] SERVICE ALERT: sigqscluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1590165672] SERVICE ALERT: sigqscluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1590165690] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (3958696 kB) free!
[1590165748] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (3162748 kB) free!
[1590165811] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393824 kB) free!
[1590165859] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 660 processes
[1590165868] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (1020184 kB) free!
[1590165931] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (412920 kB) free!
[1590165988] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.8% (2073332 kB) free!
[1590166531] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.6% (24542156 kB) free!
[1590167131] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399936 kB) free!
[1590167394] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 4 warn out of 347 processes
[1590167396] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.51, 4.91, 4.99
[1590167514] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1590167731] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16933748 kB) free!
[1590167788] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (797044 kB) free!
[1590168015] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1697 processes
[1590168331] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (424884 kB) free!
[1590168388] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.8% (2085352 kB) free!
[1590168931] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.6% (21947604 kB) free!
[1590168993] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.51, 1.75, 11.29
[1590169531] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1393356 kB) free!
[1590169615] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(2) < Channel error on connection <0.26110.620> (10.147.212.158:35808 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1590169735] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1590169797] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.01, 5.78, 5.34
[1590169917] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.79, 5.38, 5.25
[1590170037] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1590170037] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.49, 5.83, 5.42
[1590170060] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 775 processes
[1590170180] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 777 processes
[1590170188] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3308112 kB) free.
[1590170193] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 1.76, 1.72, 4.42
[1590170300] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 779 processes
[1590170415] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1706 processes
[1590171016] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1664 processes
[1590171131] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1590171237] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.16, 4.32, 4.59
[1590171388] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3248560 kB) free!
[1590171508] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (772160 kB) free!
[1590171628] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.0% (2268556 kB) free!
[1590171931] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14648756 kB) free!