[1566326243] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 91.9% (45307944 kB) free.
[1566326249] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.05, 1.32, 3.61
[1566326606] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (2042344 kB) free!
[1566326706] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193760 kB) free!
[1566326726] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (1270140 kB) free!
[1566326826] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190120 kB) free!
[1566326846] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193012 kB) free!
[1566326924] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 0.95, 1.60, 3.59
[1566326946] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181832 kB) free!
[1566328647] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 42.6% (10509568 kB) free.
[1566329508] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178452 kB) free!
[1566329568] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 31.7% (7771628 kB) free.
[1566330443] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (234092 kB) free!
[1566330463] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (390584 kB) free!
[1566330563] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (1841120 kB) free!
[1566330583] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1790784 kB) free!
[1566330683] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1566330683] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (1750440 kB) free!
[1566330703] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1793828 kB) free!
[1566334047] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (1251100 kB) free!
[1566334167] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190052 kB) free!
[1566334287] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (205600 kB) free!
[1566334605] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (234568 kB) free!
[1566334725] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (226592 kB) free!
[1566334845] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (215812 kB) free!
[1566335249] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.44, 6.83, 6.33
[1566335369] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.37, 4.70, 5.61
[1566335445] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.3% (45009716 kB) free.
[1566335489] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.35, 3.31, 5.00
[1566337146] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.8% (10240784 kB) free.
[1566337182] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 43.6% (10685612 kB) free.
[1566338483] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1566344489] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 70.72, 21.54, 7.88
[1566344609] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.76, 14.49, 6.95
[1566344729] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.62, 9.82, 6.15
[1566344998] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 177.00, 70.71, 27.74
[1566345118] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 198.72, 114.18, 48.88
[1566345238] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 201.27, 143.13, 67.41
[1566345329] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.35, 1.67, 3.40
[1566345929] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 50.35, 11.99, 5.73
[1566346049] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 159.69, 66.51, 26.47
[1566346169] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 174.36, 102.90, 44.66
[1566346782] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 147.55, 96.79, 42.95
[1566346902] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 148.36, 113.87, 55.73
[1566347022] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 28.19, 83.16, 51.85
[1566347038] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.56, 3.40, 17.78
[1566347484] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 82.0% (40398392 kB) free.
[1566347969] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1566347969] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.56, 1.24, 13.97
[1566348222] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.41, 2.85, 15.14
[1566348687] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 70.8% (17477208 kB) free.
[1566348793] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 570 processes
[1566348838] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.55, 2.60, 4.81
[1566348853] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 563 processes
[1566349169] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.58, 0.71, 4.44
[1566349303] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 65.3% (86041396 kB) free.
[1566349422] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.24, 0.37, 4.50
[1566350353] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 576 processes
[1566350413] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 587 processes
[1566350473] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 596 processes
[1566350813] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1566350933] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1566350946] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (1416936 kB) free!
[1566350986] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171428 kB) free!
[1566351053] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1566351066] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (1401224 kB) free!
[1566351087] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (1511140 kB) free!
[1566351113] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174732 kB) free!
[1566351186] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (1385952 kB) free!
[1566351207] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (1497144 kB) free!
[1566351233] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185800 kB) free!
[1566351327] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (1483020 kB) free!
[1566351877] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 375 processes
[1566351997] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 381 processes
[1566354569] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.33, 6.28, 6.16
[1566354689] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.63, 4.40, 5.48
[1566354809] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.87, 3.17, 4.90
[1566354866] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1566354986] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1566360209] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1566363809] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.40, 17.15, 8.92
[1566363929] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.64, 11.65, 7.90
[1566364049] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.15, 8.13, 7.06
[1566364156] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.4% (14523672 kB) free.
[1566364213] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.8% (21285104 kB) free.
[1566364277] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.7% (40772556 kB) free.
[1566364319] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.0% (22297004 kB) free.
[1566364374] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 77.0% (19008512 kB) free.
[1566364649] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.45, 1.48, 3.93
[1566364721] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1566364841] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1566364962] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1566365320] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1566365394] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1566366594] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1566367303] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (3367992 kB) free!
[1566367423] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (2516956 kB) free!
[1566367543] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (2580188 kB) free!
[1566367956] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 65 users currently logged in
[1566367983] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198352 kB) free!
[1566368016] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 65 users currently logged in
[1566368076] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 65 users currently logged in
[1566368103] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (162860 kB) free!
[1566368136] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 65 users currently logged in
[1566368144] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 62.3% (82111832 kB) free.
[1566368223] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192436 kB) free!
[1566368362] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.2% (2640664 kB) free!
[1566368482] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.8% (2073668 kB) free!
[1566368602] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (1552944 kB) free!
[1566369936] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 67 users currently logged in
[1566370120] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1566370124] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.81, 6.40, 5.62
[1566370244] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.12, 5.66, 5.46
[1566370365] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.78, 6.57, 5.81
[1566370438] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 24.29, 7.37, 4.38
[1566370558] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 234.37, 93.32, 36.24
[1566370678] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 262.14, 150.64, 64.17
[1566370720] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1566371022] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 136.56, 73.09, 30.25
[1566371142] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 139.69, 95.31, 43.58
[1566371262] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 140.43, 110.29, 55.35
[1566372462] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.52, 4.44, 19.43
[1566373079] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.89, 3.10, 14.65
[1566373649] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.79, 9.27, 7.28
[1566373769] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.73, 6.45, 6.49
[1566373889] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.71, 4.57, 5.79
[1566374262] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.80, 0.90, 3.73
[1566374489] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1566374489] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.83, 1.35, 3.44
[1566374879] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.12, 3.00, 4.57
[1566375081] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 13.2% (6499784 kB) free!
[1566375201] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 11.3% (5564172 kB) free!
[1566375321] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4692984 kB) free!
[1566376135] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 358 processes
[1566376236] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 65 users currently logged in
[1566376255] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 358 processes
[1566376375] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1566376536] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1566377522] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 205 processes
[1566377637] SERVICE ALERT: sig4-vm2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 286 processes
[1566377642] SERVICE ALERT: sig5-vm2;CPU_Procs;OK;HARD;1;CPU OK: 204 processes
[1566377757] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 285 processes
[1566379923] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.5% (3066224 kB) free!
[1566380043] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.7% (3100924 kB) free!
[1566380163] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1783348 kB) free!
[1566380994] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1566382246] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9316340 kB) free!
[1566382289] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 33.46, 22.02, 9.40
[1566382366] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.6% (9147376 kB) free!
[1566382410] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.23, 15.04, 8.38
[1566382486] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.2% (8963236 kB) free!
[1566382530] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.96, 10.45, 7.50
[1566383130] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.50, 2.38, 4.41
[1566383664] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 917 processes
[1566383784] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 374 processes
[1566383879] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.83, 13.15, 11.54
[1566383999] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.93, 9.72, 10.48
[1566384119] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.74, 7.41, 9.54
[1566384122] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.92, 5.44, 5.30
[1566384242] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.78, 5.26, 5.25
[1566384362] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.85, 5.13, 5.20
[1566385319] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.88, 3.04, 4.75
[1566385486] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4164344 kB) free!
[1566385530] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 72.08, 30.56, 12.24
[1566385562] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.65, 4.82, 4.95
[1566385650] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.50, 20.75, 10.87
[1566385770] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.66, 13.96, 9.57
[1566386162] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.91, 5.15, 5.10
[1566386282] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.22, 5.27, 5.15
[1566386402] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.17, 5.17, 5.12
[1566386970] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.30, 1.36, 3.41
[1566387002] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1566387002] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.49, 4.74, 4.94
[1566387602] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.07, 5.29, 5.15
[1566387722] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.34, 5.34, 5.19
[1566387775] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1566387842] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.65, 5.15, 5.15
[1566387895] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1566388770] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 93.28, 27.86, 10.76
[1566388890] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 13.69, 19.07, 9.61
[1566388965] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.15, 4.26, 4.86
[1566389010] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.73, 13.07, 8.55
[1566389610] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 281.60, 164.12, 72.16
[1566389695] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1566389815] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1566390162] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1566390202] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 84.6% (13754496 kB) free.
[1566390286] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 96.0% (47300220 kB) free.
[1566390321] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.8% (40327788 kB) free.
[1566390363] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 88.4% (21666528 kB) free.
[1566390465] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.9% (22181428 kB) free.
[1566390719] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 170.56, 52.49, 20.08
[1566390762] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1566390839] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 244.38, 119.83, 48.69
[1566390882] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1566390959] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 255.39, 165.30, 73.98
[1566391002] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1566391965] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.64, 5.52, 5.11
[1566392002] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (183232 kB) free!
[1566392010] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.50, 1.84, 14.29
[1566392085] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.98, 4.79, 4.89
[1566392123] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (169056 kB) free!
[1566392215] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 363 processes
[1566392243] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (187088 kB) free!
[1566392335] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1566392685] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.26, 6.19, 5.47
[1566392763] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.9% (3409972 kB) free!
[1566392806] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.87, 5.76, 5.41
[1566392883] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.5% (22673228 kB) free.
[1566392912] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.2% (14341740 kB) free.
[1566392920] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1566392926] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.94, 5.56, 5.38
[1566393336] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 64 users currently logged in
[1566393359] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1566393359] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.84, 3.29, 15.15
[1566393402] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1566393810] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.95, 1.18, 4.49
[1566393936] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 58 users currently logged in
[1566394253] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1566394684] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (1516476 kB) free!
[1566394721] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.1% (2776396 kB) free!
[1566394804] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181604 kB) free!
[1566394841] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.7% (2059556 kB) free!
[1566394924] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169748 kB) free!
[1566394961] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1566394961] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1314640 kB) free!
[1566395159] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.85, 3.02, 4.77
[1566395610] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.93, 12.28, 8.87
[1566395730] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.13, 8.57, 7.92
[1566395850] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.97, 6.04, 7.07
[1566396321] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.5% (7661768 kB) free!
[1566396441] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 13.8% (6819252 kB) free!
[1566396450] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.99, 1.78, 4.24
[1566396561] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 12.3% (6056344 kB) free!
[1566397161] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (1845156 kB) free!
[1566397225] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4810412 kB) free!
[1566397345] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (4724164 kB) free!
[1566397465] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.1% (4670364 kB) free!
[1566398274] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4686696 kB) free!
[1566398394] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4665260 kB) free!
[1566398514] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (4638600 kB) free!
[1566398850] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.62, 12.19, 8.28
[1566398970] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.19, 8.42, 7.37
[1566399090] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.54, 6.18, 6.68
[1566399242] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1566399691] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.40, 1.71, 3.97
[1566400486] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.2% (8981680 kB) free!
[1566400606] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.8% (8765304 kB) free!
[1566400726] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.4% (8574424 kB) free!
[1566401926] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.77, 4.08, 4.92
[1566402091] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.28, 12.75, 8.29
[1566402211] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.13, 8.84, 7.40
[1566402331] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.66, 6.11, 6.57
[1566402865] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2311880 kB) free!
[1566402931] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.22, 1.71, 3.91
[1566403961] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1566404326] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4754360 kB) free!
[1566404334] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.9% (4876156 kB) free!
[1566404454] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.8% (4858236 kB) free!
[1566404574] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.8% (4844952 kB) free!
[1566405331] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.90, 15.59, 10.02
[1566405360] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1566405451] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.07, 10.65, 8.88
[1566405571] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.50, 7.32, 7.88
[1566406171] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.97, 1.82, 4.56
[1566406885] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9802164 kB) free!
[1566407005] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.8% (9736744 kB) free!
[1566407125] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.6% (9673176 kB) free!
[1566407760] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 81.68, 28.71, 11.84
[1566407880] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 13.62, 20.20, 10.76
[1566408000] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.24, 14.44, 9.80
[1566408571] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.29, 18.83, 10.66
[1566408692] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.68, 12.88, 9.47
[1566408812] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.20, 8.94, 8.43
[1566409200] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.63, 2.86, 4.66
[1566409412] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.80, 2.04, 4.88
[1566409914] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2446356 kB) free!