[1563561061] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (16633184 kB) free.
[1563561178] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.22, 0.69, 3.43
[1563561412] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563561455] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7849320 kB) free!
[1563561510] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190508 kB) free!
[1563561533] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563561575] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11722088 kB) free.
[1563561622] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178732 kB) free!
[1563561630] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189508 kB) free!
[1563561661] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (12576680 kB) free!
[1563561680] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186492 kB) free!
[1563561742] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187748 kB) free!
[1563561750] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186376 kB) free!
[1563561781] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (8732652 kB) free!
[1563561783] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.65, 9.98, 7.33
[1563561800] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181436 kB) free!
[1563561862] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176596 kB) free!
[1563561901] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (4754172 kB) free!
[1563561903] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.25, 6.75, 6.48
[1563561920] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169008 kB) free!
[1563561999] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 20.3% (4981144 kB) free.
[1563562023] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.19, 4.60, 5.72
[1563562132] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563562175] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8029864 kB) free!
[1563562252] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563562295] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11713256 kB) free.
[1563562372] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563562501] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (15619884 kB) free.
[1563562623] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.11, 0.81, 3.11
[1563562847] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 365 processes
[1563562895] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8158836 kB) free!
[1563562967] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 364 processes
[1563562972] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563563015] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11722876 kB) free.
[1563563092] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563563101] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (11466928 kB) free!
[1563563221] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (7809072 kB) free!
[1563563341] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (17335604 kB) free.
[1563563615] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8288756 kB) free!
[1563563692] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563563735] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11713672 kB) free.
[1563563812] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563564335] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8407516 kB) free!
[1563564413] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563564455] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.5% (6655180 kB) free!
[1563564533] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563564541] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (9798260 kB) free!
[1563564575] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11720244 kB) free.
[1563564661] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (6003728 kB) free!
[1563564721] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4718904 kB) free!
[1563564781] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (16994832 kB) free.
[1563564781] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4722764 kB) free!
[1563564842] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 19.3% (4725352 kB) free!
[1563564902] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 19.3% (4724352 kB) free!
[1563564931] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1563565133] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563565253] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563565381] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (12323384 kB) free!
[1563565501] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.6% (8261548 kB) free!
[1563565621] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (17362860 kB) free.
[1563565775] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8431536 kB) free!
[1563565853] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563565895] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.6% (6721012 kB) free!
[1563565973] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563566015] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11720624 kB) free.
[1563566573] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563566615] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7046940 kB) free!
[1563566693] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563566735] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11711852 kB) free.
[1563566821] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (9921152 kB) free!
[1563566941] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (6033756 kB) free!
[1563567061] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (16712188 kB) free.
[1563567294] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563567335] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7299408 kB) free!
[1563567414] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563567455] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11718052 kB) free.
[1563567661] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (12092396 kB) free!
[1563567781] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.9% (8460992 kB) free!
[1563567901] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (17379904 kB) free.
[1563568014] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563568055] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7654436 kB) free!
[1563568134] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563568175] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11718264 kB) free.
[1563568734] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563568775] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (8000456 kB) free!
[1563568854] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563568895] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11713628 kB) free.
[1563568974] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563569101] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (9800652 kB) free!
[1563569221] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (5950160 kB) free!
[1563569341] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (16664400 kB) free.
[1563569496] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7987992 kB) free!
[1563569574] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563569616] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11718928 kB) free.
[1563569694] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563569941] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (11992708 kB) free!
[1563570061] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (7987104 kB) free!
[1563570181] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (17406192 kB) free.
[1563570216] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7996096 kB) free!
[1563570295] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563570336] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11712920 kB) free.
[1563570415] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563570936] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8156380 kB) free!
[1563571015] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563571056] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11717228 kB) free.
[1563571135] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563571381] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (8980776 kB) free!
[1563571502] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (5358324 kB) free!
[1563571622] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (16123128 kB) free.
[1563571656] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8377272 kB) free!
[1563571735] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563571776] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.6% (6721568 kB) free!
[1563571855] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563571896] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11708376 kB) free.
[1563572222] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (11950064 kB) free!
[1563572342] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.5% (8206816 kB) free!
[1563572455] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563572462] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (17407060 kB) free.
[1563572496] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (6994652 kB) free!
[1563572575] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563572578] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563572616] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11716732 kB) free.
[1563573176] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563573216] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7389460 kB) free!
[1563573296] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563573337] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11710068 kB) free.
[1563573662] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (9633984 kB) free!
[1563573782] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (5903076 kB) free!
[1563573896] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563573902] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (16428332 kB) free.
[1563573936] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7699660 kB) free!
[1563574016] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563574056] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11716284 kB) free.
[1563574502] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (11818660 kB) free!
[1563574616] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563574622] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (7980268 kB) free!
[1563574656] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7843268 kB) free!
[1563574736] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563574742] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (17319280 kB) free.
[1563574776] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11717252 kB) free.
[1563575336] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563575376] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7893808 kB) free!
[1563575456] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563575496] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11706620 kB) free.
[1563575577] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563575942] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (8684448 kB) free!
[1563576062] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (17317748 kB) free.
[1563576096] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7996800 kB) free!
[1563576176] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563576216] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11717308 kB) free.
[1563576296] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563576662] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563576816] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8142288 kB) free!
[1563576896] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563576936] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11706320 kB) free.
[1563577016] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563577263] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (8885996 kB) free!
[1563577383] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (17326956 kB) free.
[1563577536] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8286964 kB) free!
[1563577616] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563577656] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11680904 kB) free.
[1563577736] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563578257] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8397796 kB) free!
[1563578336] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563578377] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.6% (6717040 kB) free!
[1563578457] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563578497] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11707092 kB) free.
[1563578583] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (8873804 kB) free!
[1563578703] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (17327852 kB) free.
[1563579056] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563579097] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6808260 kB) free!
[1563579176] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1563579217] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11716324 kB) free.
[1563579776] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563579817] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7025560 kB) free!
[1563579896] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563579903] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (9656696 kB) free!
[1563579937] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11705892 kB) free.
[1563580023] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (6049924 kB) free!
[1563580143] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (16348588 kB) free.
[1563580496] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563580537] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7204952 kB) free!
[1563580616] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563580657] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11713296 kB) free.
[1563580743] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (11391896 kB) free!
[1563580863] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.0% (7895748 kB) free!
[1563580983] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (5013424 kB) free!
[1563581583] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (17261992 kB) free.
[1563581817] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563581937] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563582057] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563582458] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8837832 kB) free!
[1563582578] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.2% (7469260 kB) free!
[1563582657] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1563582698] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11674576 kB) free.
[1563582999] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (4523348 kB) free!
[1563583119] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.7% (4343468 kB) free!
[1563583239] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (4246136 kB) free!
[1563583257] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563583377] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563583497] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563583898] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8709548 kB) free!
[1563583983] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (12004096 kB) free!
[1563584018] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7125400 kB) free!
[1563584097] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1563584097] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563584103] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.1% (8636360 kB) free!
[1563584138] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11368556 kB) free.
[1563584223] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (5233688 kB) free!
[1563584697] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563584818] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563584823] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1563584823] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (16725260 kB) free.
[1563584938] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563585338] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8623200 kB) free!
[1563585458] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7145652 kB) free!
[1563585538] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563585578] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11641660 kB) free.
[1563586138] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1563586178] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7390744 kB) free!
[1563586258] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563586298] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11630368 kB) free.
[1563586623] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (12539292 kB) free!
[1563586743] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (8982544 kB) free!
[1563586858] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563586863] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (5479944 kB) free!
[1563586899] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7531676 kB) free!
[1563586929] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (3249492 kB) free!
[1563586978] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563587019] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11639260 kB) free.
[1563587049] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (3248792 kB) free!
[1563587169] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 20.0% (3244504 kB) free!
[1563587463] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (15796596 kB) free.
[1563587579] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563587619] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7680316 kB) free!
[1563587699] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563587739] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11629020 kB) free.
[1563588040] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 20.1% (4924652 kB) free.
[1563588063] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (11402328 kB) free!
[1563588183] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.0% (7898932 kB) free!
[1563588299] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563588303] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (16655312 kB) free.
[1563588339] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7808576 kB) free!
[1563588419] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563588459] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11639308 kB) free.
[1563588539] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563588969] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1414652 kB) free!
[1563589059] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6720516 kB) free!
[1563589179] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11636284 kB) free.
[1563589504] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (11669388 kB) free!
[1563589530] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 846.01, 404.47, 159.29
[1563589624] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.4% (8183392 kB) free!
[1563589650] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 883.66, 564.99, 247.84
[1563589739] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1563589744] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (4712288 kB) free!
[1563589770] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 457.07, 561.38, 288.31
[1563589779] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6826932 kB) free!
[1563589859] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563589899] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11135352 kB) free.
[1563589967] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 362 processes
[1563590087] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1563590344] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (16141316 kB) free.
[1563590944] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (12726392 kB) free!
[1563591064] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (9390204 kB) free!
[1563591184] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (6012844 kB) free!
[1563591659] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563591699] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7664392 kB) free!
[1563591779] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563591784] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (16490376 kB) free.
[1563591819] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11637036 kB) free.
[1563592379] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563592419] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7780264 kB) free!
[1563592499] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563592539] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11628084 kB) free.
[1563592619] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563592770] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.50, 2.77, 14.18
[1563593139] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7836056 kB) free!
[1563593219] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563593259] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11634128 kB) free.
[1563593340] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1563593440] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4841376 kB) free!
[1563593560] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (4793684 kB) free!
[1563593680] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (4695660 kB) free!
[1563593859] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8011452 kB) free!
[1563593939] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563593979] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11637564 kB) free.
[1563594059] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563594184] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (10631800 kB) free!
[1563594304] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.2% (6708656 kB) free!
[1563594424] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (2911400 kB) free!
[1563594571] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.88, 2.88, 4.47
[1563594579] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8253564 kB) free!
[1563594659] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563594699] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.2% (6509644 kB) free!
[1563594779] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563594819] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11633136 kB) free.
[1563595024] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (15785748 kB) free.
[1563595379] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563595499] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563595624] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (11978980 kB) free!
[1563595684] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.71, 8.26, 6.02
[1563595745] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.4% (7504968 kB) free!
[1563595804] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.44, 8.98, 6.59
[1563595865] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (3536924 kB) free!
[1563595924] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.92, 8.11, 6.55
[1563596019] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8114440 kB) free!
[1563596099] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563596102] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 38.7% (9487500 kB) free.
[1563596139] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11629740 kB) free.
[1563596220] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563596371] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (264440 kB) free!
[1563596492] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (243136 kB) free!
[1563596612] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (266044 kB) free!
[1563596680] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (1352876 kB) free!
[1563596739] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8204716 kB) free!
[1563596819] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563596859] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.2% (6492188 kB) free!
[1563596939] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563596979] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11635628 kB) free.
[1563597065] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (14917836 kB) free.
[1563597262] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 47.3% (11590272 kB) free.
[1563597281] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1563597281] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 12.7% (3102944 kB) free!
[1563597539] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563597659] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563597665] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (12325520 kB) free!
[1563597785] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (7805104 kB) free!
[1563597905] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (3447488 kB) free!
[1563597920] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 55.7% (13652572 kB) free.
[1563598259] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563598379] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563598499] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563598832] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.4% (2832452 kB) free!
[1563598952] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.9% (2580904 kB) free!
[1563599072] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 14.1% (2285120 kB) free!
[1563599105] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (15640556 kB) free.
[1563599673] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169812 kB) free!
[1563599979] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8009388 kB) free!
[1563600099] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11621560 kB) free.
[1563600299] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1563600305] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (10765808 kB) free!
[1563600425] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (6019244 kB) free!
[1563600545] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (1588388 kB) free!
[1563600700] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7762552 kB) free!
[1563600778] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 13961.19, 6757.63, 2646.98
[1563600820] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11631400 kB) free.
[1563600865] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.70, 5.93, 5.42
[1563600898] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 14535.80, 9361.00, 4099.41
[1563600985] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.42, 5.83, 5.45
[1563601018] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14612.17, 11102.48, 5375.54
[1563601105] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.98, 5.23, 5.28
[1563601420] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7595480 kB) free!
[1563601540] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11619112 kB) free.
[1563601706] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 3.34, 3.49, 4.31
[1563601745] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (16154392 kB) free.
[1563602073] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 25.6% (4152948 kB) free.
[1563602099] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563602140] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7363228 kB) free!
[1563602219] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563602260] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11627900 kB) free.
[1563602362] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 15.6% (2542724 kB) free!
[1563602482] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 14.1% (2294220 kB) free!
[1563602602] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 12.8% (2079740 kB) free!
[1563602673] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 15.8% (2573964 kB) free!
[1563602793] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 13.6% (2215112 kB) free!
[1563602819] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563602860] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7020416 kB) free!
[1563602914] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1563602914] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 11.5% (1876324 kB) free!
[1563602939] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563602946] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (10905916 kB) free!
[1563602980] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11623800 kB) free.
[1563603066] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (6159992 kB) free!
[1563603186] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (1802116 kB) free!
[1563603202] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (921132 kB) free!
[1563603446] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 11.7% (863160 kB) free!
[1563603514] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (298812 kB) free!
[1563603540] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563603566] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (575596 kB) free!
[1563603580] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6618304 kB) free!
[1563603660] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1563603686] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (321216 kB) free!
[1563603700] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11619976 kB) free.
[1563603716] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.4% (22556188 kB) free.
[1563603732] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.6% (24740104 kB) free.
[1563603758] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 84.1% (13670732 kB) free.
[1563603786] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 84.5% (13735632 kB) free.
[1563603804] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 68.1% (16691300 kB) free.
[1563603897] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.6% (39700212 kB) free.
[1563603909] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 91.4% (6724064 kB) free.
[1563603928] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 76.2% (12385128 kB) free.
[1563604260] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563604300] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6366880 kB) free!
[1563604321] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (830692 kB) free!
[1563604380] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563604420] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10793424 kB) free.
[1563604441] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194920 kB) free!
[1563604500] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563604561] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177220 kB) free!
[1563604623] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.83, 21.68, 17.45
[1563604743] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.62, 14.71, 15.41
[1563604863] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.49, 9.97, 13.58
[1563604942] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (4518264 kB) free!
[1563605021] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9317268 kB) free!
[1563605062] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (267916 kB) free!
[1563605141] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7934932 kB) free!
[1563605182] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (280128 kB) free!
[1563605261] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11616432 kB) free.
[1563605306] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.87, 5.59, 5.06
[1563605426] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.03, 5.36, 5.04
[1563605546] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.09, 5.32, 5.06
[1563605663] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4817892 kB) free!
[1563605783] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (4680140 kB) free!
[1563605861] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7796820 kB) free!
[1563605903] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (4579012 kB) free!
[1563605981] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.6% (11614044 kB) free.
[1563606063] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.54, 0.63, 4.07
[1563606192] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.2% (2472092 kB) free!
[1563606312] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.6% (1890380 kB) free!
[1563606433] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1404808 kB) free!
[1563606746] SERVICE FLAPPING ALERT: smscapp1;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1563606746] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 3.83, 4.33, 4.78
[1563606887] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 361 processes
[1563606982] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (16436092 kB) free.
[1563607007] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1563607367] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 13.4% (2184120 kB) free!
[1563607487] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 11.3% (1833948 kB) free!
[1563607607] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (1443044 kB) free!
[1563607618] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.04, 5.49, 15.97
[1563608303] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2052856 kB) free!
[1563608782] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (11793272 kB) free!
[1563608902] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.1% (7279760 kB) free!
[1563609022] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (2822396 kB) free!
[1563609208] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1563609362] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 40.4% (9967448 kB) free.
[1563609921] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (4622552 kB) free!
[1563610018] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.92, 1.20, 3.43
[1563610041] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.6% (4558984 kB) free!
[1563610161] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (4498088 kB) free!
[1563610181] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7258932 kB) free!
[1563610222] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.8% (11076128 kB) free!
[1563610302] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11596440 kB) free.
[1563610822] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (17590252 kB) free.
[1563611118] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 16.5% (1212816 kB) free!
[1563611238] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 14.3% (1052112 kB) free!
[1563611358] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 12.4% (910528 kB) free!
[1563611697] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9066476 kB) free!
[1563611818] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.2% (8969620 kB) free!
[1563611938] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.0% (8886584 kB) free!
[1563611958] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1563611958] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (119632 kB) free!
[1563612622] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (11837168 kB) free!
[1563612742] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.2% (7356924 kB) free!
[1563612862] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (3297112 kB) free!
[1563612900] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1563613738] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3906180 kB) free!
[1563613761] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2405532 kB) free!
[1563614062] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (18001548 kB) free.
[1563614501] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8766792 kB) free!
[1563614621] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.8% (7269060 kB) free!
[1563614700] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563614741] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11576636 kB) free.
[1563614762] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.4% (4285228 kB) free!
[1563614808] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1563614882] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.2% (4248464 kB) free!
[1563615002] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (336912 kB) free!
[1563615300] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563615420] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563615540] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563615558] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.6% (6811044 kB) free.
[1563615862] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (12519316 kB) free!
[1563615942] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8979200 kB) free!
[1563615982] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (7800100 kB) free!
[1563616062] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.4% (7565696 kB) free!
[1563616102] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (3449076 kB) free!
[1563616140] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1563616140] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563616182] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.5% (11572400 kB) free.
[1563616618] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.14, 5.97, 5.19
[1563616739] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.07, 5.96, 5.28
[1563616741] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563616859] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.59, 5.79, 5.30
[1563616861] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563616981] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563617302] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (18378344 kB) free.
[1563617382] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8484208 kB) free!
[1563617502] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.9% (6868068 kB) free!
[1563617546] SERVICE FLAPPING ALERT: smscapp1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563617581] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563617622] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.8% (11241264 kB) free.
[1563617741] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.2% (2799116 kB) free!
[1563617861] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 15.6% (2536648 kB) free!
[1563617959] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 44.2% (10910028 kB) free.
[1563617981] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.1% (2293656 kB) free!
[1563617999] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 85.8% (21029004 kB) free.
[1563618181] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563618216] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.7% (39747424 kB) free.
[1563618222] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9613964 kB) free!
[1563618257] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.4% (13873984 kB) free.
[1563618285] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 81.5% (19975456 kB) free.
[1563618301] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563618329] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.9% (14778024 kB) free.
[1563618342] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8168996 kB) free!
[1563618407] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 360 processes
[1563618421] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563618462] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.8% (6311940 kB) free!
[1563618527] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1563618569] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188852 kB) free!
[1563618581] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (849448 kB) free!
[1563618689] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177272 kB) free!
[1563618809] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1563618809] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193428 kB) free!
[1563619062] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.9% (10796060 kB) free.
[1563619102] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (12518904 kB) free!
[1563619222] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.1% (7324716 kB) free!
[1563619342] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (2684168 kB) free!
[1563619662] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9258088 kB) free!
[1563619783] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7814988 kB) free!
[1563619903] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11506624 kB) free.
[1563619946] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.04, 5.85, 5.24
[1563620066] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.48, 6.15, 5.43
[1563620186] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.64, 6.04, 5.49
[1563620358] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 10.9% (801872 kB) free!
[1563620478] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (371508 kB) free!
[1563620598] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (127824 kB) free!
[1563621103] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9234532 kB) free!
[1563621142] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (8746768 kB) free!
[1563621223] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7911196 kB) free!
[1563621257] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.0% (2932044 kB) free!
[1563621338] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (1331932 kB) free!
[1563621343] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11486624 kB) free.
[1563621377] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.4% (2508796 kB) free!
[1563621458] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (889488 kB) free!
[1563621497] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 12.9% (2090136 kB) free!
[1563621578] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (456260 kB) free!
[1563621659] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.00, 2.22, 4.07
[1563621986] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 4.27, 4.23, 4.99
[1563622097] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (173120 kB) free!
[1563622543] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8930184 kB) free!
[1563622663] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7537292 kB) free!
[1563622781] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 10.6% (1723788 kB) free!
[1563622783] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11478540 kB) free.
[1563623221] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563623381] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1563623381] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (293752 kB) free!
[1563623619] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (224872 kB) free!
[1563623739] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (246560 kB) free!
[1563623821] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1563623859] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (214720 kB) free!
[1563623941] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563623983] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8649720 kB) free!
[1563624062] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563624104] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.5% (7157520 kB) free!
[1563624224] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11452024 kB) free.
[1563624662] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563624824] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9812576 kB) free!
[1563624944] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.3% (8508072 kB) free!
[1563625064] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.6% (7215880 kB) free!
[1563625262] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563625382] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563625502] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563625586] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.32, 5.28, 5.06
[1563625664] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1563625664] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11452992 kB) free.
[1563625706] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.66, 5.39, 5.12
[1563625826] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.75, 5.16, 5.07
[1563625942] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (4711760 kB) free!
[1563625998] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.0% (6838608 kB) free.
[1563626382] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 72.0% (11707872 kB) free.
[1563626426] SERVICE FLAPPING ALERT: smscapp1;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1563626426] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 3.08, 3.73, 4.44
[1563626702] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563627302] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1563627423] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563627527] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 361 processes
[1563627647] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1563628022] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563628142] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563628262] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563628325] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.87, 3.69, 4.55
[1563628342] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (16633876 kB) free.
[1563629009] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1563629462] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563629609] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 38.8% (9563780 kB) free.
[1563629864] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6505680 kB) free!
[1563629984] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11445812 kB) free.
[1563630123] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 319 processes
[1563630142] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (12808484 kB) free!
[1563630243] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1563630262] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.4% (8832976 kB) free!
[1563630382] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (4793000 kB) free!
[1563630584] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6806800 kB) free!
[1563630704] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11444764 kB) free.
[1563631262] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563631304] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7012340 kB) free!
[1563631383] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563631424] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11443188 kB) free.
[1563631982] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1563632024] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.1% (6959380 kB) free!
[1563632102] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563632144] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11435024 kB) free.
[1563632182] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (15924336 kB) free.
[1563632599] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1563632702] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563632744] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.5% (6652516 kB) free!
[1563632822] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563632864] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11440728 kB) free.
[1563633422] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563633464] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6335260 kB) free!
[1563633542] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563633585] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11441848 kB) free.
[1563633809] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193052 kB) free!
[1563633929] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196616 kB) free!
[1563633982] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (12375512 kB) free!
[1563634049] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177820 kB) free!
[1563634102] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (8385316 kB) free!
[1563634142] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563634222] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (4656956 kB) free!
[1563634263] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1563634784] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7810584 kB) free!
[1563634862] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563634904] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11439440 kB) free.
[1563634982] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563635382] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1563635504] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7655932 kB) free!
[1563635582] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563635624] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11437636 kB) free.
[1563635702] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563636022] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.4% (12072368 kB) free!
[1563636224] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7689972 kB) free!
[1563636302] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1563636344] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11432100 kB) free.
[1563636422] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563636622] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (15372088 kB) free.
[1563636944] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7779588 kB) free!
[1563637022] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563637064] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11436732 kB) free.
[1563637143] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1563637664] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7801128 kB) free!
[1563637742] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563637784] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11438692 kB) free.
[1563637826] SERVICE FLAPPING ALERT: smscapp1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563637862] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563638385] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7613332 kB) free!
[1563638422] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (12172588 kB) free!
[1563638462] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563638505] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563638505] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11436572 kB) free.
[1563638542] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (8985116 kB) free!
[1563638582] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1563638663] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (5786840 kB) free!
[1563639105] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7332820 kB) free!
[1563639225] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11439952 kB) free.
[1563639782] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563639825] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7039944 kB) free!
[1563639902] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1563639945] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11425040 kB) free.
[1563640503] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1563640545] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7008248 kB) free!
[1563640623] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563640665] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11435836 kB) free.
[1563641223] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563641265] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7042052 kB) free!
[1563641343] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563641385] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11434236 kB) free.
[1563641943] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563641985] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.6% (6723624 kB) free!
[1563642063] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1563642063] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1563642105] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11435580 kB) free.
[1563642663] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563642706] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6469384 kB) free!
[1563642783] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563642826] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11435092 kB) free.
[1563642863] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (15162256 kB) free.
[1563642871] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.4% (21326900 kB) free.
[1563642942] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.6% (39745152 kB) free.
[1563643015] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.2% (13854756 kB) free.
[1563643078] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.2% (14823164 kB) free.
[1563643384] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563643463] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (12206668 kB) free!
[1563643504] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1563643583] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (8419744 kB) free!
[1563643703] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (4764224 kB) free!
[1563644026] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7745588 kB) free!
[1563644104] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1563644146] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11420856 kB) free.
[1563644224] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563644678] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (1225776 kB) free!
[1563644746] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7598000 kB) free!
[1563644798] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182520 kB) free!
[1563644824] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1563644866] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11428144 kB) free.
[1563644903] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (15268184 kB) free.
[1563644918] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192268 kB) free!
[1563644944] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1563645466] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7399496 kB) free!
[1563645586] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11425936 kB) free.
[1563646144] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1563646186] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.2% (7021624 kB) free!
[1563646247] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 360 processes
[1563646265] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1563646306] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11425112 kB) free.
[1563646367] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1563646704] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (10499184 kB) free!
[1563646824] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.7% (7028852 kB) free!
[1563646864] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1563646906] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7084284 kB) free!
[1563646944] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (3434636 kB) free!
[1563646984] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1563647027] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11424580 kB) free.