[1550514869] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 1 warn out of 304 processes
[1550514900] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.59, 4.61, 4.94
[1550514989] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 301 processes
[1550515109] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 302 processes
[1550515421] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 36.09, 9.90, 4.79
[1550515500] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.91, 5.67, 5.24
[1550515541] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.71, 8.77, 5.08
[1550515620] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.78, 6.70, 5.66
[1550515662] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.12, 6.28, 4.62
[1550515740] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.43, 6.95, 5.88
[1550515925] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1550515925] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 35.4% (17464596 kB) free.
[1550516525] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (211696 kB) free!
[1550516645] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (249360 kB) free!
[1550516765] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (227920 kB) free!
[1550518565] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 20.0% (9866844 kB) free!
[1550520365] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10945912 kB) free.
[1550521109] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1550522940] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.88, 4.88, 4.97
[1550523366] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8794976 kB) free!
[1550523486] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.5% (13532520 kB) free.
[1550523509] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 300 processes
[1550524086] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8442208 kB) free!
[1550524206] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.5% (13527992 kB) free.
[1550524709] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 296 processes
[1550525406] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9504824 kB) free!
[1550525526] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.4% (13523576 kB) free.
[1550526126] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9318840 kB) free!
[1550526246] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.5% (13532460 kB) free.
[1550526846] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8952292 kB) free!
[1550526966] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13415116 kB) free.
[1550527109] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 300 processes
[1550527566] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8767220 kB) free!
[1550527687] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13427900 kB) free.
[1550528286] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8473768 kB) free!
[1550528406] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13413716 kB) free.
[1550528909] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1550528909] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 298 processes
[1550529006] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8115228 kB) free!
[1550529126] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13424788 kB) free.
[1550530326] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9233792 kB) free!
[1550530446] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13425028 kB) free.
[1550531046] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8987180 kB) free!
[1550531166] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13412736 kB) free.
[1550531309] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550531766] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8737864 kB) free!
[1550531886] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1550531886] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13424356 kB) free.
[1550532487] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8436916 kB) free!
[1550532607] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13412616 kB) free.
[1550533109] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550533207] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8684848 kB) free!
[1550533229] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550533327] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13423120 kB) free.
[1550533349] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1550533927] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8358148 kB) free!
[1550534047] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13413216 kB) free.
[1550534549] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550534669] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550535247] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9382720 kB) free!
[1550535269] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550535367] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13412044 kB) free.
[1550535389] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550535967] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9184560 kB) free!
[1550535990] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550536087] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13419144 kB) free.
[1550536110] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1550536688] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8977788 kB) free!
[1550536710] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550536808] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13407784 kB) free.
[1550536830] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1550537408] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8759008 kB) free!
[1550537430] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550537528] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13419656 kB) free.
[1550537550] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550538128] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8381068 kB) free!
[1550538248] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13408292 kB) free.
[1550538750] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550538871] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550539448] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9372888 kB) free!
[1550539470] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550539568] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13404900 kB) free.
[1550539590] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1550540168] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9175436 kB) free!
[1550540190] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550540288] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13417296 kB) free.
[1550540310] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550540889] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9040404 kB) free!
[1550540910] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550541009] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13404084 kB) free.
[1550541030] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550541609] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8943372 kB) free!
[1550541630] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550541729] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13415772 kB) free.
[1550541751] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1550542329] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8823224 kB) free!
[1550542350] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550542449] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13402840 kB) free.
[1550542470] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1550542470] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550543049] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8680656 kB) free!
[1550543070] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550543169] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13397468 kB) free.
[1550543190] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1550543769] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8334904 kB) free!
[1550543889] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13391700 kB) free.
[1550544390] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550544510] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1550544689] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 20.0% (3248248 kB) free!
[1550544809] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 20.0% (3245244 kB) free!
[1550544929] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 20.0% (3242960 kB) free!
[1550545089] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.0% (9378720 kB) free!
[1550545110] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550545210] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13383128 kB) free.
[1550545231] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550545809] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.6% (9187596 kB) free!
[1550545830] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550545929] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13399232 kB) free.
[1550545950] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550546529] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9014688 kB) free!
[1550546550] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550546649] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13389088 kB) free.
[1550546670] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1550547249] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8807144 kB) free!
[1550547270] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550547369] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13213948 kB) free.
[1550547390] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550547510] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550548569] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9545784 kB) free!
[1550548689] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13383832 kB) free.
[1550548710] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1550549310] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550549430] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550549550] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550550489] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8647584 kB) free!
[1550550609] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.2% (12914284 kB) free.
[1550551809] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9324936 kB) free!
[1550551929] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.2% (13380856 kB) free.
[1550551951] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1550551951] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1550552261] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 31.22, 16.74, 7.35
[1550552381] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.19, 11.57, 6.59
[1550552501] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.04, 8.23, 5.97
[1550552551] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550552671] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 298 processes
[1550552791] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550552839] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550552959] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550553079] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550553101] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.88, 2.47, 3.85
[1550553729] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8536080 kB) free!
[1550553849] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.9% (12784772 kB) free.
[1550554072] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.6% (23191616 kB) free.
[1550554140] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.23, 5.28, 5.09
[1550554188] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.1% (14969124 kB) free.
[1550554260] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.33, 5.27, 5.10
[1550554322] SERVICE ALERT: smscapp1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1550554380] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.75, 5.46, 5.20
[1550554442] SERVICE ALERT: smscapp1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1550554562] SERVICE ALERT: smscapp1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1550555050] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9194504 kB) free!
[1550555163] SERVICE ALERT: smscapp1;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1550555170] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.1% (13353092 kB) free.
[1550555191] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1550555792] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550555912] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550556032] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1550556180] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 5.23, 4.96, 4.97
[1550556219] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 52 users currently logged in
[1550556279] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 52 users currently logged in
[1550556339] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 52 users currently logged in
[1550556399] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 53 users currently logged in
[1550556699] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 49 users currently logged in
[1550556970] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8530052 kB) free!
[1550557090] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.0% (12798500 kB) free.
[1550557240] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (947492 kB) free!
[1550557361] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 96.4% (47507672 kB) free.
[1550557599] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 54 users currently logged in
[1550557659] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 54 users currently logged in
[1550557719] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 52 users currently logged in
[1550557779] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 52 users currently logged in
[1550558013] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1550558614] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550558734] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1550558854] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550558890] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8332432 kB) free!
[1550559010] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.5% (12576776 kB) free.
[1550559454] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1550559579] SERVICE FLAPPING ALERT: webserver2;Current Users;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1550559579] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 50 users currently logged in
[1550560625] SERVICE ALERT: db1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 324 processes
[1550560654] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550560745] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 325 processes
[1550560774] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550560779] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 51 users currently logged in
[1550560839] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 51 users currently logged in
[1550560894] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550560899] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 51 users currently logged in
[1550560959] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 51 users currently logged in
[1550562010] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9238912 kB) free!
[1550562130] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.0% (13324864 kB) free.
[1550562163] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (159312 kB) free!
[1550562283] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (155804 kB) free!
[1550562404] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (151924 kB) free!
[1550563042] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 91.8% (14924940 kB) free.
[1550563215] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 341 processes
[1550563930] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8344108 kB) free!
[1550564050] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.0% (13319284 kB) free.
[1550564494] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550565094] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1550565214] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550565334] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550565759] SERVICE FLAPPING ALERT: webserver2;Current Users;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1550565850] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7923552 kB) free!
[1550565971] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 13.7% (6762828 kB) free!
[1550566091] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.1% (12851668 kB) free.
[1550566534] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1550568934] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1550569054] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550569174] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550569691] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8586900 kB) free!
[1550569774] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1550569811] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.9% (13239292 kB) free.
[1550570974] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1550571080] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550571094] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1550571214] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550571255] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550571375] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550571495] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550571680] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1550571800] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1550571814] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1550571920] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1550572211] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8255120 kB) free!
[1550572332] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.8% (13224596 kB) free.
[1550572359] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 61 users currently logged in
[1550572414] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550572534] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1550572654] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550573532] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9795208 kB) free!
[1550573652] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.4% (8061312 kB) free!
[1550573772] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.0% (6883904 kB) free!
[1550574372] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4576312 kB) free!
[1550575054] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1550575095] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550575572] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.3% (7026576 kB) free!
[1550575654] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1550575774] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550575894] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1550576772] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1550576772] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (3475044 kB) free!
[1550576867] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 60.9% (30035004 kB) free.
[1550578894] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1550579494] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1550579614] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1550579734] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1550580334] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1550580454] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550580575] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 300 processes
[1550581076] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9402552 kB) free!
[1550581196] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.3% (6052840 kB) free!
[1550581316] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (480272 kB) free!
[1550581682] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168788 kB) free!
[1550581737] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 59801 MB (26% inode=99%):
[1550581802] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196176 kB) free!
[1550581916] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 44.4% (21864736 kB) free.
[1550581922] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (218824 kB) free!
[1550582259] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 60 users currently logged in
[1550582720] SERVICE ALERT: webservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1550582941] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.4% (23126704 kB) free.
[1550583159] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 50 users currently logged in
[1550583459] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 52 users currently logged in
[1550583519] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 52 users currently logged in
[1550583579] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 52 users currently logged in
[1550583639] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 52 users currently logged in
[1550585376] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1550585976] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1550586096] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550586216] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550586339] SERVICE FLAPPING ALERT: webserver2;Current Users;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1550586339] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 49 users currently logged in
[1550586816] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 295 processes
[1550587416] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550587536] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550587656] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550589639] SERVICE FLAPPING ALERT: webserver2;Current Users;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1550589653] SERVICE ALERT: sig6;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 357 processes
[1550589808] SERVICE ALERT: sig6;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550589809] SERVICE ALERT: sig6;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'sig6' was out of bounds)
[1550589809] HOST ALERT: sig6;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.58)
[1550589809] SERVICE ALERT: sig6;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig6' was out of bounds)
[1550589848] SERVICE ALERT: sig6;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig6' was out of bounds)
[1550589852] SERVICE ALERT: sig6;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig6' was out of bounds)
[1550589891] SERVICE ALERT: sig6;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sig6' was out of bounds)
[1550589928] SERVICE ALERT: sig6;Total Processes;OK;HARD;1;PROCS OK: 319 processes
[1550589929] SERVICE ALERT: sig6;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1550589929] SERVICE ALERT: sig6;CPU Load;OK;HARD;1;OK - load average: 0.39, 0.13, 0.05
[1550589932] HOST ALERT: sig6;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1550589968] SERVICE ALERT: sig6;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1550589972] SERVICE ALERT: sig6;HDD_Root status;OK;HARD;1;DISK OK - free space: / 26875 MB (52% inode=98%):
[1550590011] SERVICE ALERT: sig6;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 212837 MB (94% inode=99%):
[1550590056] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1550590253] SERVICE ALERT: sig6;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1550590539] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 52 users currently logged in
[1550590599] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 52 users currently logged in
[1550590656] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550590659] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 53 users currently logged in
[1550590719] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 53 users currently logged in
[1550590776] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550590896] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550591019] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 49 users currently logged in
[1550591496] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1550592096] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1550592217] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1550592337] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 299 processes
[1550593316] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1550594019] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 51 users currently logged in
[1550594079] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 51 users currently logged in
[1550594139] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 51 users currently logged in
[1550594199] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 51 users currently logged in
[1550595656] SERVICE ALERT: sig6;CPU_Procs;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU_Procs' on host 'sig6' was out of bounds)
[1550595658] HOST ALERT: sig6;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1550595689] SERVICE ALERT: sig6;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig6' was out of bounds)
[1550595722] HOST ALERT: sig6;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1550595776] SERVICE ALERT: sig6;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 319 processes
[1550595809] SERVICE ALERT: sig6;Sensors Status;OK;HARD;1;SENSORS OK
[1550595896] SERVICE ALERT: sig6;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1550596106] HOST ALERT: sig6;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1550596127] SERVICE ALERT: sig6;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig6' was out of bounds)
[1550596170] HOST FLAPPING ALERT: sig6;STARTED; Host appears to have started flapping (21.1% change > 20.0% threshold)
[1550596170] HOST ALERT: sig6;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1550596247] SERVICE ALERT: sig6;Memory;OK;HARD;1;OK - 94.5% (15362940 kB) free.
[1550596537] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1550596857] SERVICE ALERT: sig6;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550596860] HOST ALERT: sig6;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.58)
[1550596924] HOST ALERT: sig6;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 603.20 ms
[1550596978] SERVICE ALERT: sig6;Memory;OK;HARD;1;OK - 94.6% (15376836 kB) free.
[1550597619] SERVICE ALERT: sig6;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550597649] HOST ALERT: sig6;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1550597706] SERVICE ALERT: sig6;CPU_Procs;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550597732] HOST ALERT: sig6;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1550597739] SERVICE ALERT: sig6;Sensors Status;OK;HARD;1;SENSORS OK
[1550597826] SERVICE FLAPPING ALERT: sig6;CPU_Procs;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1550597826] SERVICE ALERT: sig6;CPU_Procs;OK;HARD;1;CPU OK: 322 processes
[1550598337] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550598372] HOST FLAPPING ALERT: sig6;STOPPED; Host appears to have stopped flapping (4.4% change < 5.0% threshold)
[1550598457] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1550599057] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1550599177] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1550599777] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1550599897] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1550600138] SERVICE ALERT: sig6;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550600139] SERVICE ALERT: sig6;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1550600139] SERVICE ALERT: sig6;CPU Load;CRITICAL;SOFT;1;(Return code of 255 for service 'CPU Load' on host 'sig6' was out of bounds)
[1550600142] SERVICE ALERT: sig6;Sensors Status;CRITICAL;SOFT;1;(Return code of 255 for service 'Sensors Status' on host 'sig6' was out of bounds)
[1550600158] HOST ALERT: sig6;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1550600171] SERVICE ALERT: sig6;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig6' was out of bounds)
[1550600174] SERVICE ALERT: sig6;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig6' was out of bounds)
[1550600211] SERVICE ALERT: sig6;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sig6' was out of bounds)
[1550600222] HOST ALERT: sig6;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1550600226] SERVICE ALERT: sig6;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1550600258] SERVICE ALERT: sig6;Total Processes;OK;HARD;1;PROCS OK: 310 processes
[1550600259] SERVICE ALERT: sig6;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1550600259] SERVICE ALERT: sig6;CPU Load;OK;HARD;1;OK - load average: 0.28, 0.12, 0.05
[1550600262] SERVICE ALERT: sig6;Sensors Status;OK;HARD;1;SENSORS OK
[1550600291] SERVICE FLAPPING ALERT: sig6;LocalService-NTPD SERVICE;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1550600291] SERVICE ALERT: sig6;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1550600294] SERVICE FLAPPING ALERT: sig6;HDD_Root status;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1550600294] SERVICE ALERT: sig6;HDD_Root status;OK;HARD;1;DISK OK - free space: / 27108 MB (52% inode=98%):
[1550600331] SERVICE FLAPPING ALERT: sig6;HDD_Home status;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1550600331] SERVICE ALERT: sig6;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 212957 MB (94% inode=99%):
[1550600346] SERVICE ALERT: sig6;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1550600498] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1550600618] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes