[1546372398] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1546374198] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 86.40, 117.63, 52.97
[1546374318] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 11.99, 78.87, 46.60
[1546374438] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.29, 53.00, 41.04
[1546375638] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.70, 2.83, 12.60
[1546376838] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.50, 0.60, 3.90
[1546383438] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 2.65, 54.87, 41.47
[1546383558] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.84, 36.88, 36.50
[1546383678] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.45, 24.81, 32.13
[1546384278] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (25.6% change >= 20.0% threshold)
[1546384278] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.52, 3.84, 17.13
[1546386078] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.52, 0.59, 3.01
[1546396278] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1546398678] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 0.89, 34.40, 35.61
[1546398798] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.67, 23.22, 31.37
[1546398918] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.44, 17.52, 28.26
[1546399519] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.68, 3.02, 15.24
[1546400719] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.31, 0.58, 4.62
[1546406719] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.30, 37.20, 30.19
[1546406839] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.68, 25.08, 26.60
[1546406959] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1546406959] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.73, 16.99, 23.44
[1546407559] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.37, 4.73, 13.49
[1546408759] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.61, 0.83, 4.43
[1546409677] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1546409797] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1546409917] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1546410512] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546410569] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8496720 kB) free!
[1546410632] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546410689] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10173048 kB) free.
[1546410752] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546413596] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1546413716] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1546413836] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1546414117] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1546416559] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 2.17, 42.78, 32.15
[1546416679] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.47, 28.69, 28.28
[1546416799] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.57, 19.39, 24.92
[1546417399] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.97, 3.22, 13.37
[1546418599] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.77, 0.83, 4.21
[1546419236] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1546420436] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1546420556] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1546420676] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1546423289] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8410672 kB) free!
[1546423352] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546423409] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10173552 kB) free.
[1546423472] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1546424609] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7176500 kB) free!
[1546424730] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10174208 kB) free.
[1546427276] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 14 users currently logged in
[1546427672] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546427729] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8828364 kB) free!
[1546427792] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546427849] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.8% (7296084 kB) free!
[1546427912] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 288 processes
[1546427969] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10171644 kB) free.
[1546429399] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1546434476] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1546434596] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1546434716] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1546440116] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 14 users currently logged in
[1546444916] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1546445036] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1546445156] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 13 users currently logged in
[1546445913] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546445969] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9453224 kB) free!
[1546446033] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546446089] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.1% (7919556 kB) free!
[1546446153] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546446209] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.6% (10162580 kB) free.
[1546446799] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 24.88, 100.26, 54.49
[1546446919] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.11, 67.37, 47.99
[1546447039] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.42, 45.49, 42.32
[1546448239] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.86, 1.71, 12.31
[1546449439] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.87, 0.81, 3.96