[1589913610] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (3077512 kB) free!
[1589914259] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.0% (22354432 kB) free!
[1589914810] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1167208 kB) free!
[1589914860] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396028 kB) free!
[1589915410] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (3006380 kB) free!
[1589915460] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (18447608 kB) free!
[1589916060] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (456424 kB) free!
[1589916610] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (798304 kB) free!
[1589916660] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17385840 kB) free!
[1589917210] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.4% (2022612 kB) free!
[1589917260] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (430236 kB) free!
[1589917396] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 31443 MB (13.99% inode=100%):
[1589917456] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 28920 MB (12.86% inode=100%):
[1589917519] SERVICE ALERT: webserver2;HDD_Home status;WARNING;SOFT;3;DISK WARNING - free space: /home 25880 MB (11.51% inode=100%):
[1589917579] SERVICE ALERT: webserver2;HDD_Home status;CRITICAL;HARD;4;DISK CRITICAL - free space: /home 21833 MB (9.71% inode=100%):
[1589917860] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14138600 kB) free!
[1589918179] SERVICE ALERT: webserver2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 74237 MB (33.03% inode=100%):
[1589918255] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 80.4% (19699852 kB) free.
[1589918400] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1707 processes
[1589918430] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (231920 kB) free!
[1589918460] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (453492 kB) free!
[1589918550] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (238652 kB) free!
[1589918624] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9845952 kB) free!
[1589918670] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (236956 kB) free!
[1589918744] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2624544 kB) free!
[1589918864] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (1366532 kB) free!
[1589919610] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1236904 kB) free!
[1589920210] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.9% (2415560 kB) free!
[1589920470] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.8% (9254576 kB) free!
[1589922001] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1696 processes
[1589922060] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23251340 kB) free!
[1589922660] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388124 kB) free!
[1589922755] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.1% (3211956 kB) free!
[1589923026] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 364 processes
[1589923146] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1589923221] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 92.3% (45474020 kB) free.
[1589923261] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23254368 kB) free!
[1589923355] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187576 kB) free!
[1589923667] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.7% (45193712 kB) free.
[1589923861] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (406336 kB) free!
[1589924003] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1589924070] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 90.1% (44395780 kB) free.
[1589924110] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2417) < 172.105.59.173 [20/May/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1589924147] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(402) < 2020/05/20 03:01:32 [error] 22121#0: *89428033 no live upstreams while connecting to upstream, client: 172.105.57.57, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1589924230] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589924267] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589924384] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(316) < 172.105.57.57 [20/May/2020:03:01:32 +0530] TCP 502 0 0 60.000 "10.147.212.37:4000, backend3" "0, 0" "0, 0" "60.000, 0.000"
[1589924461] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14805936 kB) free!
[1589924505] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589925061] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (411996 kB) free!
[1589925602] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1702 processes
[1589925610] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (1539752 kB) free!
[1589926202] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1589926202] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1696 processes
[1589926210] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.0% (1952684 kB) free!
[1589926810] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.8% (4517364 kB) free.
[1589927410] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (2894220 kB) free!
[1589927430] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 34.1% (8362228 kB) free.
[1589927461] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30291172 kB) free.
[1589927530] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.7% (4498320 kB) free.
[1589927604] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1589927724] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1589927844] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1589928003] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1703 processes
[1589928030] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (247356 kB) free!
[1589928061] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7563100 kB) free!
[1589928130] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (3109508 kB) free!
[1589928150] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187840 kB) free!
[1589928181] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3521308 kB) free!
[1589928251] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 32.7% (5319232 kB) free.
[1589928270] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194256 kB) free!
[1589928301] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34029248 kB) free.
[1589928444] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1589928603] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1693 processes
[1589928901] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12213776 kB) free!
[1589929021] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8256100 kB) free!
[1589929044] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1589929141] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (4437628 kB) free!
[1589929164] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1589929284] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1589929284] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 350 processes
[1589929314] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.6% (4796956 kB) free!
[1589929434] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.4% (4759136 kB) free!
[1589929554] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.3% (4724840 kB) free!
[1589929741] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17445188 kB) free!
[1589930341] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31598032 kB) free.
[1589930670] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 39.4% (9664124 kB) free.
[1589930942] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10882888 kB) free!
[1589931062] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (6998052 kB) free!
[1589931182] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3471000 kB) free!
[1589931270] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195116 kB) free!
[1589931390] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188880 kB) free!
[1589931510] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1589931510] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (203348 kB) free!
[1589931782] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20415988 kB) free!
[1589932382] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31200368 kB) free.
[1589932668] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/20 05:19:52 [error] 22135#0: *89444444 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:141/17, bytes from/to upstream:17/141
[1589932788] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589932962] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 69.5% (17032948 kB) free.
[1589932982] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12516068 kB) free!
[1589933102] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37854076 kB) free.
[1589933158] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 42.5% (10485048 kB) free.
[1589933555] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 48.9% (11994248 kB) free.
[1589933702] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18788924 kB) free!
[1589933822] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.4% (15082636 kB) free!
[1589933942] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (11493292 kB) free!
[1589934542] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19752192 kB) free!
[1589935142] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30084884 kB) free.
[1589935203] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1706 processes
[1589935554] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1947144 kB) free!
[1589935742] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9246780 kB) free!
[1589935862] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37004224 kB) free.
[1589936462] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16721956 kB) free!
[1589936582] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12889412 kB) free!
[1589936703] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (9102060 kB) free!
[1589937003] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1698 processes
[1589937155] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163088 kB) free!
[1589937275] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181780 kB) free!
[1589937303] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18182940 kB) free!
[1589937395] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169864 kB) free!
[1589937455] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 17.1% (4197000 kB) free!
[1589937515] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 17.3% (4231392 kB) free!
[1589937575] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 17.4% (4257396 kB) free!
[1589937635] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 17.4% (4252628 kB) free!
[1589937903] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29538236 kB) free.
[1589937961] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196392 kB) free!
[1589938064] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 73.7% (97099168 kB) free.
[1589938081] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198600 kB) free!
[1589938188] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/20 06:52:29 [error] 22133#0: *89454936 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:112/17, bytes from/to upstream:17/112
[1589938201] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (185008 kB) free!
[1589938203] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1702 processes
[1589938308] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589938450] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589938503] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7208828 kB) free!
[1589938535] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (200560 kB) free!
[1589938623] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35858156 kB) free.
[1589938835] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 24.8% (6066792 kB) free.
[1589938908] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/20 07:02:30 [error] 22132#0: *89456018 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:154/17, bytes from/to upstream:17/154
[1589939028] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/05/20 07:12:33 [error] 22135#0: *89457144 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:253/17, bytes from/to upstream:17/253
[1589939135] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 15.4% (3770300 kB) free!
[1589939148] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589939195] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 15.4% (3765904 kB) free!
[1589939223] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12957512 kB) free!
[1589939256] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 15.4% (3765252 kB) free!
[1589939316] SERVICE FLAPPING ALERT: webserver2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1589939316] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 15.4% (3771120 kB) free!
[1589939343] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8654696 kB) free!
[1589939463] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (4403396 kB) free!
[1589940063] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15770268 kB) free!
[1589940603] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1700 processes
[1589940663] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27822156 kB) free.
[1589940685] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589940948] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/05/20 07:44:50 [error] 22132#0: *89460789 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:251/17, bytes from/to upstream:17/251
[1589941068] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589941263] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5719140 kB) free!
[1589941383] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33744068 kB) free.
[1589941949] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.84, 7.72, 5.98
[1589941983] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11130712 kB) free!
[1589942069] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.69, 7.80, 6.21
[1589942103] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6610088 kB) free!
[1589942189] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.52, 7.89, 6.43
[1589942223] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2210352 kB) free!
[1589942404] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1702 processes
[1589942824] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14592308 kB) free!
[1589942868] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/20 08:16:16 [error] 22133#0: *89464331 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:227/17, bytes from/to upstream:17/227
[1589942910] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589942988] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589943004] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1695 processes
[1589943174] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.8% (1763120 kB) free!
[1589943424] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26726020 kB) free.
[1589943589] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/20 08:23:47 [error] 22133#0: *89465274 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:135/17, bytes from/to upstream:17/135
[1589943709] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589944024] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2891964 kB) free!
[1589944144] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (414200 kB) free!
[1589944264] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31604488 kB) free.
[1589944309] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/20 08:39:51 [error] 22133#0: *89467235 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:252/17, bytes from/to upstream:17/252
[1589944429] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/05/20 08:43:37 [error] 22135#0: *89467708 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:129/17, bytes from/to upstream:17/129
[1589944549] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;HARD;3;(1) < 2020/05/20 08:43:50 [error] 22131#0: *89467732 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:129/17, bytes from/to upstream:17/129
[1589944864] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6789444 kB) free!
[1589944984] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2371792 kB) free!
[1589945016] SERVICE FLAPPING ALERT: webserver2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589945104] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (439132 kB) free!
[1589945230] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(778) < 103.255.146.164 [20/May/2020:08:56:53 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589945350] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589945402] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.2% (4087064 kB) free!
[1589945484] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1589945522] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.6% (15848916 kB) free.
[1589946005] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1701 processes
[1589946304] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31758156 kB) free.
[1589946349] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589946390] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 28.59, 26.75, 17.98
[1589946850] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (1061144 kB) free!
[1589946904] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6103140 kB) free!
[1589946970] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.9% (3565848 kB) free.
[1589946990] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.68, 10.81, 13.51
[1589947024] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1658224 kB) free!
[1589947144] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381912 kB) free!
[1589947205] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1700 processes
[1589947549] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/20 09:31:23 [error] 22131#0: *89473974 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:111/17, bytes from/to upstream:17/111
[1589947669] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/05/20 09:36:09 [error] 22133#0: *89474783 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:233/17, bytes from/to upstream:17/233
[1589947789] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589948190] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1589948190] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 26.65, 25.88, 20.23
[1589948389] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/20 09:47:25 [error] 22132#0: *89476647 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:227/17, bytes from/to upstream:17/227
[1589948484] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 720 processes
[1589948504] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44864 MB (19% inode=99%):
[1589948509] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/05/20 09:49:55 [error] 22133#0: *89477055 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:233/17, bytes from/to upstream:17/233
[1589948624] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44864 MB (19% inode=99%):
[1589948629] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589948744] SERVICE ALERT: primary-server;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44864 MB (19% inode=99%):
[1589948770] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (2952604 kB) free!
[1589948790] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 11.43, 14.97, 17.01
[1589948890] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3274384 kB) free.
[1589948950] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 172.105.59.173 [20/May/2020:09:55:06 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1589949071] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589949084] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1589949204] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1589949229] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/20 10:01:21 [error] 22135#0: *89478236 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:253/17, bytes from/to upstream:17/253
[1589949324] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 386 processes
[1589949350] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589949545] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29688352 kB) free.
[1589949670] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4) < 203.212.70.170 [20/May/2020:10:10:17 +0530] TCP 502 0 0 91.078 "10.147.212.154:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "31.078, 60.000, 0.000"
[1589949790] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589949924] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1589949949] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/20 10:08:51 [error] 22133#0: *89479171 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:231/17, bytes from/to upstream:17/231
[1589950044] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1589950069] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589950090] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.6% (3017560 kB) free!
[1589950145] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4427460 kB) free!
[1589950164] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1589950205] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1702 processes
[1589950210] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.4% (2510436 kB) free!
[1589950265] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (402692 kB) free!
[1589950330] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.4% (1851448 kB) free!
[1589950363] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.0% (3688012 kB) free!
[1589950385] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400856 kB) free!
[1589950483] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.1% (3449696 kB) free!
[1589950603] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 13.0% (3175228 kB) free!
[1589950669] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/20 10:27:27 [error] 22125#0: *89481444 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:233/17, bytes from/to upstream:17/233
[1589950789] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589951203] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1744484 kB) free!
[1589951389] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(9) < 2020/05/20 10:37:41 [error] 22120#0: *89482793 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:207/17, bytes from/to upstream:17/207
[1589951509] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589951522] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4488504 kB) free!
[1589951642] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.5% (3937008 kB) free!
[1589951707] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1589951762] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.0% (3378568 kB) free!
[1589951827] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1589951947] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1589952363] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (592560 kB) free!
[1589952390] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 29.51, 28.05, 23.24
[1589952547] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1589952605] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1695 processes
[1589952894] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.7290.599> (10.147.212.158:48788 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1589953014] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589953212] SERVICE ALERT: sig4-vm2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 577 processes
[1589953332] SERVICE ALERT: sig4-vm2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 573 processes
[1589953452] SERVICE ALERT: sig4-vm2;Total Processes;WARNING;HARD;3;PROCS WARNING: 583 processes
[1589953544] SERVICE ALERT: primary-server;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 117197 MB (52% inode=99%):
[1589953561] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 62.6% (15350860 kB) free.
[1589953808] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1769 processes
[1589953909] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/20 11:21:14 [error] 22135#0: *89491472 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1589953923] SERVICE ALERT: sigqscluster1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1589953974] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (210596 kB) free!
[1589954029] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/05/20 11:23:12 [error] 22135#0: *89491890 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1589954046] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 1767 processes
[1589954149] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589954530] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.2% (4417328 kB) free.
[1589954749] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/20 11:32:09 [error] 22135#0: *89493909 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1589954869] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589954944] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 14.4% (3544460 kB) free!
[1589954994] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 12.5% (4390836 kB) free!
[1589955064] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 12.6% (3108428 kB) free!
[1589955114] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 11.1% (3871772 kB) free!
[1589955184] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 10.7% (2648960 kB) free!
[1589955234] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (3327284 kB) free!
[1589955470] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/05/20 11:45:17 [error] 22135#0: *89497890 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1589955590] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589955704] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.22, 5.74, 5.04
[1589955784] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (245644 kB) free!
[1589955785] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.8% (15540944 kB) free!
[1589955791] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(697) < 103.255.146.164 [20/May/2020:11:50:51 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589955824] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 2.92, 4.82, 4.80
[1589955911] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589956067] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9795428 kB) free!
[1589956187] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.3% (9523596 kB) free!
[1589956221] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9574332 kB) free!
[1589956307] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.8% (9282664 kB) free!
[1589956341] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (9313016 kB) free!
[1589956385] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (431552 kB) free!
[1589956435] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.6% (31732584 kB) free.
[1589956461] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (9085976 kB) free!
[1589956563] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.4% (16042904 kB) free.
[1589956764] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1589956790] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/20 12:03:45 [error] 22127#0: *89504029 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1589956794] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 339 processes
[1589956910] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/05/20 12:10:10 [error] 22132#0: *89506407 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1589956914] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1589956930] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.8% (2238076 kB) free!
[1589956986] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16732700 kB) free!
[1589957030] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;HARD;3;(1) < 2020/05/20 12:13:29 [error] 22133#0: *89507447 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1589957050] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.9% (3726632 kB) free.
[1589957111] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 172.105.59.173 [20/May/2020:12:07:46 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1589957231] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589957507] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4312748 kB) free!
[1589957584] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.1% (22988908 kB) free.
[1589957586] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391980 kB) free!
[1589957629] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.20750.6065> (10.147.212.143:42104 -> 10.147.212.160:5672, vhost: '/', user: 'guest'), channel 1:
[1589957653] SERVICE ALERT: sigqscluster1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1589957749] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589957776] SERVICE ALERT: sigqscluster1;CPU_Procs;OK;HARD;1;CPU OK: 1782 processes
[1589958186] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20310884 kB) free!
[1589958225] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 13.27, 8.88, 6.45
[1589958230] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589958250] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3125372 kB) free!
[1589958345] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.22, 8.96, 6.77
[1589958370] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.2% (1651100 kB) free!
[1589958465] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.54, 7.80, 6.62
[1589958490] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.6% (3990776 kB) free.
[1589958786] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (419508 kB) free!
[1589958831] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/20 12:38:41 [error] 22135#0: *89515721 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/dashboard HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/dashboard", host: "altsmsc.tesync.net", referrer: "http://wap.tesync.net/tesync/AnalysisSMS.jsp"
[1589958861] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4085092 kB) free!
[1589958951] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589958961] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.4% (4520516 kB) free!
[1589959081] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.7% (4331352 kB) free!
[1589959201] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.9% (4140704 kB) free!
[1589959690] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (3086068 kB) free!
[1589959810] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.1% (2948556 kB) free!
[1589959930] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (2806056 kB) free!
[1589960401] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1736716 kB) free!
[1589960790] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1589961130] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1589961130] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3362888 kB) free.
[1589961186] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18173700 kB) free!
[1589961351] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/20 13:22:57 [error] 22135#0: *89529484 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1589961471] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589961786] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390072 kB) free!
[1589962014] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.9108.635> (10.147.212.158:56218 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1589962098] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 690 processes
[1589962134] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589962330] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.9% (2420660 kB) free!
[1589962450] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.1% (2622144 kB) free!
[1589962570] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.4% (2826956 kB) free!
[1589962665] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.63, 4.49, 4.75
[1589962748] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1589962764] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1589962868] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1589962884] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1589962988] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1589963004] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1589963265] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.02, 5.47, 5.19
[1589963271] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/20 13:49:35 [error] 22121#0: *89537661 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1589963385] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.62, 5.02, 5.08
[1589963391] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589963505] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.86, 5.15, 5.12
[1589963514] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 342 processes
[1589963586] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15211836 kB) free!
[1589963588] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1589963588] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1589963604] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1589963634] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1589964001] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1589964001] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 59.4% (14548432 kB) free.
[1589964186] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (408596 kB) free!
[1589964608] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1589964786] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30262436 kB) free.
[1589964963] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.3% (3437016 kB) free!
[1589964970] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.2% (3450848 kB) free.
[1589965083] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.5% (3038856 kB) free!
[1589965204] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.5% (2581756 kB) free!
[1589965386] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2961180 kB) free!
[1589965506] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (396368 kB) free!
[1589965570] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.2% (1817192 kB) free!
[1589965627] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (378252 kB) free!
[1589965690] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.3% (2329400 kB) free!
[1589965804] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (383832 kB) free!
[1589965810] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.0% (2920164 kB) free!
[1589966149] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.24352.6053> (10.147.212.53:37290 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1589966269] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1589967004] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1589967004] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 36.7% (8243376 kB) free.
[1589967010] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (846540 kB) free!
[1589967213] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 642 processes
[1589967333] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 641 processes
[1589967427] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.2% (21384712 kB) free!
[1589967431] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(474) < 103.255.146.164 [20/May/2020:15:02:05 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1589967453] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 356 processes
[1589967551] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589967834] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 340 processes
[1589967954] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1589968027] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (449824 kB) free!
[1589968191] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/20 15:16:18 [error] 22135#0: *89564386 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1589968210] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.6% (2702300 kB) free!
[1589968311] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589968627] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.4% (24243792 kB) free!
[1589968801] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.5% (4779864 kB) free!
[1589968810] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1044024 kB) free!
[1589968911] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/20 15:21:54 [error] 22135#0: *89566290 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1589968922] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.7% (4592484 kB) free!
[1589968974] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 12.2% (1985744 kB) free!
[1589969031] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589969042] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.9% (4381496 kB) free!
[1589969227] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (422224 kB) free!
[1589969631] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/20 15:39:44 [error] 22131#0: *89571714 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net"
[1589969752] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589969827] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28557672 kB) free.
[1589969951] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 203.212.70.170 [20/May/2020:15:47:57 +0530] TCP 502 0 0 60.000 "10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.000, 60.000, 0.000"
[1589970071] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589970242] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1905728 kB) free!
[1589970427] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (410416 kB) free!
[1589970453] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 803 processes
[1589970547] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (405736 kB) free!
[1589970574] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 804 processes
[1589970610] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.7% (1743632 kB) free!
[1589970667] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (457928 kB) free!
[1589970671] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 203.212.70.170 [20/May/2020:15:57:29 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.000, 60.000, 0.000"
[1589970694] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 377 processes
[1589970791] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589971267] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17362704 kB) free!
[1589971867] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (424560 kB) free!
[1589972467] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.5% (24375208 kB) free!
[1589972754] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 343 processes
[1589972874] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1589973067] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (412252 kB) free!
[1589973610] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (316504 kB) free!
[1589973668] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32352144 kB) free.
[1589974210] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.3% (1674140 kB) free!
[1589974268] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (1980304 kB) free!
[1589974388] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (429612 kB) free!
[1589974508] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (450324 kB) free!
[1589974804] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4417316 kB) free!
[1589974810] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3471776 kB) free.
[1589974924] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.3% (4106308 kB) free!
[1589974988] SERVICE FLAPPING ALERT: primary-server;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589975044] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.9% (3799728 kB) free!
[1589975108] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26884008 kB) free.
[1589975410] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (469772 kB) free!
[1589975530] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.6% (1727372 kB) free!
[1589975592] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(24) < 203.212.70.170 [20/May/2020:17:21:38 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1589975650] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3303228 kB) free.
[1589975708] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (446184 kB) free!
[1589975712] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589975828] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (442772 kB) free!
[1589975948] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (457140 kB) free!
[1589976244] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (1838128 kB) free!
[1589976312] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 203.212.70.170 [20/May/2020:17:26:47 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "0.000, 60.000, 0.000"
[1589976432] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589976548] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17927536 kB) free!
[1589976850] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.4% (2348680 kB) free!
[1589976951] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/05/20 17:41:48 [error] 22135#0: *89605569 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1589976970] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.5% (3653708 kB) free.
[1589977071] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589977148] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (376040 kB) free!
[1589977571] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.4% (392860 kB) free!
[1589977691] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.7% (1897148 kB) free!
[1589977748] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.7% (25937404 kB) free!
[1589977811] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.2% (3453936 kB) free.
[1589978045] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.2% (16206896 kB) free.
[1589978348] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (417900 kB) free!
[1589978872] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/20 18:10:00 [error] 22135#0: *89612269 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1589978948] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34469788 kB) free.
[1589978992] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1589979011] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.5% (2025388 kB) free!
[1589979131] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3582252 kB) free.
[1589979548] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (4764120 kB) free!
[1589979668] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (403240 kB) free!
[1589979731] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.9% (1773968 kB) free!
[1589979789] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (439560 kB) free!
[1589979851] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3431620 kB) free.
[1589980389] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (25861820 kB) free!
[1589980989] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (414324 kB) free!
[1589981042] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1589981051] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.0% (1944060 kB) free!
[1589981172] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3388636 kB) free.
[1589981589] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33599160 kB) free.
[1589982189] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.5% (2001884 kB) free!
[1589982309] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (393240 kB) free!
[1589982371] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (1310692 kB) free!
[1589982429] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407016 kB) free!
[1589982491] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (3069916 kB) free!
[1589982611] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3760284 kB) free.
[1589983029] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.1% (25120128 kB) free!
[1589983211] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (740208 kB) free!
[1589983331] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.5% (2195960 kB) free!
[1589983451] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.7% (3856868 kB) free.
[1589983629] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398240 kB) free!
[1589983905] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.79, 4.50, 4.88
[1589984229] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34285100 kB) free.
[1589984651] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (1238400 kB) free!
[1589984771] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.4% (2668748 kB) free!
[1589984829] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3252688 kB) free!
[1589984891] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3694940 kB) free.
[1589984949] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (401624 kB) free!
[1589985069] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407952 kB) free!
[1589985105] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.62, 5.96, 5.40
[1589985225] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.68, 5.27, 5.23
[1589985345] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 2.23, 4.35, 4.91
[1589985491] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (196008 kB) free!
[1589985611] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (1609216 kB) free!
[1589985669] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20850432 kB) free!
[1589985731] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3329280 kB) free.
[1589985945] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.55, 5.61, 5.34
[1589986065] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.22, 5.11, 5.19
[1589986185] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.21, 4.70, 5.02
[1589986269] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (445128 kB) free!
[1589986870] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26143304 kB) free!
[1589987470] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (434812 kB) free!
[1589987531] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.1% (505904 kB) free!
[1589987651] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3394676 kB) free.
[1589988070] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27491656 kB) free.
[1589988245] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1589988251] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1589988670] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (452788 kB) free!
[1589988790] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (401092 kB) free!
[1589988910] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (395460 kB) free!
[1589989510] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (25759504 kB) free!
[1589990110] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396628 kB) free!
[1589990710] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.9% (26281284 kB) free!
[1589991310] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (427360 kB) free!
[1589991585] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.25, 4.24, 5.00
[1589991910] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31905576 kB) free.
[1589992185] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.46, 6.48, 5.84
[1589992305] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.82, 5.65, 5.62
[1589992425] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1589992425] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.35, 5.42, 5.53
[1589992451] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3235164 kB) free!
[1589992510] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.2% (2940656 kB) free!
[1589992572] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3292080 kB) free.
[1589992630] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (406780 kB) free!
[1589992750] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (404052 kB) free!
[1589993171] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3218124 kB) free!
[1589993291] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (3200280 kB) free!
[1589993350] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19658292 kB) free!
[1589993411] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (493464 kB) free!
[1589993507] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 13.4% (6620764 kB) free!
[1589993625] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 5.40, 4.78, 4.97
[1589993951] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (400276 kB) free!
[1589994011] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3271948 kB) free.
[1589994107] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 27.7% (13665688 kB) free.
[1589994225] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.29, 6.08, 5.54
[1589994345] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.15, 6.25, 5.68
[1589994390] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.13, 6.62, 16.97
[1589994465] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.20, 5.69, 5.53
[1589994551] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.6% (24478512 kB) free!
[1589994611] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (3008496 kB) free!
[1589994731] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (3143824 kB) free!
[1589994851] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (3136752 kB) free!
[1589994861] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 28.2% (13902128 kB) free.
[1589995151] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (432136 kB) free!
[1589995751] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27506592 kB) free.
[1589996351] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (477024 kB) free!
[1589996471] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (455804 kB) free!
[1589996591] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (450404 kB) free!
[1589996651] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1589996651] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.7% (4504576 kB) free.
[1589996865] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.45, 4.43, 4.77
[1589997191] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16473404 kB) free!
[1589997251] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.8% (2250708 kB) free!
[1589997371] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.0% (2600364 kB) free!
[1589997390] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.49, 4.07, 4.58
[1589997465] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.33, 5.02, 5.17
[1589997491] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (3015248 kB) free!
[1589997585] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.08, 5.34, 5.25
[1589997705] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.87, 5.27, 5.22
[1589997791] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (447788 kB) free!
[1589998092] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.8% (4351160 kB) free.
[1589998391] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17531152 kB) free!
[1589998905] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.43, 3.84, 4.51
[1589998991] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398600 kB) free!