[1591728071] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181280 kB) free!
[1591728599] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 37.62, 32.12, 36.10
[1591729199] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 59.84, 50.25, 42.20
[1591729799] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 47.03, 45.35, 43.48
[1591730399] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1591730399] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 34.31, 46.54, 46.33
[1591731493] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174076 kB) free!
[1591731553] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183552 kB) free!
[1591731600] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 46.89, 44.19, 44.82
[1591731614] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (195948 kB) free!
[1591731674] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (177072 kB) free!
[1591732200] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 62.68, 47.32, 44.73
[1591732274] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 82.7% (20268404 kB) free.
[1591732449] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396164 kB) free!
[1591732675] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 198751 MB (18% inode=99%):
[1591732769] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (4883916 kB) free!
[1591732795] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 198751 MB (18% inode=99%):
[1591732889] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (4862856 kB) free!
[1591732915] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 198751 MB (18% inode=99%):
[1591733010] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4833644 kB) free!
[1591734810] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.9% (29529968 kB) free.
[1591735800] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 23.30, 36.84, 42.23
[1591736400] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 74.48, 49.97, 45.01
[1591736806] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 12.2% (6020140 kB) free!
[1591737406] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (2320396 kB) free!
[1591737600] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 43.50, 42.13, 44.19
[1591738006] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 88.2% (43486464 kB) free.
[1591738200] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 79.04, 46.26, 44.08
[1591738205] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(427) < 103.255.146.164 [10/Jun/2020:03:00:05 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1591738209] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.0% (44862344 kB) free.
[1591738325] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1108) < 172.105.59.173 [10/Jun/2020:03:01:32 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1591738446] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591738469] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(187) < 103.16.101.72 [10/Jun/2020:03:00:12 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1591738486] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(243) < 2020/06/10 03:01:22 [error] 12485#0: *94071612 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1591738590] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591738606] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591738800] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 45.12, 44.11, 43.60
[1591739647] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 20.6% (5046328 kB) free.
[1591740247] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197236 kB) free!
[1591740367] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198060 kB) free!
[1591740487] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193004 kB) free!
[1591744687] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.6% (10200672 kB) free.
[1591745887] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186592 kB) free!
[1591746008] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176524 kB) free!
[1591746128] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1591746128] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178312 kB) free!
[1591746600] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 63.78, 51.20, 46.10
[1591746671] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 71.8% (17607132 kB) free.
[1591746737] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 40.0% (9872884 kB) free.
[1591747338] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.4% (3052712 kB) free!
[1591747458] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187744 kB) free!
[1591747578] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175172 kB) free!
[1591748778] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 20.1% (4963036 kB) free.
[1591749105] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.0% (6412264 kB) free!
[1591749378] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186620 kB) free!
[1591749498] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191960 kB) free!
[1591749618] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1591749618] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187684 kB) free!
[1591749705] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (222404 kB) free!
[1591750029] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 70.3% (17236200 kB) free.
[1591752115] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 279509 MB (26% inode=99%):
[1591752249] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 55.9% (73699836 kB) free.
[1591752600] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1591756928] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1591756961] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.59, 5.64, 5.33
[1591757081] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.84, 5.42, 5.27
[1591757201] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.41, 5.71, 5.39
[1591759601] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 4.63, 4.31, 4.81
[1591760202] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.95, 7.28, 5.98
[1591760322] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.38, 7.14, 6.09
[1591760442] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 8.14, 7.62, 6.40
[1591761018] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1591763472] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1591763592] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1591763829] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.8% (3633624 kB) free!
[1591763949] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.6% (3328424 kB) free!
[1591764069] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.3% (3004816 kB) free!
[1591764071] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196772 kB) free!
[1591764192] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195104 kB) free!
[1591764192] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1591764312] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198940 kB) free!
[1591764312] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1591764432] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1591764460] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1591764581] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1591764669] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1739964 kB) free!
[1591764701] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1591765032] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1591765301] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1591766477] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 10.7% (2393868 kB) free!
[1591766597] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (1928972 kB) free!
[1591766717] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1428364 kB) free!
[1591767069] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (3893036 kB) free!
[1591767521] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 52.3% (12813208 kB) free.
[1591767669] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1591767669] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1869300 kB) free!
[1591767689] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 12.5% (2034024 kB) free!
[1591768033] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 344 processes
[1591768153] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 344 processes
[1591768273] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1591768606] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/06/10 11:20:28 [error] 12486#0: *94092461 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1591768726] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591768809] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.6% (8672952 kB) free!
[1591768929] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.7% (8245764 kB) free!
[1591769049] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.0% (7870880 kB) free!
[1591769914] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 17.8% (6245456 kB) free!
[1591770034] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 16.2% (5684864 kB) free!
[1591770154] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 14.6% (5102996 kB) free!
[1591770249] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3673484 kB) free!
[1591770401] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1591770521] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1591770521] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 13.9% (3400332 kB) free!
[1591770641] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 12.4% (3038076 kB) free!
[1591770673] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1591770673] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 356 processes
[1591770755] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.3% (31613324 kB) free.
[1591770761] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 11.0% (2686612 kB) free!
[1591770917] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.2% (16220064 kB) free.
[1591771157] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 17.7% (4364224 kB) free!
[1591771273] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1591771277] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 16.1% (3977256 kB) free!
[1591771332] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.58, 5.38, 5.14
[1591771361] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (411084 kB) free!
[1591771397] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 14.5% (3577144 kB) free!
[1591771452] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.69, 5.75, 5.29
[1591771572] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.22, 5.22, 5.16
[1591771842] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.14, 3.74, 4.67
[1591771997] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.1% (22734636 kB) free.
[1591772473] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1591772645] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(11) < 103.13.97.241 [10/Jun/2020:12:31:35 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1591772766] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591773877] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 17.1% (5588920 kB) free!
[1591773997] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 16.3% (5339576 kB) free!
[1591774050] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(5) < Channel error on connection <0.14004.4057> (10.147.212.158:38206 -> 10.147.212.203:5672, vhost: '/', user: 'guest'), channel 1:
[1591774117] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 15.6% (5112300 kB) free!
[1591774170] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1591774717] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 20.9% (6845560 kB) free.
[1591775317] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 17.2% (5640444 kB) free!
[1591775437] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 16.5% (5391864 kB) free!
[1591775557] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 15.7% (5134316 kB) free!
[1591775806] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9307676 kB) free!
[1591775926] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (9042056 kB) free!
[1591776046] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (8750660 kB) free!
[1591776074] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 363 processes
[1591776366] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4) < 49.50.86.152 [10/Jun/2020:13:32:27 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1591776486] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591776642] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.48, 5.59, 5.04
[1591776674] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1591776757] SERVICE FLAPPING ALERT: mrq4;Memory;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1591776757] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (2574652 kB) free!
[1591776762] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.47, 6.09, 5.30
[1591776794] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1591776882] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.61, 6.15, 5.43
[1591776914] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1591776917] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.9% (2902392 kB) free!
[1591777037] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 10.9% (2450780 kB) free!
[1591777157] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1953860 kB) free!
[1591777572] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.40, 3.31, 4.33
[1591777957] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 12.0% (3925576 kB) free!
[1591777994] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 761 processes
[1591778056] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 753 processes
[1591778114] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 754 processes
[1591778177] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 754 processes
[1591778286] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 203.212.70.170 [10/Jun/2020:14:05:22 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1591778297] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 347 processes
[1591778406] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591778446] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4264424 kB) free!
[1591778557] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (3165844 kB) free!
[1591778714] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 391 processes
[1591779069] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1591781321] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1591781406] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 185.255.8.58 [10/Jun/2020:15:00:03 +0530] TCP 502 0 0 0.001 "10.147.212.154:13001, 10.147.212.142:13001, 10.147.212.155:13001, 10.147.212.143:13001" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 0.001, 0.000, 0.000"
[1591781441] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 337 processes
[1591781526] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591781561] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1591781957] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.8% (16345704 kB) free.
[1591782314] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 341 processes
[1591782434] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 341 processes
[1591782555] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 342 processes
[1591783361] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 337 processes
[1591783481] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1591783755] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1591789439] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 605 processes
[1591789512] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (4453320 kB) free!
[1591789559] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 357 processes
[1591789926] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(8) < 203.212.70.170 [10/Jun/2020:17:12:33 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1591789957] SERVICE FLAPPING ALERT: mrq4;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1591790046] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1591790112] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175608 kB) free!
[1591790357] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.8% (3327008 kB) free!
[1591790477] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.7% (3067688 kB) free!
[1591790597] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.1% (2717860 kB) free!
[1591791197] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (1321380 kB) free!
[1591792398] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1591792398] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.7% (16322504 kB) free.
[1591792512] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.6% (2593100 kB) free!
[1591793112] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1591793112] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (203536 kB) free!
[1591794312] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 27.4% (6705832 kB) free.
[1591794555] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1591794912] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.6% (3327268 kB) free!
[1591795032] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.9% (3161864 kB) free!
[1591795152] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 12.2% (2986640 kB) free!
[1591795752] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2200788 kB) free!
[1591797881] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1591798001] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1591801400] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.1% (3613500 kB) free!
[1591801520] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.3% (3437920 kB) free!
[1591801640] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.9% (3118308 kB) free!
[1591802240] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1963156 kB) free!
[1591802529] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.2% (3716244 kB) free!
[1591803158] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 68.2% (22334412 kB) free.
[1591803440] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.3% (16230764 kB) free.
[1591804929] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2164252 kB) free!
[1591806882] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 3.87, 4.42, 4.93
[1591807152] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1591812600] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 44.08, 41.79, 43.44