[1553711413] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553711474] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3874708 kB) free!
[1553711533] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553712074] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.5% (8108384 kB) free!
[1553712674] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11385948 kB) free.
[1553712733] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1553713274] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.1% (4998872 kB) free!
[1553713333] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553713394] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (3553344 kB) free!
[1553713453] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553713513] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 48.9% (7945644 kB) free.
[1553713514] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11392048 kB) free.
[1553713573] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1553713671] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.3% (3773812 kB) free!
[1553713791] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.3% (3772524 kB) free!
[1553713912] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.3% (3773272 kB) free!
[1553714114] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.4% (6090040 kB) free!
[1553714173] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553714235] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (4816804 kB) free!
[1553714293] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553714355] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3241260 kB) free!
[1553714414] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553714955] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.0% (7401816 kB) free!
[1553715555] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11396416 kB) free.
[1553715614] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1553716155] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4498508 kB) free!
[1553716214] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553716275] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (2694872 kB) free!
[1553716334] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553716395] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11389328 kB) free.
[1553716934] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553716995] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.5% (5680292 kB) free!
[1553717054] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553717115] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (4376268 kB) free!
[1553717174] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553717235] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (2505148 kB) free!
[1553717835] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.0% (6882812 kB) free!
[1553718375] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553718435] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11386064 kB) free.
[1553718810] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (2202368 kB) free!
[1553718975] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553719036] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (4208572 kB) free!
[1553719095] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553719156] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (2220540 kB) free!
[1553719215] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553719276] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11388764 kB) free.
[1553719410] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.8% (6825132 kB) free!
[1553719815] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553719876] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.9% (5372936 kB) free!
[1553719935] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553719996] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (4100624 kB) free!
[1553720010] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (1929204 kB) free!
[1553720055] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553720116] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11382432 kB) free.
[1553720610] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.4% (7113460 kB) free!
[1553720716] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6472284 kB) free!
[1553720836] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.5% (5170340 kB) free!
[1553720956] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3830580 kB) free!
[1553721210] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (1896876 kB) free!
[1553721256] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1553721556] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.5% (8111168 kB) free!
[1553721810] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.6% (5212028 kB) free!
[1553721856] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553721976] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553722096] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1553722156] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11386908 kB) free.
[1553722696] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553722757] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.2% (5003248 kB) free!
[1553722816] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553722877] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (3564752 kB) free!
[1553722936] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553722997] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11390624 kB) free.
[1553723010] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3156824 kB) free!
[1553723597] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.6% (6199136 kB) free!
[1553723610] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.5% (6137056 kB) free!
[1553723717] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (4907948 kB) free!
[1553723837] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3408220 kB) free!
[1553724137] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1553724437] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.4% (7585256 kB) free!
[1553724737] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553724810] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3864888 kB) free!
[1553724857] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553724977] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553725037] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11337996 kB) free.
[1553725410] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.3% (7519236 kB) free!
[1553725577] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553725637] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4601392 kB) free!
[1553725697] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553725757] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (2934312 kB) free!
[1553725817] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1553725877] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11322904 kB) free.
[1553726010] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (2696124 kB) free!
[1553726417] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553726478] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.8% (5827744 kB) free!
[1553726538] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553726598] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (4558876 kB) free!
[1553726610] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.2% (6016060 kB) free!
[1553726658] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553726718] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (2848504 kB) free!
[1553727211] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (2279872 kB) free!
[1553727318] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.3% (7070948 kB) free!
[1553727811] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.5% (8137392 kB) free!
[1553727858] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553727918] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11323396 kB) free.
[1553728458] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553728518] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (4333480 kB) free!
[1553728578] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553728638] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (2447708 kB) free!
[1553728698] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553728758] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11332128 kB) free.
[1553729298] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553729358] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.1% (5481864 kB) free!
[1553729419] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553729478] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (4165904 kB) free!
[1553729539] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553729598] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (2166744 kB) free!
[1553730199] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.4% (6601028 kB) free!
[1553730739] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1553730799] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11332712 kB) free.
[1553731339] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553731399] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (3770448 kB) free!
[1553731459] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553731475] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 48.76, 24.69, 10.25
[1553731519] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11319436 kB) free.
[1553731579] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1553731595] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 50.14, 33.19, 15.11
[1553731715] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 50.88, 39.09, 19.47
[1553732119] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.8% (6301508 kB) free!
[1553732179] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553732239] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.2% (5006012 kB) free!
[1553732299] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553732315] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.89, 6.71, 11.53
[1553732359] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3594528 kB) free!
[1553732420] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1553732959] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.8% (7786472 kB) free!
[1553733515] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.54, 0.68, 3.60
[1553733709] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 43.66, 28.03, 13.10
[1553733811] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (2119340 kB) free!
[1553733829] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 43.87, 33.25, 16.82
[1553733949] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 43.89, 36.80, 20.12
[1553734159] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11330156 kB) free.
[1553734411] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.1% (6447196 kB) free!
[1553734549] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.55, 7.62, 12.82
[1553734759] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3123400 kB) free!
[1553734820] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553734879] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11311164 kB) free.
[1553735420] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553735480] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.2% (6525564 kB) free!
[1553735540] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553735586] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 25.60, 10.02, 4.09
[1553735600] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.7% (5250048 kB) free!
[1553735606] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.07, 10.73, 8.65
[1553735660] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553735706] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 29.03, 16.49, 7.17
[1553735720] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3940172 kB) free!
[1553735726] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.51, 7.34, 7.66
[1553735743] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.9% (4385992 kB) free!
[1553735749] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.71, 1.98, 4.86
[1553735826] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 29.32, 20.74, 9.86
[1553735846] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.34, 5.00, 6.76
[1553735863] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.8% (3883184 kB) free!
[1553735983] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 26.4% (6480904 kB) free.
[1553736211] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (3450412 kB) free!
[1553736320] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.5% (8121548 kB) free!
[1553736446] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.44, 1.10, 3.75
[1553736811] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.9% (8331828 kB) free!
[1553736860] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1553736920] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11326496 kB) free.
[1553737026] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.94, 1.19, 3.76
[1553737460] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553737520] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.1% (4990348 kB) free!
[1553737580] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553737640] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (3593452 kB) free!
[1553737700] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553737760] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11336420 kB) free.
[1553738360] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.5% (6146024 kB) free!
[1553738480] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (4818808 kB) free!
[1553738600] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3230548 kB) free!
[1553738900] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1553739201] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.1% (7455688 kB) free!
[1553739500] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553739620] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1553739801] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11332772 kB) free.
[1553739811] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (3455956 kB) free!
[1553740220] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553740340] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553740401] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (4506088 kB) free!
[1553740411] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.0% (8383800 kB) free!
[1553740460] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553740521] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (2740968 kB) free!
[1553740641] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11327056 kB) free.
[1553741011] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4386672 kB) free!
[1553741241] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.5% (5652368 kB) free!
[1553741361] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (4344416 kB) free!
[1553741481] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (2473820 kB) free!
[1553741611] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.4% (8062448 kB) free!
[1553741660] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1553742081] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.9% (6846648 kB) free!
[1553742211] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3159608 kB) free!
[1553742260] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1553742380] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553742501] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1553742681] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11321844 kB) free.
[1553742812] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.9% (5391752 kB) free!
[1553743100] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553743220] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553743281] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4120944 kB) free!
[1553743340] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1553743401] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (2122716 kB) free!
[1553743522] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11323920 kB) free.
[1553744012] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4873904 kB) free!
[1553744121] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.4% (5133748 kB) free!
[1553744241] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (3860612 kB) free!
[1553744361] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11321212 kB) free.
[1553744612] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.6% (7198964 kB) free!
[1553744961] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (3849456 kB) free!
[1553745081] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11312200 kB) free.
[1553745681] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (3622916 kB) free!
[1553745801] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11307384 kB) free.
[1553745812] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (3982504 kB) free!
[1553746401] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3432656 kB) free!
[1553746412] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.8% (6308384 kB) free!
[1553746521] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11306168 kB) free.
[1553747121] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.5% (3226152 kB) free!
[1553747242] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11303152 kB) free.
[1553747612] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2945052 kB) free!
[1553747841] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (3013380 kB) free!
[1553747961] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11305088 kB) free.
[1553748212] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.1% (5468828 kB) free!
[1553748561] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (2708212 kB) free!
[1553748681] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11081252 kB) free.
[1553749281] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (2442760 kB) free!
[1553749340] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1553749401] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10941472 kB) free.
[1553749412] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (2284072 kB) free!
[1553749471] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 51.8% (25519300 kB) free.
[1553749574] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.0% (21212992 kB) free.
[1553749616] SERVICE ALERT: sig6-vm1;Memory;OK;HARD;1;OK - 94.8% (6975540 kB) free.
[1553749940] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553750001] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (2092728 kB) free!
[1553750060] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1553750121] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.3% (10502748 kB) free.
[1553750661] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553750721] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (1157580 kB) free!
[1553750781] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553750841] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.8% (10747848 kB) free.
[1553751381] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553751501] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553751513] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (148528 kB) free!
[1553751633] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (161472 kB) free!
[1553751753] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161416 kB) free!
[1553751983] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175424 kB) free!
[1553752074] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (3675124 kB) free!
[1553752101] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553752104] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190560 kB) free!
[1553752194] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11521404 kB) free.
[1553752221] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1553752224] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169568 kB) free!
[1553752794] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (3153340 kB) free!
[1553752821] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553752824] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.9% (21196124 kB) free.
[1553752830] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (2546176 kB) free!
[1553752834] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (2471936 kB) free!
[1553752890] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 41.8% (20616792 kB) free.
[1553752941] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553752945] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.6% (14890372 kB) free.
[1553753093] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 64232 MB (28% inode=99%):
[1553753542] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553753662] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1553753671] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.1% (7916436 kB) free!
[1553753791] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.3% (12966000 kB) free.
[1553754094] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.4% (5111300 kB) free!
[1553754214] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (3815200 kB) free!
[1553754262] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 301 processes
[1553754334] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11533632 kB) free.
[1553754382] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553754391] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.5% (8107288 kB) free!
[1553754511] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.1% (12866384 kB) free.
[1553754934] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4142916 kB) free!
[1553754982] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553755055] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (2292376 kB) free!
[1553755102] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1553755175] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11059244 kB) free.
[1553755702] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553755775] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (2497216 kB) free!
[1553755822] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553755895] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11141420 kB) free.
[1553756423] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1553756495] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (2588612 kB) free!
[1553756543] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1553756543] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1553756615] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11176964 kB) free.
[1553757143] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553757215] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (2668900 kB) free!
[1553757263] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1553757335] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11178184 kB) free.
[1553757863] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553757935] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3495944 kB) free!
[1553757983] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553758055] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11510264 kB) free.
[1553758103] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553758655] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4137388 kB) free!
[1553758775] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (2184564 kB) free!
[1553758896] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10648940 kB) free.
[1553759495] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (4006524 kB) free!
[1553759615] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.4% (10569012 kB) free.
[1553760815] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (4610936 kB) free!
[1553760935] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (2969840 kB) free!
[1553761055] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11414584 kB) free.
[1553761111] SERVICE FLAPPING ALERT: db1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1553761655] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3525088 kB) free!
[1553761775] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11470692 kB) free.
[1553762375] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (4246000 kB) free!
[1553762495] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (2381636 kB) free!
[1553762615] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10929664 kB) free.
[1553763216] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (2969084 kB) free!
[1553763336] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11349344 kB) free.
[1553763936] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (2794312 kB) free!
[1553764056] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11166128 kB) free.
[1553764656] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (2902580 kB) free!
[1553764776] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1553764776] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11262472 kB) free.
[1553765376] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (2574196 kB) free!
[1553765496] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.4% (11016960 kB) free.
[1553766096] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (3009260 kB) free!
[1553766216] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11510388 kB) free.
[1553766816] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (3661636 kB) free!
[1553766937] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11494392 kB) free.
[1553767536] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (3628120 kB) free!
[1553767656] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11477920 kB) free.
[1553768256] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (3559996 kB) free!
[1553768376] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11487228 kB) free.
[1553768976] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (3343388 kB) free!
[1553769096] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11463228 kB) free.
[1553769163] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1553769283] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1553769403] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1553769696] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (2942820 kB) free!
[1553769816] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11312144 kB) free.
[1553770416] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (2822072 kB) free!
[1553770536] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11288680 kB) free.
[1553771137] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (2843988 kB) free!
[1553771257] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11293944 kB) free.
[1553771857] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (2599820 kB) free!
[1553771903] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1553771977] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.5% (11100524 kB) free.
[1553772503] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553772577] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (3645168 kB) free!
[1553772623] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1553772697] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11458340 kB) free.
[1553772743] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1553772956] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 316 processes
[1553773240] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 363 processes
[1553773297] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4094952 kB) free!
[1553773417] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (2187964 kB) free!
[1553773537] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.5% (10587616 kB) free.
[1553774737] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6540128 kB) free!
[1553774857] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.1% (5453584 kB) free!
[1553774978] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4174504 kB) free!
[1553775143] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1553775578] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.5% (9098540 kB) free!
[1553775743] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553775863] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553775983] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1553775983] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1553776178] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.3% (11467892 kB) free.
[1553776778] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5604768 kB) free!
[1553776898] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (4168600 kB) free!
[1553777018] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1553777018] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (2143960 kB) free!
[1553777113] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3232932 kB) free!
[1553777233] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.9% (3229804 kB) free!
[1553777353] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.9% (3226512 kB) free!
[1553777784] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 315 processes
[1553778218] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.2% (6027024 kB) free!
[1553778384] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1553778504] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553778624] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 316 processes
[1553778818] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11299384 kB) free.
[1553779224] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1553779344] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553779418] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (2954052 kB) free!
[1553779464] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1553779538] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11456432 kB) free.
[1553780064] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553780139] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.8% (5829660 kB) free!
[1553780185] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1553780259] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (4525988 kB) free!
[1553780305] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553780379] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (2704500 kB) free!
[1553780979] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.3% (6559828 kB) free!
[1553781356] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 320 processes
[1553782105] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1553782179] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11444388 kB) free.
[1553782705] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553782779] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.2% (5036040 kB) free!
[1553782825] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1553782899] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (3531712 kB) free!
[1553782945] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1553783019] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11456840 kB) free.
[1553783619] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.5% (6159308 kB) free!
[1553783739] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (4841884 kB) free!
[1553783860] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (3236716 kB) free!
[1553784145] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1553784460] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.2% (7481860 kB) free!
[1553784745] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553784865] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553784985] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553785060] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11452616 kB) free.
[1553785585] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553785660] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (3627144 kB) free!
[1553785705] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553785780] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11442200 kB) free.
[1553785825] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1553786380] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (4416236 kB) free!
[1553786426] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553786500] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.1% (2508880 kB) free!
[1553786546] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1553786620] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11341168 kB) free.
[1553787146] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553787220] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (2781808 kB) free!
[1553787266] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1553787340] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11131260 kB) free.
[1553787866] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553787940] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (2534220 kB) free!
[1553787986] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553788061] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11446608 kB) free.
[1553788586] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1553788660] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 11.4% (5632224 kB) free!
[1553788706] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553788780] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (4323856 kB) free!
[1553788827] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553788900] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (2372896 kB) free!
[1553789500] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.8% (6805524 kB) free!
[1553790027] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553790041] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 335 processes
[1553790100] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11440824 kB) free.
[1553790627] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1553790700] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (3928244 kB) free!
[1553790747] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553790820] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11434636 kB) free.
[1553790867] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1553791420] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6448240 kB) free!
[1553791467] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553791540] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 10.4% (5128260 kB) free!
[1553791587] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553791660] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3717656 kB) free!
[1553791708] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1553792260] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.1% (7929216 kB) free!
[1553792861] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11439356 kB) free.
[1553792908] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1553793461] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.5% (5176060 kB) free!
[1553793508] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553793581] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (3833228 kB) free!
[1553793628] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553793701] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11434944 kB) free.
[1553793748] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1553794301] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.8% (6789880 kB) free!
[1553794348] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1553794421] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 11.4% (5597720 kB) free!
[1553794468] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1553794541] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4323092 kB) free!
[1553794589] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1553795141] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.7% (9205488 kB) free!
[1553795156] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 321 processes
[1553795442] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 337 processes
[1553795741] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11437332 kB) free.
[1553795756] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 323 processes
[1553795789] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553796341] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.2% (5995800 kB) free!
[1553796389] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1553796461] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (4670224 kB) free!
[1553796509] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1553796582] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (2910964 kB) free!
[1553796629] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1553797182] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.6% (7201100 kB) free!
[1553797229] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1553797349] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1553797470] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1553797782] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.2% (11435696 kB) free.