[1588789828] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (9066056 kB) free!
[1588789948] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5341548 kB) free!
[1588790548] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (18469756 kB) free!
[1588790919] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (1429176 kB) free!
[1588791039] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (1575236 kB) free!
[1588791148] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33453892 kB) free.
[1588791159] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.7% (1745348 kB) free!
[1588791748] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11769912 kB) free!
[1588791830] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (178936 kB) free!
[1588791876] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7924768 kB) free!
[1588791996] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (4272848 kB) free!
[1588792359] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (159296 kB) free!
[1588792596] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17248624 kB) free!
[1588792730] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 76.0% (18620356 kB) free.
[1588793196] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31443632 kB) free.
[1588793796] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10558812 kB) free!
[1588793916] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.3% (6929328 kB) free!
[1588794036] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37613444 kB) free.
[1588794159] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.2% (1977944 kB) free!
[1588794636] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15239160 kB) free!
[1588794756] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (11316524 kB) free!
[1588794876] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (7722812 kB) free!
[1588795101] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 12.1% (5970620 kB) free!
[1588795221] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 12.1% (5969048 kB) free!
[1588795341] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 12.1% (5962032 kB) free!
[1588795359] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (791160 kB) free!
[1588795476] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19094084 kB) free!
[1588796076] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32054004 kB) free.
[1588796135] SERVICE FLAPPING ALERT: sigqscluster1;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1588796676] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10424284 kB) free!
[1588796797] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6655624 kB) free!
[1588796917] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36181136 kB) free.
[1588797141] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 43.1% (21234548 kB) free.
[1588797159] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.9% (4046932 kB) free.
[1588797517] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13899464 kB) free!
[1588797637] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (10205964 kB) free!
[1588797757] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (6763488 kB) free!
[1588798357] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17474664 kB) free!
[1588798957] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28391084 kB) free.
[1588799172] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 91.2% (44960108 kB) free.
[1588799557] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7794616 kB) free!
[1588799588] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.0% (45812872 kB) free.
[1588799677] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (13924600 kB) free!
[1588799797] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32910200 kB) free.
[1588800397] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11042520 kB) free!
[1588800517] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7275384 kB) free!
[1588800637] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37243248 kB) free.
[1588800758] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1588800938] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(493) < 2020/05/07 03:01:32 [error] 22129#0: *87046662 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1588801045] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(381) < 49.50.86.152 [07/May/2020:03:00:37 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1588801059] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588801106] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2925) < 203.212.70.170 [07/May/2020:03:01:42 +0530] TCP 502 0 0 60.000 "10.147.212.143:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1588801165] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588801226] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588801237] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14334132 kB) free!
[1588801357] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (10424896 kB) free!
[1588801477] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37265840 kB) free.
[1588802078] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16665028 kB) free!
[1588802198] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (13072832 kB) free!
[1588802318] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (9866560 kB) free!
[1588802918] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19275888 kB) free!
[1588803518] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26891916 kB) free.
[1588804095] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.7% (9973056 kB) free.
[1588804695] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (249744 kB) free!
[1588804718] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.2% (16080428 kB) free!
[1588804815] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (285976 kB) free!
[1588804838] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.7% (15473716 kB) free!
[1588804935] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (253660 kB) free!
[1588804958] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (33980028 kB) free.
[1588805558] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10757708 kB) free!
[1588805678] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (5642460 kB) free!
[1588805798] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35226556 kB) free.
[1588806398] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11649044 kB) free!
[1588806518] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6821196 kB) free!
[1588806638] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35349820 kB) free.
[1588806735] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 39.3% (9642664 kB) free.
[1588807239] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12482792 kB) free!
[1588807359] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7237348 kB) free!
[1588807359] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1588807479] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35099812 kB) free.
[1588807935] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (227660 kB) free!
[1588808055] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (236932 kB) free!
[1588808079] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12563312 kB) free!
[1588808175] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1588808175] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (222620 kB) free!
[1588808199] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37069800 kB) free.
[1588808799] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18010560 kB) free!
[1588808919] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.8% (14266436 kB) free!
[1588809039] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38087636 kB) free.
[1588809276] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 75.8% (18576916 kB) free.
[1588809639] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18119872 kB) free!
[1588809716] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 55.4% (13669528 kB) free.
[1588809725] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 53.9% (13210640 kB) free.
[1588809759] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.9% (14337404 kB) free!
[1588809879] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36169880 kB) free.
[1588810479] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17290556 kB) free!
[1588810599] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37403012 kB) free.
[1588811199] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16762084 kB) free!
[1588811319] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36210616 kB) free.
[1588811919] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14913284 kB) free!
[1588812040] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36631832 kB) free.
[1588812639] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16398516 kB) free!
[1588812759] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38342360 kB) free.
[1588813359] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19232288 kB) free!
[1588813479] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.8% (15579376 kB) free!
[1588813599] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36447560 kB) free.
[1588814199] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16414700 kB) free!
[1588814319] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37999336 kB) free.
[1588814331] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 18.8% (4597740 kB) free!
[1588814391] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 18.8% (4597172 kB) free!
[1588814451] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 18.8% (4596216 kB) free!
[1588814511] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 18.7% (4589984 kB) free!
[1588814516] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190756 kB) free!
[1588814636] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178040 kB) free!
[1588814756] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187044 kB) free!
[1588814919] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19112464 kB) free!
[1588815039] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.1% (25116080 kB) free!
[1588815159] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35264400 kB) free.
[1588815759] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14960592 kB) free!
[1588815879] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36111536 kB) free.
[1588816479] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (17027044 kB) free!
[1588816599] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37460316 kB) free.
[1588817200] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18119056 kB) free!
[1588817320] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.5% (23099972 kB) free!
[1588817440] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35405828 kB) free.
[1588818040] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16679736 kB) free!
[1588818160] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1588818160] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37147448 kB) free.
[1588818760] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17387132 kB) free!
[1588818880] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.4% (16309768 kB) free!
[1588818975] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1588819000] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34957576 kB) free.
[1588819600] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14255844 kB) free!
[1588819720] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36458400 kB) free.
[1588820320] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15231544 kB) free!
[1588820440] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (36041060 kB) free.
[1588821040] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16232572 kB) free!
[1588821160] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.2% (12157556 kB) free!
[1588821280] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34755100 kB) free.
[1588821880] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13160188 kB) free!
[1588822001] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34908892 kB) free.
[1588822311] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 2.7% (655036 kB) free!
[1588822361] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 352 processes
[1588822481] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 346 processes
[1588822600] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11495156 kB) free!
[1588822601] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1588822681] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 20.66, 15.59, 10.00
[1588822720] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6833480 kB) free!
[1588822801] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 22.02, 20.27, 12.58
[1588822840] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33048320 kB) free.
[1588822921] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 24.63, 22.09, 14.19
[1588823258] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/07 09:11:51 [error] 22133#0: *87063916 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/senderid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/senderid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/senderid"
[1588823378] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588823440] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10171828 kB) free!
[1588823560] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (5888988 kB) free!
[1588823680] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31493252 kB) free.
[1588824125] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4851688 kB) free!
[1588824245] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (4693364 kB) free!
[1588824280] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8722216 kB) free!
[1588824365] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (4493816 kB) free!
[1588824400] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4129160 kB) free!
[1588824520] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30145796 kB) free.
[1588825120] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5490328 kB) free!
[1588825240] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (1307756 kB) free!
[1588825360] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29685016 kB) free.
[1588825778] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/07 09:59:26 [error] 22135#0: *87066850 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1588825898] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/05/07 10:01:28 [error] 22135#0: *87066995 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1588825960] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (390200 kB) free!
[1588826018] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588826080] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (374760 kB) free!
[1588826165] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1705636 kB) free!
[1588826201] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38139560 kB) free.
[1588826201] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 347 processes
[1588826325] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1588826445] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 348 processes
[1588826565] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1588826618] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/05/07 10:12:44 [error] 22135#0: *87067925 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1588826676] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166912 kB) free!
[1588826738] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588826796] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186928 kB) free!
[1588826800] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8640912 kB) free!
[1588826916] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192140 kB) free!
[1588826920] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4351020 kB) free!
[1588827040] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (17063772 kB) free!
[1588827166] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1588827401] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1588827938] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2020/05/07 10:29:25 [error] 22132#0: *87069690 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /temtapp/index.php/dashboard HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/dashboard", host: "altsmsc.tesync.net"
[1588828058] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588828840] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33584744 kB) free.
[1588828880] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.1% (3844852 kB) free!
[1588829000] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.1% (3379320 kB) free!
[1588829120] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.8% (2869324 kB) free!
[1588829440] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8954620 kB) free!
[1588829560] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (4415492 kB) free!
[1588829680] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (385172 kB) free!
[1588829720] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (383212 kB) free!
[1588829735] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1588829801] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1588829801] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 349 processes
[1588829855] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1588830223] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 55.3% (13545308 kB) free.
[1588830880] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1588830880] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19631256 kB) free!
[1588831001] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1588831480] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30616216 kB) free.
[1588832080] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5854316 kB) free!
[1588832200] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1735716 kB) free!
[1588832321] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31065152 kB) free.
[1588832426] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1000) < 103.255.146.164 [07/May/2020:11:44:57 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588832546] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588832802] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1588832920] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5539324 kB) free!
[1588833040] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (1112964 kB) free!
[1588833160] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (412104 kB) free!
[1588833321] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.0% (16175532 kB) free.
[1588834360] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27840684 kB) free.
[1588834362] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3196076 kB) free!
[1588834482] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.6% (4808692 kB) free.
[1588834960] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.9% (3823260 kB) free!
[1588835023] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.6% (4060720 kB) free!
[1588835080] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34367444 kB) free.
[1588835144] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.0% (3917040 kB) free!
[1588835264] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 15.2% (3719100 kB) free!
[1588835680] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10221400 kB) free!
[1588835800] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5450248 kB) free!
[1588835920] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (995924 kB) free!
[1588836188] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9702616 kB) free!
[1588836308] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.2% (9441848 kB) free!
[1588836428] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.7% (9208284 kB) free!
[1588836464] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2029656 kB) free!
[1588836520] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15255624 kB) free!
[1588836746] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(167) < 103.255.146.164 [07/May/2020:13:01:40 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588836867] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588836882] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.0% (2433076 kB) free!
[1588837002] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.8% (1914600 kB) free!
[1588837120] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29116692 kB) free.
[1588837122] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.0% (2755928 kB) free!
[1588837246] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 24.4% (3973048 kB) free.
[1588837720] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (2976428 kB) free!
[1588837840] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (404392 kB) free!
[1588837961] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37554648 kB) free.
[1588838322] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3358756 kB) free.
[1588838560] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9597528 kB) free!
[1588838680] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4774516 kB) free!
[1588838800] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (414904 kB) free!
[1588838828] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4638040 kB) free!
[1588838922] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (1593416 kB) free!
[1588839042] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.2% (3763464 kB) free.
[1588839321] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 10.3% (2304728 kB) free!
[1588839400] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16699560 kB) free!
[1588839441] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (1609164 kB) free!
[1588839561] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (908124 kB) free!
[1588839876] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.27322.107> (10.147.212.147:37980 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1588839996] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1588840000] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36819904 kB) free.
[1588840242] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (2801608 kB) free!
[1588840362] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.9% (2739384 kB) free!
[1588840482] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.6% (2698604 kB) free!
[1588840600] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9269888 kB) free!
[1588840659] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/07 14:02:56 [error] 22134#0: *87096684 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/uuid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/uuid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/uuid"
[1588840720] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4751080 kB) free!
[1588840779] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588840840] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (432876 kB) free!
[1588840846] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (174860 kB) free!
[1588840966] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (167636 kB) free!
[1588841066] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(819) < 103.255.146.164 [07/May/2020:14:13:32 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588841086] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (208748 kB) free!
[1588841172] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (8616300 kB) free!
[1588841186] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588841292] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8302604 kB) free!
[1588841412] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 16.2% (8007200 kB) free!
[1588841440] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.7% (19381872 kB) free!
[1588842040] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397916 kB) free!
[1588842215] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 13.4% (4678156 kB) free!
[1588842335] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 11.4% (4009304 kB) free!
[1588842455] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (3419196 kB) free!
[1588842524] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4696296 kB) free!
[1588842644] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 17.4% (4284492 kB) free!
[1588842680] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 20.0% (26403644 kB) free!
[1588842765] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 16.2% (3998568 kB) free!
[1588842800] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 26.7% (35228800 kB) free.
[1588842882] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1588842882] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (1328036 kB) free!
[1588842986] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(829) < 103.255.146.164 [07/May/2020:14:44:08 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1588843002] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1588843107] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588843212] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4560568 kB) free!
[1588843240] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.7% (15456048 kB) free!
[1588843482] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.9% (2267020 kB) free!
[1588843596] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.19526.234> (10.147.212.158:52840 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1588843716] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1588843779] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/05/07 14:58:47 [error] 22120#0: *87103621 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1588843840] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (1529124 kB) free!
[1588843899] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588843965] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1714232 kB) free!
[1588844082] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.1% (3266640 kB) free.
[1588844256] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.9% (31847576 kB) free.
[1588844306] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(153) < 185.255.8.58 [07/May/2020:15:00:04 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1588844361] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 73.2% (16442600 kB) free.
[1588844426] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588844441] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.9% (22315632 kB) free!
[1588844682] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.4% (1683728 kB) free!
[1588844802] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.3% (2165260 kB) free!
[1588844855] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1588844922] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.0% (2593716 kB) free!
[1588844975] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1588845041] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (411136 kB) free!
[1588845165] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.5% (22825172 kB) free.
[1588845522] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1318684 kB) free!
[1588845641] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27061988 kB) free.
[1588846122] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (2847144 kB) free!
[1588846241] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (387500 kB) free!
[1588846361] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (397208 kB) free!
[1588846481] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.0% (43563880 kB) free.
[1588847000] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.4% (20283948 kB) free!
[1588847081] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10415700 kB) free!
[1588847120] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.6% (28417168 kB) free.
[1588847201] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4892708 kB) free!
[1588847321] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384064 kB) free!
[1588847322] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (985088 kB) free!
[1588847921] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.2% (25247812 kB) free!
[1588847922] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.1% (2457652 kB) free!
[1588848521] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (1047372 kB) free!
[1588849121] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.9% (22219384 kB) free!
[1588849721] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42545960 kB) free.
[1588849826] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(434) < 103.255.146.164 [07/May/2020:16:31:47 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1588849947] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588850321] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12357776 kB) free!
[1588850323] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.1% (3600228 kB) free.
[1588850442] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7463868 kB) free!
[1588850562] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2820548 kB) free!
[1588850923] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3235216 kB) free!
[1588851043] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (3167164 kB) free!
[1588851162] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.2% (25316636 kB) free!
[1588851163] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (3135716 kB) free!
[1588851321] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.6% (25781636 kB) free!
[1588851441] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.7% (18011372 kB) free!
[1588851561] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.1% (29185896 kB) free.
[1588851561] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.8% (3316096 kB) free!
[1588851681] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.6% (2823172 kB) free!
[1588851762] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (384940 kB) free!
[1588851763] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (240120 kB) free!
[1588851802] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.5% (2362844 kB) free!
[1588852362] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21441760 kB) free!
[1588852402] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (200060 kB) free!
[1588852487] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 11.8% (1914752 kB) free!
[1588852922] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 31.70, 18.32, 13.27
[1588852962] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402572 kB) free!
[1588852963] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.2% (4907336 kB) free.
[1588853361] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 10.2% (13504884 kB) free!
[1588853375] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1588853481] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.6% (27124192 kB) free.
[1588853495] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1588853522] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.44, 10.03, 11.56
[1588853562] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.7% (19349140 kB) free!
[1588853563] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (1437092 kB) free!
[1588853683] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (3241304 kB) free!
[1588853803] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.7% (4984788 kB) free.
[1588854162] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (373724 kB) free!
[1588854403] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.0% (2281120 kB) free!
[1588854523] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.0% (1953668 kB) free!
[1588854643] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1066116 kB) free!
[1588854681] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 11.0% (14495876 kB) free!
[1588854762] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18115144 kB) free!
[1588854801] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 16.4% (21662356 kB) free!
[1588854803] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1588854803] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 73.5% (16498432 kB) free.
[1588854921] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (12877996 kB) free!
[1588855244] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 30.9% (5022188 kB) free.
[1588855362] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (42050924 kB) free.
[1588855487] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (307128 kB) free!
[1588855521] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.8% (22148160 kB) free!
[1588855962] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13426040 kB) free!
[1588856082] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8422296 kB) free!
[1588856202] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (4055320 kB) free!
[1588856444] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.9% (2428680 kB) free!
[1588856564] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.2% (3606748 kB) free.
[1588856802] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.6% (25870036 kB) free!
[1588857403] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2382056 kB) free!
[1588857764] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (1309528 kB) free!
[1588857884] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.9% (2579840 kB) free!
[1588858003] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.7% (24622904 kB) free!
[1588858004] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.2% (3940848 kB) free.
[1588858322] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.47, 3.12, 4.81
[1588858603] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (411452 kB) free!
[1588858922] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.87, 6.00, 5.44
[1588859042] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.22, 6.11, 5.56
[1588859163] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1588859163] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.24, 6.03, 5.59
[1588859203] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17740332 kB) free!
[1588859204] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (1510304 kB) free!
[1588859324] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (3023904 kB) free!
[1588859444] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3751060 kB) free.
[1588859803] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399060 kB) free!
[1588860044] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (909476 kB) free!
[1588860164] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.1% (2449128 kB) free!
[1588860285] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3762904 kB) free.
[1588860321] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 27.4% (36166936 kB) free.
[1588860403] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (16016328 kB) free!
[1588860884] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (519060 kB) free!
[1588861003] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41573220 kB) free.
[1588861004] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.5% (2026588 kB) free!
[1588861124] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.5% (3651384 kB) free.
[1588861603] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13427704 kB) free!
[1588861723] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8363516 kB) free!
[1588861724] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.8% (2569932 kB) free!
[1588861843] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (3608600 kB) free!
[1588861844] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.6% (3994252 kB) free.
[1588862443] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27620128 kB) free.
[1588862444] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (686504 kB) free!
[1588862564] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.0% (2109676 kB) free!
[1588862684] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3521448 kB) free.
[1588863043] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.6% (2095856 kB) free!
[1588863163] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (399904 kB) free!
[1588863283] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41867012 kB) free.
[1588863363] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 4.80, 4.32, 4.97
[1588863883] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14927384 kB) free!
[1588863884] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (1486132 kB) free!
[1588864004] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (9857132 kB) free!
[1588864004] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (2976408 kB) free!
[1588864010] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1588864124] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (5077268 kB) free!
[1588864124] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.6% (3675976 kB) free.
[1588864403] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.7% (3748796 kB) free!
[1588864523] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.7% (3523584 kB) free!
[1588864643] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.8% (3317448 kB) free!
[1588864724] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23082764 kB) free!
[1588864724] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.3% (374904 kB) free!
[1588864844] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.2% (1815908 kB) free!
[1588864965] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3474056 kB) free.
[1588865324] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392308 kB) free!
[1588865499] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(41) < 2020/05/07 21:00:30 [error] 22130#0: *87135277 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1588865619] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1588865844] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 73.3% (16458216 kB) free.
[1588866165] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (1276496 kB) free!
[1588866285] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (1349092 kB) free!
[1588866405] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (2578116 kB) free!
[1588866524] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.6% (24498392 kB) free!
[1588867005] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3393444 kB) free.
[1588867124] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (373856 kB) free!
[1588867724] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31959264 kB) free.
[1588868205] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.8% (2574036 kB) free!
[1588868324] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (4964612 kB) free!
[1588868325] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.5% (3336248 kB) free.
[1588868444] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (511516 kB) free!
[1588868564] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (407436 kB) free!
[1588868925] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.5% (2358136 kB) free!
[1588869045] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (3191520 kB) free!
[1588869095] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1588869164] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18224092 kB) free!
[1588869165] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3173144 kB) free!
[1588869216] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1588869764] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (431832 kB) free!
[1588870364] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14740344 kB) free!
[1588870365] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.5% (402872 kB) free!
[1588870536] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1528 processes
[1588870964] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396936 kB) free!
[1588870966] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.5% (3493112 kB) free.
[1588871136] SERVICE ALERT: sigqscluster1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1752 processes
[1588871163] SERVICE FLAPPING ALERT: sigqscluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1588871565] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14715780 kB) free!
[1588872165] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38923872 kB) free.
[1588872765] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12927620 kB) free!
[1588872885] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8750924 kB) free!
[1588873005] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (4895188 kB) free!
[1588873605] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (22810216 kB) free!
[1588874205] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (2538160 kB) free!
[1588874805] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17539324 kB) free!
[1588875405] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (36033964 kB) free.
[1588876005] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12947600 kB) free!
[1588876125] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (8967724 kB) free!