[1580150167] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1580150172] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.4% (30272416 kB) free.
[1580150619] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187896 kB) free!
[1580150739] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181836 kB) free!
[1580150768] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1580150859] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175168 kB) free!
[1580150888] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1580151488] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1580151608] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1580151729] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1580152329] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1580152630] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 40.8% (10069804 kB) free.
[1580153529] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1580153649] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1580154249] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1580154370] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1580154430] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (201900 kB) free!
[1580154550] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186192 kB) free!
[1580154670] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198456 kB) free!
[1580154987] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.2% (7009372 kB) free!
[1580155107] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 14.2% (7008056 kB) free!
[1580155227] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.2% (6991184 kB) free!
[1580156259] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.2% (10104548 kB) free.
[1580156673] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.5% (20607688 kB) free.
[1580156745] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.0% (22794216 kB) free.
[1580156794] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580156798] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/28 01:56:34 [error] 18345#0: *26888214 access forbidden by rule while initializing session, client: 103.209.99.7, server: 0.0.0.0:13000
[1580156918] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(20) < 2020/01/28 01:58:35 [error] 18341#0: *26888288 access forbidden by rule while initializing session, client: 103.209.99.7, server: 0.0.0.0:13000
[1580157038] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(20) < 2020/01/28 02:00:35 [error] 18346#0: *26888368 access forbidden by rule while initializing session, client: 103.209.99.7, server: 0.0.0.0:13000
[1580158484] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172684 kB) free!
[1580158604] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (834896 kB) free!
[1580158724] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1580158724] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180728 kB) free!
[1580160692] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3309) < 172.105.59.173 [28/Jan/2020:03:01:32 +0530] TCP 502 0 0 60.001 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.001, 0.000"
[1580160812] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(60) < 49.50.64.122 [28/Jan/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1580160932] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580161169] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(226) < 2020/01/28 03:01:31 [error] 18360#0: *44576043 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580161201] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(179) < 103.16.101.72 [28/Jan/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1580161289] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580161322] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580163969] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1580165946] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1580166066] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1580167124] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1580168259] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180076 kB) free!
[1580168380] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176980 kB) free!
[1580168500] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194220 kB) free!
[1580168925] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.1% (21731016 kB) free.
[1580172200] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 39.4% (51978700 kB) free.
[1580182866] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1580182986] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1580187028] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4826776 kB) free!
[1580190693] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < webmachine error: path="/api/bindings/%2f/e/amq.direct/q/1181_n"
[1580190813] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1580190933] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(30) < 203.212.70.170 [28/Jan/2020:11:21:51 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1580191053] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580192015] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.4% (40109572 kB) free.
[1580192078] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.5% (21946360 kB) free.
[1580192586] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1580192706] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1580194125] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (416660 kB) free!
[1580194245] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (293256 kB) free!
[1580194365] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170860 kB) free!
[1580195820] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.3% (21287772 kB) free.
[1580198146] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.2% (4714336 kB) free!
[1580198266] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.8% (4608072 kB) free!
[1580198387] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.6% (4550664 kB) free!
[1580199306] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1580199426] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 398 processes
[1580200465] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.5% (21699032 kB) free.
[1580201228] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.4% (3789196 kB) free!
[1580201348] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.9% (3676588 kB) free!
[1580201468] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.4% (3548652 kB) free!
[1580201489] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2020/01/28 14:12:22 [error] 18370#0: *44762382 open() "/usr/share/nginx/html/api/whoami" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /api/whoami HTTP/1.1", host: "qservercluster1.com:8091", referrer: "http://qservercluster1.com:8091/"
[1580201610] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580202068] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1580202068] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (943648 kB) free!
[1580202348] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180628 kB) free!
[1580202369] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1580202406] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1580202469] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197444 kB) free!
[1580202489] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1580202527] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1580202589] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199516 kB) free!
[1580202610] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1580202647] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1580202887] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (776564 kB) free!
[1580203007] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184600 kB) free!
[1580203127] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190136 kB) free!
[1580203210] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1580203493] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.4% (21074588 kB) free.
[1580203551] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.7% (21739480 kB) free.
[1580203719] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.4% (21663156 kB) free.
[1580205402] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 340 processes
[1580205522] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1580205613] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1580205733] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1580205853] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1580205897] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174388 kB) free!
[1580205959] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178388 kB) free!
[1580206017] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180748 kB) free!
[1580206079] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178524 kB) free!
[1580206106] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.1% (20999088 kB) free.
[1580206149] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.6% (21476436 kB) free.
[1580206409] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(26) < 2020/01/28 15:43:17 [error] 18355#0: *44811793 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1580206453] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1580206529] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580207830] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.5% (9620952 kB) free!
[1580207950] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9434656 kB) free!
[1580208070] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.7% (9199588 kB) free!
[1580208213] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < webmachine error: path="/api/bindings/%2f/e/amq.direct/q/1183_n"
[1580208334] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1580208929] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/28 16:25:07 [error] 18370#0: *44837973 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/consolidate/lrnreport HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/consolidate/lrnreport", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/consolidate/lrnreport"
[1580209049] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580209109] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (4403072 kB) free!
[1580209229] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.2% (3741296 kB) free!
[1580209349] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (2771236 kB) free!
[1580209418] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.5% (8616940 kB) free!
[1580209538] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.5% (8617976 kB) free!
[1580209658] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.5% (8617792 kB) free!
[1580209949] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1839716 kB) free!
[1580210355] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.5% (20594028 kB) free.
[1580210396] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 90.1% (44393652 kB) free.
[1580210424] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.4% (40091564 kB) free.
[1580211453] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(361) < 49.50.64.122 [28/Jan/2020:17:01:17 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1580211574] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580212157] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.1% (4227668 kB) free!
[1580212173] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(690) < 103.255.146.164 [28/Jan/2020:17:10:37 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1580212277] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.9% (4168808 kB) free!
[1580212293] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(3473) < 103.255.146.164 [28/Jan/2020:17:21:29 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1580212397] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.7% (4129812 kB) free!
[1580212413] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(8831) < 103.255.146.164 [28/Jan/2020:17:23:33 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1580212735] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4881712 kB) free!
[1580212855] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (4840116 kB) free!
[1580212975] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 21.2% (5206852 kB) free.
[1580212997] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (374200 kB) free!
[1580213306] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.3% (22045080 kB) free.
[1580214213] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580215026] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1580215146] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1580215429] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9623076 kB) free!
[1580215453] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1580215549] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (9548272 kB) free!
[1580215573] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1580215622] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 89.0% (43851124 kB) free.
[1580215693] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1580215975] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (4820552 kB) free!
[1580216096] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (4802916 kB) free!
[1580216216] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (4781668 kB) free!
[1580216456] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.5% (22180424 kB) free.
[1580216893] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1580217013] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1580217133] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1580217559] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.4% (1850508 kB) free!
[1580217679] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.4% (1852684 kB) free!
[1580217799] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 11.4% (1855360 kB) free!
[1580219133] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < webmachine error: path="/api/bindings/%2f/e/amq.direct/q/1184_n"
[1580219254] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1580219533] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1580224109] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (4387288 kB) free!
[1580224229] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.7% (4369372 kB) free!
[1580224349] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.6% (4350756 kB) free!
[1580224649] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/01/28 20:40:14 [error] 18370#0: *44977675 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1580224653] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < webmachine error: path="/api/bindings/%2f/e/amq.direct/q/1185_n"
[1580224769] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580224773] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1580230349] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1580230413] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 203.212.70.170 [28/Jan/2020:22:16:44 +0530] TCP 502 0 0 63.007 "10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "3.007, 60.000, 0.000"
[1580230533] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1580230949] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2254860 kB) free!
[1580234772] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 15.5% (7629744 kB) free!
[1580234892] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (216280 kB) free!
[1580235012] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235472 kB) free!
[1580235200] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7658404 kB) free!
[1580235320] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (395472 kB) free!
[1580235440] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1157908 kB) free!
[1580236120] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.1% (21250568 kB) free.
[1580236178] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.9% (30025892 kB) free.