[1573842971] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 28.3% (6354264 kB) free.
[1573843011] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1573843131] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 365 processes
[1573843181] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2019/11/16 00:08:00 [error] 23638#0: *3682959 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:172/17, bytes from/to upstream:17/172
[1573843301] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/11/16 00:11:31 [error] 23640#0: *3684864 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:224/17, bytes from/to upstream:17/224
[1573843421] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;HARD;3;(1) < 2019/11/16 00:12:46 [error] 23636#0: *3685537 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:184/17, bytes from/to upstream:17/184
[1573843477] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186436 kB) free!
[1573843506] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189836 kB) free!
[1573843535] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184216 kB) free!
[1573843597] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179380 kB) free!
[1573843626] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185288 kB) free!
[1573843656] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197180 kB) free!
[1573843717] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163540 kB) free!
[1573843747] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186560 kB) free!
[1573843776] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171172 kB) free!
[1573844131] SERVICE ALERT: sig4-vm2;CPU_Procs;OK;HARD;1;CPU OK: 228 processes
[1573844609] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174524 kB) free!
[1573844669] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168212 kB) free!
[1573844729] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (184768 kB) free!
[1573844790] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 66.7% (16349444 kB) free.
[1573845545] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.5% (29676644 kB) free.
[1573846104] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (25211024 kB) free!
[1573846224] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.8% (48539508 kB) free.
[1573846421] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573846522] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.5% (8632252 kB) free!
[1573846642] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.5% (8153576 kB) free!
[1573846762] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.7% (7728992 kB) free!
[1573846824] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573847021] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/16 01:09:25 [error] 23641#0: *3715015 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:176/17, bytes from/to upstream:17/176
[1573847141] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573847334] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1573847345] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.9% (26282200 kB) free!
[1573847465] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.3% (26784624 kB) free.
[1573847962] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (2110000 kB) free!
[1573848342] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/11/16 01:33:29 [error] 23645#0: *3727814 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:58/17, bytes from/to upstream:17/58
[1573848462] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573849662] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/16 01:52:34 [error] 23641#0: *3738119 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:195/17, bytes from/to upstream:17/195
[1573849782] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2019/11/16 01:59:15 [error] 23641#0: *3741772 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:233/17, bytes from/to upstream:17/233
[1573849902] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573850502] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/11/16 02:10:07 [error] 23643#0: *3747960 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:237/17, bytes from/to upstream:17/237
[1573850623] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573850771] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1573851860] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.1% (6469856 kB) free!
[1573851980] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 89.2% (43947424 kB) free.
[1573852956] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.0% (43363540 kB) free.
[1573854162] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(25) < 2019/11/16 03:11:08 [error] 24130#0: *6770956 connect() failed (111: Connection refused) while connecting to upstream, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573854282] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573854882] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(27) < 2019/11/16 03:18:41 [error] 24130#0: *6771274 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1573855002] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573855734] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 56.5% (13944408 kB) free.
[1573864734] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186436 kB) free!
[1573864824] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26412652 kB) free!
[1573864854] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167228 kB) free!
[1573864944] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.9% (48599328 kB) free.
[1573864974] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184328 kB) free!
[1573865376] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.1% (20871056 kB) free.
[1573865740] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.7% (22628172 kB) free.
[1573865777] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.1% (22809544 kB) free.
[1573865826] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.0% (39901752 kB) free.
[1573865877] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 66.7% (16339200 kB) free.
[1573870118] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1573870238] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1573870544] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.3% (3779944 kB) free!
[1573870664] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.3% (3767000 kB) free!
[1573870784] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.2% (3743920 kB) free!
[1573870838] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1573870958] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1573871558] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1573871678] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1573872144] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (25173248 kB) free!
[1573872264] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.9% (47348624 kB) free.
[1573872371] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.4% (4127052 kB) free!
[1573872491] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4052764 kB) free!
[1573872611] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.7% (3978036 kB) free!
[1573873784] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1573873784] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2068232 kB) free!
[1573874664] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26156736 kB) free!
[1573874785] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.0% (22415596 kB) free!
[1573874905] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.2% (45143732 kB) free.
[1573875012] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 24.7% (5552632 kB) free.
[1573875222] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(170) < 2019/11/16 09:00:39 [error] 23639#0: *3944372 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1573875342] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573875500] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.2% (22246552 kB) free.
[1573875505] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (24051276 kB) free!
[1573875625] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.1% (19961288 kB) free!
[1573875745] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.9% (44672484 kB) free.
[1573876345] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21483572 kB) free!
[1573876465] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.1% (17302168 kB) free!
[1573876585] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44355996 kB) free.
[1573877143] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/16 09:30:01 [error] 23638#0: *3958632 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1573877185] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (19971064 kB) free!
[1573877202] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(5) < 2019/11/16 09:30:53 [error] 24131#0: *6840821 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573877263] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573877305] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (15701872 kB) free!
[1573877323] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573877425] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.0% (43464580 kB) free.
[1573877863] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/16 09:47:15 [error] 23645#0: *3967122 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1573877904] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170712 kB) free!
[1573877983] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573878025] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.9% (22431132 kB) free.
[1573878025] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18023396 kB) free!
[1573878145] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.3% (13583472 kB) free!
[1573878265] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (9400324 kB) free!
[1573878354] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (198100 kB) free!
[1573878474] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (197880 kB) free!
[1573878595] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (200180 kB) free!
[1573878866] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17552788 kB) free!
[1573879092] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1573879122] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(34) < 2019/11/16 10:08:31 [error] 24130#0: *6847185 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1573879212] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1573879212] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.0% (3154720 kB) free!
[1573879242] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(5) < 2019/11/16 10:08:54 [error] 24135#0: *6847140 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573879332] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.2% (2732280 kB) free!
[1573879362] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573879452] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.1% (2276820 kB) free!
[1573879466] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28077800 kB) free.
[1573879712] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.4% (14533012 kB) free.
[1573880052] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1573880052] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (180976 kB) free!
[1573880066] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (3910808 kB) free!
[1573880186] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35356792 kB) free.
[1573880383] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/16 10:23:23 [error] 23646#0: *3984832 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1573880432] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (1555828 kB) free!
[1573880504] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573880552] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195076 kB) free!
[1573880672] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191452 kB) free!
[1573880786] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7340744 kB) free!
[1573880906] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (3133716 kB) free!
[1573881026] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (36980512 kB) free.
[1573881626] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8302080 kB) free!
[1573881746] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.0% (4020380 kB) free!
[1573881762] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/11/16 10:45:20 [error] 24132#0: *6854125 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1573881866] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1573881866] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (414764 kB) free!
[1573881882] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573882329] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1573882449] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1573882682] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.3% (4006612 kB) free!
[1573882803] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.8% (3868112 kB) free!
[1573882923] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.2% (3736488 kB) free!
[1573882942] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.05, 5.99, 5.33
[1573882956] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.4% (22304452 kB) free.
[1573883004] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 47.7% (62875928 kB) free.
[1573883062] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.36, 6.14, 5.46
[1573883068] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 27.0% (6065336 kB) free.
[1573883182] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.01, 5.25, 5.22
[1573883672] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.5% (2798652 kB) free!
[1573883782] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.49, 3.74, 4.53
[1573883792] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (2127492 kB) free!
[1573883835] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9760232 kB) free!
[1573883871] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 18.2% (6347168 kB) free!
[1573883912] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (1445740 kB) free!
[1573883955] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.3% (8539116 kB) free!
[1573883991] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 16.3% (5662496 kB) free!
[1573884075] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.0% (7882236 kB) free!
[1573884111] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 14.4% (5003968 kB) free!
[1573884465] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 10.9% (3819020 kB) free!
[1573884585] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (3161420 kB) free!
[1573884705] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (2514580 kB) free!
[1573884711] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (1660564 kB) free!
[1573884723] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2016124 kB) free!
[1573884813] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5864968 kB) free!
[1573884933] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.4% (45384492 kB) free.
[1573885482] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/11/16 11:45:25 [error] 24130#0: *6865333 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1573885533] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (12038852 kB) free!
[1573885602] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573885654] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (7998324 kB) free!
[1573885774] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (3993560 kB) free!
[1573885906] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.7% (31755388 kB) free.
[1573885912] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 92.5% (32229312 kB) free.
[1573885956] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.0% (8366984 kB) free!
[1573886076] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.2% (7994380 kB) free!
[1573886186] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.0% (4158652 kB) free!
[1573886196] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.5% (7642256 kB) free!
[1573886306] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.5% (4035140 kB) free!
[1573886313] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 28.3% (6348376 kB) free.
[1573886330] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.7% (51060124 kB) free.
[1573886365] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 64.6% (15839684 kB) free.
[1573886426] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (3905456 kB) free!
[1573887382] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.37, 6.55, 5.30
[1573887397] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3789684 kB) free!
[1573887502] SERVICE ALERT: qservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.75, 7.51, 5.78
[1573887534] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20996160 kB) free!
[1573887622] SERVICE ALERT: qservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.31, 7.36, 5.94
[1573887626] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2360924 kB) free!
[1573887654] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.5% (16474364 kB) free!
[1573887675] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3780748 kB) free!
[1573887766] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.9% (2689940 kB) free!
[1573887774] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (12388716 kB) free!
[1573887886] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.1% (2487544 kB) free!
[1573888006] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2291672 kB) free!
[1573888374] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.4% (18999336 kB) free!
[1573888974] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (396704 kB) free!
[1573889376] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4739556 kB) free!
[1573889422] SERVICE FLAPPING ALERT: qservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1573889422] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.59, 3.56, 4.26
[1573889496] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.1% (4439248 kB) free!
[1573889574] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (13931244 kB) free!
[1573889616] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (4298644 kB) free!
[1573889913] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.7% (3758536 kB) free!
[1573889943] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.8% (22156760 kB) free.
[1573889994] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.0% (22801940 kB) free.
[1573890033] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.3% (3203312 kB) free!
[1573890040] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.3% (39085644 kB) free.
[1573890094] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.3% (46948268 kB) free.
[1573890153] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.8% (2639676 kB) free!
[1573890174] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36808864 kB) free.
[1573890380] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9101708 kB) free!
[1573890500] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.0% (8873764 kB) free!
[1573890512] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.9% (2093236 kB) free!
[1573890557] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167484 kB) free!
[1573890565] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173076 kB) free!
[1573890620] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 17.6% (8650376 kB) free!
[1573890632] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.9% (2092580 kB) free!
[1573890677] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178300 kB) free!
[1573890685] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178892 kB) free!
[1573890752] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 12.9% (2095096 kB) free!
[1573890753] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (180916 kB) free!
[1573890774] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7227376 kB) free!
[1573890797] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177560 kB) free!
[1573890805] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1573890805] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187532 kB) free!
[1573890894] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3440236 kB) free!
[1573891014] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398632 kB) free!
[1573891203] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.5% (4782944 kB) free!
[1573891323] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.5% (3544744 kB) free!
[1573891443] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1573891443] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.8% (3378692 kB) free!
[1573891686] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < Channel error on connection <0.30048.7572> (10.147.212.147:35068 -> 10.147.212.203:5672, vhost: '/', user: 'guest'), channel 1:
[1573891808] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573892214] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26773336 kB) free.
[1573892616] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2370040 kB) free!
[1573892814] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (392400 kB) free!
[1573892935] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (388996 kB) free!
[1573893020] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4213708 kB) free!
[1573893055] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (386832 kB) free!
[1573893104] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (238780 kB) free!
[1573893224] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (248016 kB) free!
[1573893243] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (1418388 kB) free!
[1573893344] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1573893344] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (219396 kB) free!
[1573894255] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36224000 kB) free.
[1573894388] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.21626.604> (10.147.212.158:51848 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1573894508] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573894855] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6827992 kB) free!
[1573894866] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.6% (22103812 kB) free.
[1573894903] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/16 14:31:19 [error] 23644#0: *4128574 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: _, request: "GET /favicon.ico HTTP/1.1", host: "altsmsc.tesync", referrer: "http://altsmsc.tesync/"
[1573894906] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.5% (22667572 kB) free.
[1573894912] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 15.2% (5278508 kB) free!
[1573894954] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 28.5% (14021564 kB) free.
[1573894975] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2257672 kB) free!
[1573895008] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31145008 kB) free.
[1573895023] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573895032] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 13.3% (4639224 kB) free!
[1573895045] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 31.2% (6997256 kB) free.
[1573895084] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 31.2% (15396340 kB) free.
[1573895126] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 66.6% (16334752 kB) free.
[1573895152] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 11.4% (3981892 kB) free!
[1573895160] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.4% (22400420 kB) free.
[1573895505] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 13.6% (4773924 kB) free!
[1573895552] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1348128 kB) free!
[1573895625] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 12.1% (4230880 kB) free!
[1573895653] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.3% (4118328 kB) free!
[1573895745] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 10.5% (3666652 kB) free!
[1573895752] SERVICE FLAPPING ALERT: sig2-vm2;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1573895752] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (974772 kB) free!
[1573895773] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.7% (3531492 kB) free!
[1573895894] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.1% (2948312 kB) free!
[1573896160] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.4309.302> (10.147.212.42:55258 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1573896212] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19438828 kB) free!
[1573896280] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573896332] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (15726444 kB) free!
[1573896345] SERVICE FLAPPING ALERT: sig5-vm2;Memory;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1573896345] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (817332 kB) free!
[1573896452] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40213436 kB) free.
[1573896494] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186684 kB) free!
[1573896934] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197972 kB) free!
[1573896946] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 91.0% (31855280 kB) free.
[1573896952] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 92.8% (32335132 kB) free.
[1573896964] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.3% (3497136 kB) free!
[1573896991] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 33.72, 32.06, 30.85
[1573897052] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7672024 kB) free!
[1573897054] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187168 kB) free!
[1573897085] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.9% (3403664 kB) free!
[1573897094] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 31.4% (7053992 kB) free.
[1573897112] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 32.25, 31.78, 30.88
[1573897172] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3427576 kB) free!
[1573897174] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189752 kB) free!
[1573897205] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1573897205] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 13.5% (3319072 kB) free!
[1573897232] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 31.72, 31.76, 30.97
[1573897274] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188796 kB) free!
[1573897292] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397860 kB) free!
[1573897318] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (1353044 kB) free!
[1573897394] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191320 kB) free!
[1573897438] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (1212628 kB) free!
[1573897514] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (513068 kB) free!
[1573897558] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (1083368 kB) free!
[1573897812] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1573897892] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16148688 kB) free!
[1573897932] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1573898492] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382820 kB) free!
[1573898650] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1573898770] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1573898802] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(42) < 2019/11/16 15:36:38 [error] 24131#0: *6903757 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1573898890] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 362 processes
[1573898922] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(34) < 2019/11/16 15:38:24 [error] 24131#0: *6904064 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573899033] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 31.64, 28.44, 29.35
[1573899042] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(20) < 2019/11/16 15:40:34 [error] 24130#0: *6904120 recv() failed (104: Connection reset by peer) while proxying connection, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1573899309] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.6285.644> (10.147.212.158:56104 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1573899429] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573899490] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1573899623] SERVICE FLAPPING ALERT: qservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1573899692] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37744236 kB) free.
[1573899893] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (237364 kB) free!
[1573900013] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (238592 kB) free!
[1573900133] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (231940 kB) free!
[1573900233] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 31.58, 30.65, 30.29
[1573900242] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573900292] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6474052 kB) free!
[1573900353] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 28.80, 30.14, 30.15
[1573900364] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236812 kB) free!
[1573900413] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (2399652 kB) free!
[1573900473] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 29.78, 29.95, 30.07
[1573900485] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (247472 kB) free!
[1573900533] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (456924 kB) free!
[1573900605] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (250424 kB) free!
[1573900805] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2225184 kB) free!
[1573900844] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (9514956 kB) free!
[1573900964] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9049328 kB) free!
[1573901084] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.0% (9360848 kB) free!
[1573901133] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (15047756 kB) free!
[1573901294] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (2033868 kB) free!
[1573901415] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (1534240 kB) free!
[1573901535] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (1024540 kB) free!
[1573901733] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (450788 kB) free!
[1573902224] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(115) < 2019/11/16 16:27:22 [error] 23633#0: *4187834 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573902273] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1573902273] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 29.02, 28.87, 29.38
[1573902344] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573902429] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.19466.635> (10.147.212.158:53022 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1573902549] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573902933] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.3% (24114148 kB) free!
[1573903533] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (446288 kB) free!
[1573904073] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 34.26, 32.11, 30.76
[1573904133] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.7% (15465940 kB) free!
[1573904193] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 32.52, 31.98, 30.87
[1573904313] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 31.35, 31.59, 30.85
[1573904619] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.2% (22259668 kB) free.
[1573904646] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.7% (22723224 kB) free.
[1573904676] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 31.3% (15442272 kB) free.
[1573904710] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 32.5% (7306356 kB) free.
[1573904733] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (455676 kB) free!
[1573904740] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 30.3% (14955268 kB) free.
[1573904768] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.7% (18088148 kB) free!
[1573904771] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31460764 kB) free.
[1573904802] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 70.5% (17278868 kB) free.
[1573904825] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 94.6% (23178716 kB) free.
[1573904856] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.3% (14515556 kB) free.
[1573905280] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.25509.336> (10.147.212.42:34002 -> 10.147.212.55:5672, vhost: '/', user: 'guest'), channel 1:
[1573905401] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573905822] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.8% (3400052 kB) free!
[1573905914] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.0% (2694284 kB) free!
[1573905942] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.3% (3293452 kB) free!
[1573905952] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 15.7% (5466512 kB) free!
[1573905975] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13842884 kB) free!
[1573906034] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (2231168 kB) free!
[1573906062] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.0% (3215204 kB) free!
[1573906072] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 14.1% (4927744 kB) free!
[1573906095] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39099108 kB) free.
[1573906113] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 29.46, 28.67, 29.36
[1573906154] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1766348 kB) free!
[1573906192] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 12.7% (4409096 kB) free!
[1573906547] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 14.1% (4948940 kB) free!
[1573906667] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 12.7% (4441324 kB) free!
[1573906695] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (9155708 kB) free!
[1573906787] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 11.3% (3941700 kB) free!
[1573906792] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (1695084 kB) free!
[1573906815] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4840196 kB) free!
[1573906935] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (910512 kB) free!
[1573907387] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (1413564 kB) free!
[1573907442] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/16 17:56:38 [error] 24132#0: *6930716 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573907535] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19509344 kB) free!
[1573907563] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573907862] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2251532 kB) free!
[1573907955] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 32.7% (7352332 kB) free.
[1573907987] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 91.0% (31856136 kB) free.
[1573907993] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 93.1% (32439588 kB) free.
[1573908135] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (381508 kB) free!
[1573908284] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3916028 kB) free!
[1573908735] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (18157620 kB) free!
[1573908762] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/16 18:15:12 [error] 24130#0: *6934223 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573908882] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573909454] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174020 kB) free!
[1573909574] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176976 kB) free!
[1573909602] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165396 kB) free!
[1573909633] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189180 kB) free!
[1573909694] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192756 kB) free!
[1573909713] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 63.75, 57.92, 44.36
[1573909722] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193480 kB) free!
[1573909753] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169772 kB) free!
[1573909833] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 71.15, 62.08, 47.49
[1573909842] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167888 kB) free!
[1573909874] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (220392 kB) free!
[1573909932] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1573909935] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (9752928 kB) free!
[1573909955] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 52.10, 58.40, 47.95
[1573910052] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1573910145] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/16 18:43:38 [error] 23646#0: *4267486 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1573910264] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.8% (1761464 kB) free!
[1573910265] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573910384] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.8% (1760444 kB) free!
[1573910504] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.8% (1758032 kB) free!
[1573910535] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35692556 kB) free.
[1573910556] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.03, 15.63, 31.08
[1573910745] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (240572 kB) free!
[1573910865] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (301764 kB) free!
[1573910985] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (228360 kB) free!
[1573911136] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16175180 kB) free!
[1573911156] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 17.58, 12.05, 21.17
[1573911256] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (10642860 kB) free!
[1573911376] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (6857620 kB) free!
[1573911549] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(2) < Channel error on connection <0.2104.619> (10.147.212.158:60060 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1573911669] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573911976] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18771016 kB) free!
[1573912155] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4288792 kB) free!
[1573912275] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.8% (3989528 kB) free!
[1573912395] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.3% (3658248 kB) free!
[1573912544] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.4% (22801536 kB) free.
[1573912576] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (452404 kB) free!
[1573912577] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.0% (22546572 kB) free.
[1573912612] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 40.5% (53423688 kB) free.
[1573912636] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.9% (39856584 kB) free.
[1573912663] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 26.9% (13280012 kB) free.
[1573912698] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 69.7% (17087632 kB) free.
[1573912726] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 94.2% (23095136 kB) free.
[1573912866] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.7% (25971544 kB) free!
[1573912986] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 30.7% (40498996 kB) free.
[1573912995] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2074768 kB) free!
[1573913155] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.3% (3046640 kB) free!
[1573913275] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (1861200 kB) free!
[1573913395] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (1274228 kB) free!
[1573913586] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.5% (24401992 kB) free!
[1573913706] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 36.5% (48160312 kB) free.
[1573914413] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12172936 kB) free!
[1573914533] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8381180 kB) free!
[1573914653] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39480576 kB) free.
[1573915253] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18098824 kB) free!
[1573915373] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.7% (14080528 kB) free!
[1573915493] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (10189876 kB) free!
[1573916093] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.0% (22420976 kB) free!
[1573916693] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38183956 kB) free.
[1573917293] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15209492 kB) free!
[1573917413] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (11166824 kB) free!
[1573917533] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (7317712 kB) free!
[1573918133] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20846200 kB) free!
[1573918395] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 31.9% (7164632 kB) free.
[1573918733] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34908316 kB) free.
[1573918793] SERVICE FLAPPING ALERT: sig2-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1573919333] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12601968 kB) free!
[1573919393] SERVICE FLAPPING ALERT: sig5-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573919453] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8579008 kB) free!
[1573919505] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1074156 kB) free!
[1573919574] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.8% (24804140 kB) free!
[1573920384] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1573920533] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1573920774] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29798140 kB) free.
[1573921102] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1573921374] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8342208 kB) free!
[1573921494] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4449164 kB) free!
[1573921614] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34827024 kB) free.
[1573921672] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1573922214] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11407624 kB) free!
[1573922287] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.9% (20463888 kB) free.
[1573922317] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.3% (14517256 kB) free.
[1573922334] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7349468 kB) free!
[1573922354] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.7% (20637772 kB) free!
[1573922357] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32523036 kB) free.
[1573922556] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573922684] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9748024 kB) free!
[1573922804] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (9668328 kB) free!
[1573922925] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (9593332 kB) free!
[1573923795] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.8% (3766176 kB) free!
[1573923883] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(17) < 2019/11/16 22:27:33 [error] 24133#0: *6980773 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1573923915] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.9% (3572976 kB) free!
[1573924003] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573924035] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.1% (3379384 kB) free!
[1573924164] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12736540 kB) free!
[1573924284] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8736880 kB) free!
[1573924404] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.4% (19027680 kB) free!
[1573925004] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (377328 kB) free!
[1573925235] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1630584 kB) free!
[1573925707] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.6% (25825952 kB) free!
[1573925827] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (8002736 kB) free!
[1573925947] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (7353332 kB) free!
[1573926204] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29741604 kB) free.
[1573926356] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9127568 kB) free!
[1573926476] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.3% (9039316 kB) free!
[1573926547] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.3% (20226660 kB) free!
[1573926596] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.2% (8961736 kB) free!
[1573926804] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9948084 kB) free!
[1573926925] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6010556 kB) free!
[1573927045] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2207500 kB) free!
[1573927147] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (12850996 kB) free!
[1573927645] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14430364 kB) free!
[1573927694] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1573927747] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1573927747] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.1% (19874768 kB) free!
[1573928245] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27554456 kB) free.
[1573928347] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (12593296 kB) free!
[1573928845] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7120232 kB) free!
[1573928947] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.2% (18777880 kB) free!
[1573928965] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3319120 kB) free!