[1546972437] SERVICE FLAPPING ALERT: pservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1546984546] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 7.78, 8.02, 9.63
[1547003146] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.87, 10.90, 10.33
[1547003266] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.90, 10.81, 10.36
[1547003325] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1547003386] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.90, 10.47, 10.29
[1547003445] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1547003986] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 9.38, 9.63, 9.97
[1547004586] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.77, 10.60, 10.27
[1547004706] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.19, 10.55, 10.30
[1547004826] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.29, 10.01, 10.12
[1547005845] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547005965] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1547006085] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1547008437] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1547008557] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1547008677] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1547009685] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547009805] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547009925] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1547010525] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547010646] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547010766] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1547011366] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1547011966] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547012086] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1547012686] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547012806] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1547013406] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547013527] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1547013647] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1547014847] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547014967] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547015087] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547016826] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 8.20, 8.33, 9.20
[1547026487] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1547027688] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1547027808] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1547028408] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547028528] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1547028537] HOST ALERT: webservercluster2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1547028567] SERVICE ALERT: webservercluster2;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'webservercluster2' was out of bounds)
[1547028662] HOST ALERT: webservercluster2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1547028687] SERVICE ALERT: webservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1547029575] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547029605] HOST ALERT: webservercluster1;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1547029695] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 94.9% (15424344 kB) free.
[1547029699] HOST ALERT: webservercluster1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1547030902] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1547030903] HOST ALERT: sservercluster1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.49)
[1547030930] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sservercluster1' was out of bounds)
[1547031023] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1547031030] HOST ALERT: sservercluster1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1547031051] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1547031622] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 281 processes
[1547032720] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1547032728] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547032840] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1547032848] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 297 processes
[1547032960] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1547032968] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1547033315] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 5 warn out of 342 processes
[1547033916] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 338 processes
[1547034768] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547034889] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547035009] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1547036809] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1547037760] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1547041009] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547041129] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547041249] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1547041849] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1547041849] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1547042877] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 12 users currently logged in
[1547044249] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547044369] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1547044489] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1547050426] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.55, 11.78, 10.66
[1547050546] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.99, 10.87, 10.46
[1547050666] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.25, 10.32, 10.31
[1547051266] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 8.80, 9.23, 9.75
[1547053489] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1547053666] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.25, 10.65, 10.31
[1547053786] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.09, 10.70, 10.36
[1547053906] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 14.48, 12.05, 10.89