[1582741924] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(47) < 2020/02/27 00:01:47 [error] 12246#0: *60692884 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582742044] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582742062] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.8% (29469272 kB) free.
[1582742210] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177432 kB) free!
[1582742330] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172272 kB) free!
[1582742411] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 368 processes
[1582742450] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177532 kB) free!
[1582742531] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1582742609] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186220 kB) free!
[1582742729] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177400 kB) free!
[1582742849] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193756 kB) free!
[1582743899] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (183876 kB) free!
[1582743959] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 68.0% (16668472 kB) free.
[1582744850] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 38.9% (9597660 kB) free.
[1582746051] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169232 kB) free!
[1582746171] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197980 kB) free!
[1582746271] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.8% (8294540 kB) free!
[1582746291] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168672 kB) free!
[1582746323] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1582746391] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 14.6% (7215652 kB) free!
[1582746511] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 11.8% (5790644 kB) free!
[1582747112] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (637796 kB) free!
[1582749449] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.8% (10005792 kB) free.
[1582750491] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (213076 kB) free!
[1582750520] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7458984 kB) free!
[1582750611] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (211276 kB) free!
[1582750640] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 78.1% (38466604 kB) free.
[1582750731] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (223000 kB) free!
[1582750772] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1582750892] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1582751332] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 76.2% (37528316 kB) free.
[1582751644] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/27 02:43:36 [error] 12257#0: *60731871 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1582751764] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582752745] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2669) < 172.105.59.173 [27/Feb/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1582752865] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582752964] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(210) < 2020/02/27 03:01:22 [error] 12257#0: *60736342 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582753019] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(158) < 103.16.101.72 [27/Feb/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1582753084] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582753141] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582756684] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(33) < 2020/02/27 04:00:39 [error] 12246#0: *60751236 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582756804] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582760249] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189140 kB) free!
[1582760369] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 29.7% (7290676 kB) free.
[1582760723] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.2% (20891348 kB) free.
[1582760969] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (217900 kB) free!
[1582761089] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (218164 kB) free!
[1582761209] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (213172 kB) free!
[1582761291] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 45.8% (11293024 kB) free.
[1582770364] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 58.1% (28633780 kB) free.
[1582770395] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.7% (21503380 kB) free.
[1582770427] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.7% (14098248 kB) free.
[1582772691] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180712 kB) free!
[1582772812] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173092 kB) free!
[1582772932] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168452 kB) free!
[1582778087] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 90.16, 109.44, 52.61
[1582778207] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 26.32, 78.13, 47.99
[1582778327] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 11.20, 55.57, 43.40
[1582778453] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.6% (21622500 kB) free.
[1582778927] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 18.09, 25.44, 34.67
[1582779527] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 13.60, 16.44, 25.73
[1582780127] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 17.42, 38.57, 40.65
[1582780247] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 15.51, 30.66, 37.49
[1582780368] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1582780368] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 18.74, 26.14, 34.97
[1582780706] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1582780826] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1582780946] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1582780968] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 31.21, 135.53, 104.65
[1582781405] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/27 10:56:08 [error] 12257#0: *60879302 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: rq-2-1gbps.server, request: "GET /favicon.ico HTTP/1.1", host: "rq-2-1gbps.server:8091"
[1582781525] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582781816] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1582781868] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 709 processes
[1582781936] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 349 processes
[1582781988] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 708 processes
[1582782056] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 356 processes
[1582782109] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 709 processes
[1582782226] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.2673.3594> (10.147.212.158:45560 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1582782347] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1582782659] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 15.9% (3898100 kB) free!
[1582782756] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.1% (3394452 kB) free!
[1582782779] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.1% (3698672 kB) free!
[1582782876] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.9% (3127288 kB) free!
[1582782899] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.1% (3448820 kB) free!
[1582782996] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 11.9% (2662064 kB) free!
[1582783500] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1981752 kB) free!
[1582783596] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (446380 kB) free!
[1582783884] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.0% (3449000 kB) free!
[1582784004] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.1% (3243356 kB) free!
[1582784123] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (4622824 kB) free!
[1582784124] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.4% (3053708 kB) free!
[1582784243] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4454624 kB) free!
[1582784364] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.5% (4284304 kB) free!
[1582784546] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1582784964] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 15.6% (3844232 kB) free!
[1582785056] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 365 processes
[1582785084] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 13.1% (3225896 kB) free!
[1582785125] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(25) < 2020/02/27 12:00:43 [error] 12246#0: *60912350 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582785204] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 10.3% (2532596 kB) free!
[1582785246] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582785266] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(39) < 172.105.59.173 [27/Feb/2020:12:02:59 +0530] TCP 502 0 0 60.000 "10.147.212.143:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1582785300] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (20412020 kB) free!
[1582785325] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (1525804 kB) free!
[1582785386] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582785396] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.5% (15828428 kB) free.
[1582785420] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (15946992 kB) free!
[1582785541] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40326792 kB) free.
[1582785656] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1582785684] SERVICE ALERT: sig4-vm2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 554 processes
[1582785746] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1582785804] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 549 processes
[1582785805] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194908 kB) free!
[1582785854] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 16.14, 13.35, 8.51
[1582785866] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1582785930] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(2) < {writer,send_failed,{error,timeout}}
[1582785974] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.43, 11.97, 8.58
[1582785986] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1582786052] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1582786094] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 9.03, 11.32, 8.79
[1582786140] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14847984 kB) free!
[1582786260] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39280264 kB) free.
[1582786404] SERVICE ALERT: sig4-vm2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 551 processes
[1582786405] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.5% (22836068 kB) free.
[1582786473] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582786524] SERVICE ALERT: sig4-vm2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 555 processes
[1582786603] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582786604] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582786645] SERVICE ALERT: sig4-vm2;Total Processes;WARNING;HARD;3;PROCS WARNING: 554 processes
[1582786721] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582786727] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1582786727] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.3% (30191812 kB) free.
[1582786763] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2181240 kB) free!
[1582786845] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1582786860] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14818364 kB) free!
[1582786980] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40261924 kB) free.
[1582787067] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582787169] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 13.0% (6385104 kB) free!
[1582787197] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1582787245] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 549 processes
[1582787289] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 12.9% (6375160 kB) free!
[1582787319] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1582787409] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 12.9% (6375836 kB) free!
[1582787580] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (16011272 kB) free!
[1582787700] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38148928 kB) free.
[1582787728] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.3% (21650632 kB) free.
[1582787749] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1582787749] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 71.7% (17683472 kB) free.
[1582787778] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.5% (21704152 kB) free.
[1582787814] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.8% (39343424 kB) free.
[1582787845] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 689 processes
[1582787927] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582787932] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9770440 kB) free!
[1582787933] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1582787965] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 694 processes
[1582788050] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1582788052] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.0% (9346932 kB) free!
[1582788055] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.4% (30254424 kB) free.
[1582788085] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 695 processes
[1582788172] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.1% (8924848 kB) free!
[1582788300] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13708132 kB) free!
[1582788420] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32740960 kB) free.
[1582788687] SERVICE FLAPPING ALERT: sig4-vm2;Total Processes;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1582788687] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 229 processes
[1582789288] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1148 processes
[1582789368] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 18.37, 21.28, 31.53
[1582789408] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1152 processes
[1582789528] SERVICE ALERT: sig4-vm2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1152 processes
[1582789549] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185752 kB) free!
[1582789620] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21939084 kB) free!
[1582789669] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178224 kB) free!
[1582789740] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.3% (47797520 kB) free.
[1582789790] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195928 kB) free!
[1582789968] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 17.72, 16.92, 24.18
[1582789972] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (3438688 kB) free!
[1582790200] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (4696868 kB) free!
[1582790320] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (4487404 kB) free!
[1582790340] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (20437280 kB) free!
[1582790440] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.5% (4278532 kB) free!
[1582790460] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.1% (46251388 kB) free.
[1582791061] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19694484 kB) free!
[1582791110] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1582791181] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.0% (44846624 kB) free.
[1582791363] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.4% (45518620 kB) free.
[1582791394] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.6% (21615968 kB) free.
[1582791452] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.2% (21372496 kB) free.
[1582791781] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18579792 kB) free!
[1582791901] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.0% (43501480 kB) free.
[1582792501] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17562044 kB) free!
[1582792586] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 3 users currently logged in
[1582792621] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.2% (13498840 kB) free!
[1582792741] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38564408 kB) free.
[1582793196] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.7% (3976128 kB) free!
[1582793316] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.2% (3630192 kB) free!
[1582793341] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13327724 kB) free!
[1582793436] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.6% (3283424 kB) free!
[1582793461] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39609488 kB) free.
[1582794036] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (24.7% change >= 20.0% threshold)
[1582794036] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1684580 kB) free!
[1582794061] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (13324608 kB) free!
[1582794181] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38811288 kB) free.
[1582794400] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.7% (3622324 kB) free!
[1582794520] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.0% (3445684 kB) free!
[1582794640] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.3% (3287680 kB) free!
[1582794781] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13722512 kB) free!
[1582794813] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.7% (15868516 kB) free.
[1582794860] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.9% (21447400 kB) free.
[1582794901] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9692788 kB) free!
[1582794953] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 44.3% (58407572 kB) free.
[1582795470] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.5% (3585144 kB) free!
[1582795581] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189096 kB) free!
[1582795590] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197836 kB) free!
[1582795621] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (213920 kB) free!
[1582795701] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190516 kB) free!
[1582795710] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173496 kB) free!
[1582795741] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (244796 kB) free!
[1582795821] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189448 kB) free!
[1582795856] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1582795861] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (227996 kB) free!
[1582795976] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1582796005] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 13.2% (3265788 kB) free!
[1582796096] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1582796125] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 11.5% (2830940 kB) free!
[1582796187] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(141) < 183.82.105.122 [27/Feb/2020:15:00:23 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1582796245] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 10.1% (2492396 kB) free!
[1582796307] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582796696] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1582796696] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1582796731] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.3% (4722500 kB) free!
[1582796840] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9335780 kB) free!
[1582796845] SERVICE FLAPPING ALERT: db1-vm2;Memory;STARTED; Service appears to have started flapping (24.0% change >= 20.0% threshold)
[1582796845] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.5% (22843096 kB) free.
[1582796851] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.1% (4435580 kB) free!
[1582796866] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (201652 kB) free!
[1582796894] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 22.10, 23.67, 21.12
[1582796960] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (9135980 kB) free!
[1582796972] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.0% (4168344 kB) free!
[1582796986] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175576 kB) free!
[1582797080] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (8926264 kB) free!
[1582797106] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168548 kB) free!
[1582798094] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 16.52, 17.74, 19.94
[1582798172] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1541048 kB) free!
[1582798562] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25864796 kB) free!
[1582798682] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.2% (21353264 kB) free!
[1582798802] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17173632 kB) free!
[1582799645] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(68) < 2020/02/27 16:01:29 [error] 12257#0: *61054337 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582799696] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 373 processes
[1582799736] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3122612 kB) free!
[1582799765] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582799816] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1582799856] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.2% (3442228 kB) free.
[1582799907] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(24) < 203.212.70.170 [27/Feb/2020:16:03:30 +0530] TCP 502 0 0 120.001 "10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1582799936] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582800002] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29711304 kB) free.
[1582800027] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582800080] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4547524 kB) free!
[1582800928] SERVICE FLAPPING ALERT: sig4-vm2;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582800986] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1582801056] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (2799708 kB) free!
[1582801106] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1582801176] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (3236424 kB) free!
[1582801202] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13452184 kB) free!
[1582801226] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1582801296] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.2% (3604740 kB) free.
[1582801322] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39607216 kB) free.
[1582801368] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582801736] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 352 processes
[1582801896] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (1518744 kB) free!
[1582801922] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15307456 kB) free!
[1582802017] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.5% (1870336 kB) free!
[1582802042] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11136652 kB) free!
[1582802137] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.5% (2194472 kB) free!
[1582802163] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37609512 kB) free.
[1582802294] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 91.59, 34.07, 19.05
[1582802336] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1582802728] SERVICE ALERT: sig4-vm2;Total Processes;OK;HARD;1;PROCS OK: 228 processes
[1582802737] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.3% (3454120 kB) free.
[1582802762] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.9% (14425776 kB) free!
[1582802882] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (12416228 kB) free!
[1582802894] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1582802894] SERVICE ALERT: sig4-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.06, 4.60, 10.00
[1582803002] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34702764 kB) free.
[1582803337] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.4% (2659844 kB) free!
[1582803457] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3155648 kB) free!
[1582803537] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1582803577] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (3109040 kB) free!
[1582803602] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12240688 kB) free!
[1582803664] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1582803722] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36324900 kB) free.
[1582803752] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.1% (21598932 kB) free.
[1582803784] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582803786] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.0% (21212604 kB) free.
[1582803820] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1582803831] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 91.7% (45207636 kB) free.
[1582803866] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.7% (39273216 kB) free.
[1582803904] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.6% (21715604 kB) free.
[1582803905] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582803934] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.4% (21174748 kB) free.
[1582804095] SERVICE ALERT: sig4-vm2;CPU Load;OK;HARD;1;OK - load average: 0.07, 0.13, 2.78
[1582804177] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1582804177] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.7% (4338348 kB) free.
[1582804322] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13747812 kB) free!
[1582804442] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (10013508 kB) free!
[1582804540] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186560 kB) free!
[1582804562] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35038820 kB) free.
[1582804660] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172884 kB) free!
[1582804737] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 362 processes
[1582804780] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (183984 kB) free!
[1582804857] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 362 processes
[1582804977] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 358 processes
[1582804992] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.1% (2500284 kB) free!
[1582805105] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1582805112] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183784 kB) free!
[1582805162] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11706320 kB) free!
[1582805233] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174480 kB) free!
[1582805282] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7946792 kB) free!
[1582805377] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.5% (2197296 kB) free!
[1582805402] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32546440 kB) free.
[1582805497] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.6% (2853760 kB) free!
[1582805577] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1582805617] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3517588 kB) free.
[1582806002] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9486588 kB) free!
[1582806122] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32816688 kB) free.
[1582806177] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1582806217] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.0% (1951212 kB) free!
[1582806220] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.1% (3605584 kB) free!
[1582806297] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1582806337] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.9% (2743892 kB) free!
[1582806340] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.3% (3433320 kB) free!
[1582806417] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1582806457] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3366280 kB) free.
[1582806460] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.6% (3284664 kB) free!
[1582806722] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10452888 kB) free!
[1582806843] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (6494876 kB) free!
[1582806963] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31579268 kB) free.
[1582807017] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1582807058] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.8% (1750848 kB) free!
[1582807061] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.7% (15870324 kB) free.
[1582807138] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1582807178] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.2% (2478448 kB) free!
[1582807258] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1582807298] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3356412 kB) free.
[1582807505] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1582807563] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8429652 kB) free!
[1582807626] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1582807683] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4668448 kB) free!
[1582807803] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29725128 kB) free.
[1582807858] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 362 processes
[1582807898] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (1447072 kB) free!
[1582807978] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1582808018] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.4% (2343612 kB) free!
[1582808138] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (3041300 kB) free!
[1582808245] SERVICE FLAPPING ALERT: db1-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582808403] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6527108 kB) free!
[1582808523] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2778600 kB) free!
[1582808578] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 354 processes
[1582808643] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28923072 kB) free.
[1582808698] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1582808738] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1186452 kB) free!
[1582808771] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9362700 kB) free!
[1582808891] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9128816 kB) free!
[1582809011] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.0% (8866888 kB) free!
[1582809243] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5579316 kB) free!
[1582809298] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1582809338] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (3097360 kB) free!
[1582809363] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (1830748 kB) free!
[1582809418] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1582809483] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28242776 kB) free.
[1582809538] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1582809938] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.4% (4608324 kB) free.
[1582810083] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5301840 kB) free!
[1582810138] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1582810203] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (1543996 kB) free!
[1582810259] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1582810323] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28654852 kB) free.
[1582810362] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.0% (4160812 kB) free!
[1582810379] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1582810482] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.3% (3992456 kB) free!
[1582810602] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 15.8% (3867984 kB) free!
[1582810924] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6594688 kB) free!
[1582811044] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (2761112 kB) free!
[1582811070] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9201184 kB) free!
[1582811164] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28572636 kB) free.
[1582811190] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (9797060 kB) free!
[1582811311] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (9797888 kB) free!
[1582811738] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 20.0% (3258936 kB) free!
[1582811764] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7379380 kB) free!
[1582811859] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.3% (3623828 kB) free.
[1582811884] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3479396 kB) free!
[1582812004] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2821152 kB) free!
[1582812012] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4645928 kB) free!
[1582812026] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1582812402] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2179912 kB) free!
[1582812779] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1582813058] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.6% (2217180 kB) free!
[1582813178] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.9% (2585192 kB) free!
[1582813204] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1582813204] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14603332 kB) free!
[1582813298] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (2820092 kB) free!
[1582813379] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582813499] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1582813619] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 353 processes
[1582813804] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34072324 kB) free.
[1582813898] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 33.2% (5404204 kB) free.
[1582814165] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(207) < 2020/02/27 20:02:07 [error] 12257#0: *61183402 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1582814219] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582814285] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582814404] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12086032 kB) free!
[1582814427] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4) < 203.212.70.170 [27/Feb/2020:20:01:40 +0530] TCP 502 0 0 60.000 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1582814524] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38293920 kB) free.
[1582814548] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582814819] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1582814833] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1582814895] SERVICE FLAPPING ALERT: sig4-vm2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1582815124] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15792088 kB) free!
[1582815244] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11713116 kB) free!
[1582815364] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.7% (14081224 kB) free!
[1582815419] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1582815511] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4327716 kB) free!
[1582815540] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1582815660] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1582816602] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 26.3% (6446204 kB) free.
[1582816625] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1582816746] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582816866] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1582817164] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36271896 kB) free.
[1582817202] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163460 kB) free!
[1582817322] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177592 kB) free!
[1582817442] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1582817442] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181336 kB) free!
[1582817460] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1582817498] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2983668 kB) free!
[1582817618] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.2% (2802948 kB) free!
[1582817738] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.1% (2612964 kB) free!
[1582817764] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10231576 kB) free!
[1582817884] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6060936 kB) free!
[1582818004] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35662520 kB) free.
[1582818060] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1582818180] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1582818300] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 352 processes
[1582818310] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (4526432 kB) free!
[1582818339] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 27.3% (4432592 kB) free.
[1582818430] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (4405796 kB) free!
[1582818550] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.6% (4305504 kB) free!
[1582818604] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9486116 kB) free!
[1582818724] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.9% (5114592 kB) free!
[1582818844] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34892748 kB) free.
[1582818900] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1582819444] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8882056 kB) free!
[1582819564] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4702780 kB) free!
[1582819684] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (628092 kB) free!
[1582820100] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582820221] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582820341] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1582820350] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2205864 kB) free!
[1582820442] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 64.1% (15713756 kB) free.
[1582820466] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1582820884] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23142144 kB) free!
[1582820941] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1582821061] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 366 processes
[1582821181] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1582821484] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33911764 kB) free.
[1582821781] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1582822084] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12195428 kB) free!
[1582822204] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8424484 kB) free!
[1582822325] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (4805280 kB) free!
[1582822381] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1582822501] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1582822925] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13724268 kB) free!
[1582823101] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1582823221] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582823342] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582824125] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33258400 kB) free.
[1582824725] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9828936 kB) free!
[1582824845] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6036188 kB) free!
[1582824965] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32399576 kB) free.
[1582825142] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 349 processes
[1582825565] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (10029224 kB) free!
[1582825685] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6387904 kB) free!
[1582825742] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582825805] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32839620 kB) free.
[1582826405] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10269412 kB) free!
[1582826492] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 405 processes
[1582826525] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6540524 kB) free!
[1582826612] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 405 processes
[1582826645] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33026104 kB) free.
[1582827053] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (247892 kB) free!
[1582827173] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (221784 kB) free!
[1582827245] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9647984 kB) free!
[1582827289] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7244332 kB) free!
[1582827293] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (260392 kB) free!
[1582827365] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6077732 kB) free!
[1582827409] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8420744 kB) free!
[1582827486] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31841292 kB) free.
[1582827529] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (9766456 kB) free!
[1582827812] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1582827933] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1582828085] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10148160 kB) free!
[1582828142] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes