[1584038220] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11477024 kB) free!
[1584038333] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198260 kB) free!
[1584038340] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (8054716 kB) free!
[1584038453] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (209528 kB) free!
[1584038460] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (4850924 kB) free!
[1584038573] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183284 kB) free!
[1584039060] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19473800 kB) free!
[1584039660] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (2783940 kB) free!
[1584039888] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184836 kB) free!
[1584039948] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187924 kB) free!
[1584040011] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.8% (189728 kB) free!
[1584040071] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (187420 kB) free!
[1584040265] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15837032 kB) free!
[1584040671] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 72.2% (17697164 kB) free.
[1584040865] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32020188 kB) free.
[1584041024] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584041465] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11109612 kB) free!
[1584041585] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7606888 kB) free!
[1584041705] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (4346804 kB) free!
[1584042305] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16588044 kB) free!
[1584042905] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33327044 kB) free.
[1584043158] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (240152 kB) free!
[1584043278] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236724 kB) free!
[1584043398] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (226300 kB) free!
[1584043505] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11198856 kB) free!
[1584043625] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (7815284 kB) free!
[1584043745] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (4633676 kB) free!
[1584044345] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16174356 kB) free!
[1584044623] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.4% (1845532 kB) free!
[1584044743] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (396432 kB) free!
[1584044863] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410196 kB) free!
[1584044945] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29914000 kB) free.
[1584045198] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 62.6% (30869072 kB) free.
[1584045346] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 4751.24, 3032.00, 1321.89
[1584045346] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 425 processes
[1584045448] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(85) < 203.212.70.170 [13/Mar/2020:02:04:51 +0530] TCP 502 0 0 240.001 "10.147.212.143:3001, 10.147.212.155:3001, 10.147.212.142:3001, 10.147.212.154:3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.001, 60.000, 60.000"
[1584045466] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1483.61, 2612.53, 1394.36
[1584045466] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 473 processes
[1584045546] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10214656 kB) free!
[1584045569] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584045586] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 201.33, 1748.60, 1225.47
[1584045666] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.2% (6885912 kB) free!
[1584045786] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (3895012 kB) free!
[1584045918] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 703 processes
[1584046038] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 705 processes
[1584046082] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < webmachine error: path="/api/queues"
[1584046094] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9049800 kB) free!
[1584046158] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 705 processes
[1584046186] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.35, 235.76, 642.80
[1584046202] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1584046214] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (4898412 kB) free!
[1584046334] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (222344 kB) free!
[1584046386] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14358540 kB) free!
[1584046432] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9219896 kB) free!
[1584046552] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (4739504 kB) free!
[1584046672] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (474028 kB) free!
[1584046786] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.55, 32.93, 337.54
[1584047534] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 30.7% (15114052 kB) free.
[1584047586] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (7291932 kB) free!
[1584047904] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.6% (9951200 kB) free.
[1584048186] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (17068840 kB) free!
[1584048472] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.1% (42903340 kB) free.
[1584048558] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 696 processes
[1584048752] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(363) < 2020/03/13 03:01:32 [error] 12255#0: *69821801 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1584048773] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 13.8% (3377564 kB) free!
[1584048786] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29180156 kB) free.
[1584048872] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584048943] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(286) < 103.16.101.72 [13/Mar/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1584049065] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584049168] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3465) < 203.212.70.170 [13/Mar/2020:03:01:42 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1584049288] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584049386] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9631012 kB) free!
[1584049506] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6382820 kB) free!
[1584049626] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32827920 kB) free.
[1584050226] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (4643832 kB) free!
[1584050346] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (871080 kB) free!
[1584050466] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (423040 kB) free!
[1584051666] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.8% (23526756 kB) free!
[1584052267] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (4286380 kB) free!
[1584052373] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2203936 kB) free!
[1584052867] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19848944 kB) free!
[1584053467] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29829884 kB) free.
[1584054667] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20332852 kB) free!
[1584054787] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (16913388 kB) free!
[1584054907] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29167120 kB) free.
[1584055507] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21060824 kB) free!
[1584055627] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.4% (17724372 kB) free!
[1584055747] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32599916 kB) free.
[1584056304] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198420 kB) free!
[1584056347] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (22759668 kB) free!
[1584056426] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192216 kB) free!
[1584056467] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.8% (19556124 kB) free!
[1584056546] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169416 kB) free!
[1584056587] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39084596 kB) free.
[1584057173] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 78.7% (19293876 kB) free.
[1584057187] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (23603400 kB) free!
[1584057307] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.5% (20402040 kB) free!
[1584057428] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39104832 kB) free.
[1584057669] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 50.4% (12434612 kB) free.
[1584058027] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23776540 kB) free!
[1584058147] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.6% (20501664 kB) free!
[1584058267] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39091160 kB) free.
[1584058867] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23382320 kB) free!
[1584058987] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.1% (19943184 kB) free!
[1584059107] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39082040 kB) free.
[1584059672] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584059707] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (21965408 kB) free!
[1584059793] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584059827] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.0% (18440456 kB) free!
[1584059947] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (38969412 kB) free.
[1584060547] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21774264 kB) free!
[1584060667] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.8% (18251888 kB) free!
[1584060787] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38872952 kB) free.
[1584060808] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584060934] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 63.7% (31411732 kB) free.
[1584061311] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 376 processes
[1584061387] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (21022908 kB) free!
[1584061431] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 372 processes
[1584061507] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.4% (17601760 kB) free!
[1584061627] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38214712 kB) free.
[1584062228] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21437944 kB) free!
[1584062263] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.7% (21988580 kB) free!
[1584062348] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (18099384 kB) free!
[1584062441] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.4% (21424408 kB) free.
[1584062466] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 29.3% (14430500 kB) free.
[1584062468] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37809436 kB) free.
[1584062495] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 86.3% (21146172 kB) free.
[1584062863] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401228 kB) free!
[1584063068] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17817832 kB) free!
[1584063069] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (4489508 kB) free!
[1584063188] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.5% (13895040 kB) free!
[1584063189] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (4471444 kB) free!
[1584063308] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37161032 kB) free.
[1584063309] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (4461232 kB) free!
[1584063463] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 42.9% (56616388 kB) free.
[1584063908] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17912860 kB) free!
[1584063909] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (218192 kB) free!
[1584064028] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (13188620 kB) free!
[1584064148] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36769204 kB) free.
[1584064748] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18034272 kB) free!
[1584064868] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (38977200 kB) free.
[1584065468] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23279468 kB) free!
[1584065588] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.1% (19856112 kB) free!
[1584065708] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38910316 kB) free.
[1584066308] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22399680 kB) free!
[1584066429] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (18894476 kB) free!
[1584066549] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38108064 kB) free.
[1584066688] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(184) < 172.105.59.173 [13/Mar/2020:07:56:30 +0530] TCP 502 0 0 120.001 "10.147.212.142:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1584066808] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584067149] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18650040 kB) free!
[1584067269] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38347580 kB) free.
[1584067869] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (21979056 kB) free!
[1584067989] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.1% (18548588 kB) free!
[1584068109] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37842572 kB) free.
[1584068681] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1584068681] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.4% (21068072 kB) free.
[1584068709] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14915576 kB) free!
[1584068829] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (11142580 kB) free!
[1584068949] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584068949] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37417100 kB) free.
[1584069549] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19784252 kB) free!
[1584069669] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.3% (16178464 kB) free!
[1584069789] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36934396 kB) free.
[1584070352] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584070389] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (17167688 kB) free!
[1584070482] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584070509] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (13036324 kB) free!
[1584070538] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584070577] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584070584] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584070602] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584070607] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 22.65, 21.06, 14.66
[1584070630] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34430736 kB) free.
[1584070668] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584070695] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584070700] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584070707] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584070708] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1584070727] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584070798] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584070820] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584070825] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584070830] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584070837] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584070950] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584070955] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584070960] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584071080] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584071212] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584071229] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14261180 kB) free!
[1584071314] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584071332] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584071336] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 22.75, 25.05, 20.77
[1584071349] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (10571424 kB) free!
[1584071399] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 63.1% (31120972 kB) free.
[1584071435] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1584071443] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584071462] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584071469] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33353496 kB) free.
[1584071582] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584072009] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584072009] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584072053] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584072069] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13242176 kB) free!
[1584072072] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.3% (9009104 kB) free!
[1584072139] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584072139] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1584072164] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 194855 MB (86% inode=99%):
[1584072168] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2283 processes
[1584072174] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584072189] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9265796 kB) free!
[1584072192] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584072192] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (9009052 kB) free!
[1584072269] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584072269] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1584072309] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32237328 kB) free.
[1584072312] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584072312] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.3% (9007608 kB) free!
[1584072644] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584072765] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1584072887] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2294 processes
[1584072909] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10961316 kB) free!
[1584073029] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7374324 kB) free!
[1584073149] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31507684 kB) free.
[1584073374] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584073374] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584073378] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584073383] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584073494] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 188087 MB (83% inode=99%):
[1584073494] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1584073500] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2296 processes
[1584073504] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1584073742] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584073749] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10932320 kB) free!
[1584073863] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 66.46, 41.65, 30.53
[1584073869] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (7112888 kB) free!
[1584073983] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 31.76, 36.86, 30.14
[1584073989] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3422920 kB) free!
[1584074070] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38774 MB (75% inode=98%):
[1584074070] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1584074070] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.6% (29864400 kB) free.
[1584074583] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 23.01, 22.37, 25.18
[1584075190] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.1% (22531840 kB) free!
[1584075284] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1584075790] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (3202872 kB) free!
[1584075884] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.6% (3112756 kB) free!
[1584075912] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (3357200 kB) free!
[1584076004] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.8% (2920400 kB) free!
[1584076010] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1584076010] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 43.7% (57585544 kB) free.
[1584076059] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 77.9% (38387484 kB) free.
[1584076124] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176432 kB) free!
[1584076243] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.2% (4216744 kB) free!
[1584076363] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.8% (3875012 kB) free!
[1584076484] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.1% (3457324 kB) free!
[1584076649] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.1% (20760676 kB) free.
[1584077084] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (1473944 kB) free!
[1584077255] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.1% (21581196 kB) free.
[1584077813] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24496876 kB) free!
[1584077933] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.0% (50144732 kB) free.
[1584078488] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 11.8% (2645212 kB) free!
[1584078558] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 711 processes
[1584078608] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 10.0% (2245532 kB) free!
[1584078678] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 711 processes
[1584078703] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.7% (4092964 kB) free!
[1584078728] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1869568 kB) free!
[1584078776] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.6% (3815448 kB) free!
[1584078783] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 27.26, 32.10, 30.05
[1584078798] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 710 processes
[1584078823] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.6% (3833944 kB) free!
[1584078896] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.8% (3616896 kB) free!
[1584078903] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 21.57, 28.17, 28.86
[1584078943] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.7% (3594164 kB) free!
[1584079016] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 13.8% (3385052 kB) free!
[1584079133] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22915728 kB) free!
[1584079253] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (17516052 kB) free!
[1584079305] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 16.0% (3939260 kB) free!
[1584079373] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (12490792 kB) free!
[1584079425] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 13.7% (3381264 kB) free!
[1584079503] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 28.86, 31.00, 30.12
[1584079543] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2081540 kB) free!
[1584079545] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 11.4% (2810092 kB) free!
[1584079616] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2154976 kB) free!
[1584079623] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 26.58, 29.82, 29.83
[1584079653] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172408 kB) free!
[1584079773] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177692 kB) free!
[1584079893] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193588 kB) free!
[1584080070] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7338780 kB) free!
[1584080145] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186484 kB) free!
[1584080190] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 13.7% (6763500 kB) free!
[1584080223] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 39.91, 36.56, 33.00
[1584080266] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 342 processes
[1584080310] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 12.5% (6179312 kB) free!
[1584080343] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 41.23, 38.17, 34.03
[1584080382] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 19.0% (6634096 kB) free!
[1584080386] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 341 processes
[1584080463] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 42.38, 39.57, 35.03
[1584080502] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 17.5% (6100796 kB) free!
[1584080506] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1584080623] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 16.0% (5572940 kB) free!
[1584080910] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3401532 kB) free!
[1584080954] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42242 MB (18% inode=99%):
[1584081074] SERVICE ALERT: primary-server;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46273 MB (20% inode=99%):
[1584081128] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.0% (15728236 kB) free.
[1584081194] SERVICE ALERT: primary-server;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45921 MB (20% inode=99%):
[1584081223] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 91.7% (31949704 kB) free.
[1584081280] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584081330] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1584081330] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.0% (20476660 kB) free.
[1584081368] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 48.7% (64212664 kB) free.
[1584081400] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38773 MB (75% inode=98%):
[1584081431] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.3% (30216616 kB) free.
[1584081462] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.7% (20763460 kB) free.
[1584081463] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.3% (25482692 kB) free!
[1584081494] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 83.9% (20561152 kB) free.
[1584081583] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.4% (29562812 kB) free.
[1584082064] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.6% (4309060 kB) free!
[1584082066] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 10.8% (5337988 kB) free!
[1584082183] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 20.0% (26344076 kB) free!
[1584082184] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.8% (3878912 kB) free!
[1584082186] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 13.5% (6652124 kB) free!
[1584082303] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.0% (32995720 kB) free.
[1584082304] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1584082304] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.1% (3452876 kB) free!
[1584082307] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.5% (6652916 kB) free!
[1584082545] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 91.7% (22645484 kB) free.
[1584082672] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.0% (7899820 kB) free!
[1584082792] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.1% (7465472 kB) free!
[1584082903] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.2% (25329972 kB) free!
[1584082904] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (1425688 kB) free!
[1584082907] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1584082907] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.4% (2636792 kB) free!
[1584082912] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 14.3% (7068384 kB) free!
[1584083023] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.2% (31898440 kB) free.
[1584083134] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (4550916 kB) free!
[1584083174] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14910208 kB) free!
[1584083244] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1584083254] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.8% (4397476 kB) free!
[1584083294] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11029228 kB) free!
[1584083364] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1584083374] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.2% (4247092 kB) free!
[1584083414] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (7197920 kB) free!
[1584083472] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1584083472] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 27.28, 26.27, 29.28
[1584083484] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1584083506] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 347 processes
[1584083623] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.4% (22920684 kB) free!
[1584083743] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 11.7% (15486056 kB) free!
[1584083863] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.0% (30347088 kB) free.
[1584083904] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173516 kB) free!
[1584084015] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.3% (24063912 kB) free!
[1584084024] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192000 kB) free!
[1584084072] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 37.85, 30.99, 30.25
[1584084112] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (3333352 kB) free!
[1584084139] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8568384 kB) free!
[1584084144] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1584084144] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188764 kB) free!
[1584084192] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 39.22, 34.74, 31.78
[1584084259] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8217552 kB) free!
[1584084313] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 40.29, 36.67, 32.85
[1584084322] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.4% (21675128 kB) free.
[1584084367] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.2% (39038184 kB) free.
[1584084379] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 16.0% (7894768 kB) free!
[1584084401] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 65.8% (32436460 kB) free.
[1584084429] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 63.8% (15628772 kB) free.
[1584084463] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.5% (24383664 kB) free!
[1584084574] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2036228 kB) free!
[1584084583] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.2% (33215660 kB) free.
[1584084615] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (2204004 kB) free!
[1584084684] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1584084706] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 353 processes
[1584084794] SERVICE ALERT: primary-server;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 112874 MB (50% inode=99%):
[1584084897] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584084916] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 32.32, 47.71, 43.32
[1584084922] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1584085017] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2299 processes
[1584085043] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1584085517] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 24.23, 32.74, 37.97
[1584085579] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4907292 kB) free!
[1584085783] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.3% (20212204 kB) free!
[1584085903] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (12643056 kB) free!
[1584086023] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.7% (28594044 kB) free.
[1584086415] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.4% (13721828 kB) free!
[1584086623] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.6% (25862080 kB) free!
[1584086717] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 19.36, 21.52, 27.06
[1584086743] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.2% (33254136 kB) free.
[1584086772] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 12.2% (6014500 kB) free!
[1584086865] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4761384 kB) free!
[1584086892] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 12.0% (5902504 kB) free!
[1584086985] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.0% (4405544 kB) free!
[1584087012] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 11.6% (5703776 kB) free!
[1584087015] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389056 kB) free!
[1584087105] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.8% (4116696 kB) free!
[1584087262] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 60.7% (14974852 kB) free.
[1584087309] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 91.3% (44989564 kB) free.
[1584087342] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.1% (22519284 kB) free!
[1584087345] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31632552 kB) free.
[1584087613] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4921548 kB) free!
[1584088305] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1584088305] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (399456 kB) free!
[1584088424] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (3107300 kB) free!
[1584088543] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.7% (20726672 kB) free!
[1584088544] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.7% (3519844 kB) free.
[1584088554] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10102976 kB) free!
[1584088663] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 10.4% (13768616 kB) free!
[1584088674] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32454788 kB) free.
[1584088784] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.2% (30540204 kB) free.
[1584088873] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.0% (38928520 kB) free.
[1584088919] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.0% (20581424 kB) free.
[1584088928] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4299992 kB) free!
[1584089049] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4024920 kB) free!
[1584089144] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.8% (2087348 kB) free!
[1584089169] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.9% (3784692 kB) free!
[1584089264] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.1% (2621592 kB) free!
[1584089274] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9420292 kB) free!
[1584089384] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3176984 kB) free!
[1584089394] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (4986084 kB) free!
[1584089514] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (824976 kB) free!
[1584089662] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177896 kB) free!
[1584089782] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176320 kB) free!
[1584089903] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (329036 kB) free!
[1584089983] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.9% (24976024 kB) free!
[1584090084] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8073088 kB) free!
[1584090103] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.4% (32210488 kB) free.
[1584090114] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26156568 kB) free!
[1584090204] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 12.4% (6089240 kB) free!
[1584090324] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.9% (5392788 kB) free!
[1584090369] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1584090369] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (966004 kB) free!
[1584090434] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4647716 kB) free!
[1584090554] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.5% (4535244 kB) free!
[1584090584] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.2% (4251512 kB) free.
[1584090674] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (4436048 kB) free!
[1584090714] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (1839156 kB) free!
[1584090924] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4685312 kB) free!
[1584090945] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4714020 kB) free!
[1584091065] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4429696 kB) free!
[1584091185] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 16.8% (4154532 kB) free!
[1584091314] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.0% (23720844 kB) free!
[1584091442] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.8% (2562368 kB) free!
[1584091562] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.8% (2573016 kB) free!
[1584091633] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8513940 kB) free!
[1584091683] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 15.7% (2556276 kB) free!
[1584091753] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.5% (8155968 kB) free!
[1584091754] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;1;CRITICAL - 1.4% (334956 kB) free!
[1584091874] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.1% (29634956 kB) free.
[1584091874] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;2;CRITICAL - 1.3% (330880 kB) free!
[1584091914] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389796 kB) free!
[1584091994] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 1.3% (328392 kB) free!
[1584092008] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(154) < 92.42.104.166 [13/Mar/2020:15:00:23 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1584092126] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.1% (3947820 kB) free!
[1584092128] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584092169] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.2% (15760520 kB) free.
[1584092246] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.3% (3748980 kB) free!
[1584092271] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 18.3% (4477436 kB) free!
[1584092331] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 18.3% (4477344 kB) free!
[1584092366] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.7% (3594032 kB) free!
[1584092384] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (3034200 kB) free!
[1584092385] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (1445524 kB) free!
[1584092391] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 18.2% (4471892 kB) free!
[1584092451] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 18.2% (4468440 kB) free!
[1584092473] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584092474] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191968 kB) free!
[1584092504] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.2% (3778116 kB) free.
[1584092515] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17238832 kB) free!
[1584092523] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182672 kB) free!
[1584092643] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (209792 kB) free!
[1584092763] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166860 kB) free!
[1584092985] SERVICE FLAPPING ALERT: db1-vm2;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1584092985] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.2% (22767604 kB) free.
[1584093103] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 11.6% (15279684 kB) free!
[1584093115] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403704 kB) free!
[1584093223] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.5% (24431044 kB) free!
[1584093343] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.9% (18258308 kB) free!
[1584093566] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (1582876 kB) free!
[1584094304] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.5% (2194836 kB) free!
[1584094425] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (188744 kB) free!
[1584094545] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.0% (1782320 kB) free!
[1584094915] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29213712 kB) free.
[1584095145] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1584095145] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1267280 kB) free!
[1584095515] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (3507176 kB) free!
[1584095635] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (380568 kB) free!
[1584095743] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.0% (29012892 kB) free.
[1584095755] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393184 kB) free!
[1584096345] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3321596 kB) free.
[1584096355] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16962620 kB) free!
[1584096955] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (423192 kB) free!
[1584097517] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584097545] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (2952656 kB) free!
[1584097555] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27761612 kB) free.
[1584097665] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.4% (4285276 kB) free.
[1584097906] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1584098143] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9477240 kB) free!
[1584098155] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (2787576 kB) free!
[1584098263] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 12.9% (16962708 kB) free!
[1584098275] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (442692 kB) free!
[1584098383] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (9117180 kB) free!
[1584098395] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (469416 kB) free!
[1584098865] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (1443380 kB) free!
[1584098983] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1584098983] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.6% (16639820 kB) free!
[1584098985] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (3094784 kB) free!
[1584098995] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17819016 kB) free!
[1584099105] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.5% (4151020 kB) free.
[1584099106] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 350 processes
[1584099503] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1584099596] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382160 kB) free!
[1584099705] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (793944 kB) free!
[1584099825] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (178520 kB) free!
[1584099945] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.4% (3960728 kB) free.
[1584100196] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.6% (21878824 kB) free!
[1584100546] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (1040800 kB) free!
[1584100563] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1584100666] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.9% (3891080 kB) free.
[1584100796] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (444052 kB) free!
[1584101396] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26145948 kB) free!
[1584101804] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9860600 kB) free!
[1584101835] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.7% (20417472 kB) free.
[1584101866] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (148224 kB) free!
[1584101886] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.7% (21743776 kB) free.
[1584101924] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.4% (9568856 kB) free!
[1584101964] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.0% (38434020 kB) free.
[1584101986] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.1% (1799292 kB) free!
[1584101996] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (448528 kB) free!
[1584102014] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.0% (20833944 kB) free.
[1584102045] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.9% (9318472 kB) free!
[1584102074] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 83.8% (20542152 kB) free.
[1584102106] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.1% (3924312 kB) free.
[1584102131] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42804072 kB) free.
[1584102169] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 89.7% (21977260 kB) free.
[1584102198] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584102369] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.2% (3854608 kB) free!
[1584102448] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194860 kB) free!
[1584102489] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.1% (3620288 kB) free!
[1584102568] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170136 kB) free!
[1584102609] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584102609] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.6% (15638124 kB) free.
[1584102688] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175004 kB) free!
[1584102798] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 703 processes
[1584102918] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1584103066] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.6% (20377668 kB) free.
[1584103182] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (426996 kB) free!
[1584103251] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 1.0% (249452 kB) free!
[1584103256] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.2% (39042352 kB) free.
[1584103331] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18272360 kB) free!
[1584103451] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.8% (14220892 kB) free!
[1584103571] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (10183180 kB) free!
[1584103673] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187368 kB) free!
[1584103717] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40172892 kB) free.
[1584103782] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.0% (20236444 kB) free.
[1584103865] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (254476 kB) free!
[1584103906] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (2940116 kB) free!
[1584103986] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (238984 kB) free!
[1584104027] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.3% (2162644 kB) free!
[1584104094] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 83.6% (20480212 kB) free.
[1584104099] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1584104106] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (388024 kB) free!
[1584104147] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.8% (3535996 kB) free.
[1584104158] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.3% (39065304 kB) free.
[1584104219] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1584104283] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (251812 kB) free!
[1584104385] SERVICE FLAPPING ALERT: db1-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584104386] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185604 kB) free!
[1584104400] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.7% (14092032 kB) free.
[1584104442] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.4% (19843080 kB) free.
[1584104446] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (4704612 kB) free!
[1584104628] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.7% (44208096 kB) free.
[1584104672] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/03/13 18:26:10 [error] 12257#0: *70139297 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1584104719] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1584104792] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584104840] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1584104925] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14747496 kB) free!
[1584105045] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (11006308 kB) free!
[1584105046] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188580 kB) free!
[1584105165] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (7296960 kB) free!
[1584105166] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196864 kB) free!
[1584105286] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1584105286] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168996 kB) free!
[1584105304] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 16.1% (3945636 kB) free!
[1584105363] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (2596572 kB) free!
[1584105364] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 16.1% (3942160 kB) free!
[1584105424] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 16.1% (3943828 kB) free!
[1584105483] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (2376568 kB) free!
[1584105484] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 16.0% (3933452 kB) free!
[1584105603] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4852428 kB) free!
[1584105625] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189736 kB) free!
[1584105745] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193804 kB) free!
[1584105865] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182928 kB) free!
[1584105947] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3224928 kB) free!
[1584106067] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 26.4% (4296160 kB) free.
[1584106285] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1584106528] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(31) < 103.255.146.164 [13/Mar/2020:19:02:31 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1584106648] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1584106693] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1584106965] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23234876 kB) free!
[1584107565] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (392580 kB) free!
[1584108166] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27697500 kB) free.
[1584108467] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584108583] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 27.3% (36029604 kB) free.
[1584108766] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (452468 kB) free!
[1584108886] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (446200 kB) free!
[1584109006] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (450120 kB) free!
[1584109067] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.3% (2814084 kB) free!
[1584109188] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.1% (3759444 kB) free.
[1584109606] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26771964 kB) free.
[1584109787] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.6% (2702816 kB) free!
[1584109907] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3506224 kB) free.
[1584110206] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (405620 kB) free!
[1584110326] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (444240 kB) free!
[1584110446] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (454424 kB) free!
[1584110507] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.4% (1845592 kB) free!
[1584110627] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.0% (2604708 kB) free!
[1584110747] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (3161148 kB) free!
[1584110983] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.2% (22622164 kB) free!
[1584111046] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (23179888 kB) free!
[1584111103] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.6% (31130268 kB) free.
[1584111232] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 60.5% (14930820 kB) free.
[1584111269] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 58.2% (28701012 kB) free.
[1584111300] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.6% (10207856 kB) free.
[1584111646] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (424032 kB) free!
[1584111703] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.3% (25431768 kB) free!
[1584111824] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.6% (33757092 kB) free.
[1584111834] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.0% (3199176 kB) free!
[1584111904] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183520 kB) free!
[1584111947] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3687064 kB) free.
[1584111955] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163104 kB) free!
[1584112024] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176824 kB) free!
[1584112075] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194188 kB) free!
[1584112084] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 53.5% (13122936 kB) free.
[1584112144] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (162144 kB) free!
[1584112246] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (33959512 kB) free.
[1584112476] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (4057688 kB) free!
[1584112547] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (1450404 kB) free!
[1584112596] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.9% (3913428 kB) free!
[1584112667] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.7% (1907884 kB) free!
[1584112716] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3726620 kB) free!
[1584112787] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.5% (2516992 kB) free!
[1584112846] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6569684 kB) free!
[1584112966] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2246656 kB) free!
[1584113086] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401812 kB) free!
[1584113387] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1584113387] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3391640 kB) free.
[1584113686] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33126816 kB) free.
[1584114287] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5823624 kB) free!
[1584114407] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (1681732 kB) free!
[1584114527] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (424264 kB) free!
[1584114823] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1584115127] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28873720 kB) free.
[1584115693] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.5% (4790204 kB) free!
[1584115727] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.0% (3915692 kB) free!
[1584115813] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.4% (4761796 kB) free!
[1584115847] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (410576 kB) free!
[1584115933] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.3% (4731084 kB) free!
[1584115967] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (446772 kB) free!
[1584116105] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.3% (20302524 kB) free.
[1584116289] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27587656 kB) free.
[1584116322] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 56.4% (27803952 kB) free.
[1584116358] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 57.6% (14123776 kB) free.
[1584116705] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191208 kB) free!
[1584116825] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191792 kB) free!
[1584116924] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 10.8% (5343012 kB) free!
[1584116945] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164908 kB) free!
[1584117044] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (3806676 kB) free!
[1584117164] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (2708104 kB) free!
[1584117494] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12412036 kB) free!
[1584117614] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8477852 kB) free!
[1584117735] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.8% (24832752 kB) free!
[1584118306] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1584118335] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (399164 kB) free!
[1584120735] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14809516 kB) free!
[1584121335] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (444152 kB) free!
[1584121935] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.7% (16764960 kB) free!
[1584122535] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (439868 kB) free!
[1584123135] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19918788 kB) free!
[1584123735] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382244 kB) free!