[1572892639] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181712 kB) free!
[1572892691] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.5% (27004772 kB) free.
[1572892759] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191396 kB) free!
[1572892879] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176156 kB) free!
[1572893291] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11159944 kB) free!
[1572893311] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1572893412] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34425988 kB) free.
[1572893431] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1572894011] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17896056 kB) free!
[1572894131] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.9% (14424312 kB) free!
[1572894251] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (11097776 kB) free!
[1572894851] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19254668 kB) free!
[1572895452] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27859712 kB) free.
[1572896052] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11092700 kB) free!
[1572896172] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33058832 kB) free.
[1572896772] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15167940 kB) free!
[1572896892] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28715636 kB) free.
[1572897492] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18161756 kB) free!
[1572897612] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.1% (14627792 kB) free!
[1572897732] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35252712 kB) free.
[1572898332] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16335748 kB) free!
[1572898452] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29467128 kB) free.
[1572899052] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18175860 kB) free!
[1572899172] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.1% (14572016 kB) free!
[1572899292] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35530184 kB) free.
[1572899489] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(72) < 2019/11/05 01:57:18 [error] 15228#0: *1344001 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:201/17, bytes from/to upstream:17/201
[1572899609] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572899742] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 41.8% (10321936 kB) free.
[1572899892] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16412272 kB) free!
[1572900012] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35660792 kB) free.
[1572900612] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17828428 kB) free!
[1572900726] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (2243448 kB) free!
[1572900733] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.9% (14379504 kB) free!
[1572900846] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (244708 kB) free!
[1572900853] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34560716 kB) free.
[1572900942] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172664 kB) free!
[1572900966] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (232188 kB) free!
[1572901062] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193168 kB) free!
[1572901182] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183772 kB) free!
[1572901225] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 70.3% (34625552 kB) free.
[1572901409] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(85) < 2019/11/05 02:26:02 [error] 15228#0: *1346396 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1572901453] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15289664 kB) free!
[1572901529] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572901566] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.0% (45337540 kB) free.
[1572901573] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35045788 kB) free.
[1572902173] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (16005960 kB) free!
[1572902293] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (36042360 kB) free.
[1572902893] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17409812 kB) free!
[1572903013] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36154892 kB) free.
[1572903613] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18753380 kB) free!
[1572903632] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1572903733] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (15198936 kB) free!
[1572903740] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1938) < 2019/11/05 03:11:41 [error] 24133#0: *2628432 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:67/0, bytes from/to upstream:0/67
[1572903752] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1572903781] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1638) < 203.212.70.170 [05/Nov/2019:03:11:32 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1572903853] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34507500 kB) free.
[1572903860] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(7) < 2019/11/05 03:12:22 [error] 24130#0: *2629046 upstream timed out (110: Connection timed out) while connecting to upstream, client: 202.65.140.29, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1572903902] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572903929] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1117) < 2019/11/05 03:11:21 [error] 15229#0: *1351799 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1572903980] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572903982] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1003) < 103.255.146.164 [05/Nov/2019:03:10:13 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1572904050] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572904103] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572904352] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1572904453] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16517900 kB) free!
[1572904472] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1572904573] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35583440 kB) free.
[1572905072] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 401 processes
[1572905173] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17469460 kB) free!
[1572905193] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1572905293] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36448956 kB) free.
[1572905893] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17882144 kB) free!
[1572906014] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36520588 kB) free.
[1572906613] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18083196 kB) free!
[1572906733] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36623616 kB) free.
[1572907333] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17779568 kB) free!
[1572907453] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36689540 kB) free.
[1572908053] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17343332 kB) free!
[1572908173] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36083120 kB) free.
[1572908773] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (16961544 kB) free!
[1572908893] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35190480 kB) free.
[1572909493] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19808520 kB) free!
[1572909613] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1572909613] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34995440 kB) free.
[1572909992] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1572910112] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1572910213] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17207360 kB) free!
[1572910333] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36605212 kB) free.
[1572910933] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19310436 kB) free!
[1572911053] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (15926792 kB) free!
[1572911173] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35403096 kB) free.
[1572911774] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17350664 kB) free!
[1572911894] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36859724 kB) free.
[1572912494] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18750292 kB) free!
[1572912614] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37463680 kB) free.
[1572913214] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20792728 kB) free!
[1572913334] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.4% (17711652 kB) free!
[1572913454] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37563896 kB) free.
[1572914054] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19212896 kB) free!
[1572914174] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.7% (15437212 kB) free!
[1572914294] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36167808 kB) free.
[1572914634] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 79.9% (19718036 kB) free.
[1572914680] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.8% (22256004 kB) free.
[1572914726] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.2% (40527412 kB) free.
[1572914762] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 66.8% (16381680 kB) free.
[1572914809] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 94.2% (23095136 kB) free.
[1572914894] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16544676 kB) free!
[1572915014] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37701676 kB) free.
[1572915239] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (206140 kB) free!
[1572915359] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176524 kB) free!
[1572915479] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191144 kB) free!
[1572915614] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17496652 kB) free!
[1572915734] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37786188 kB) free.
[1572916334] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17601752 kB) free!
[1572916455] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37852480 kB) free.
[1572917054] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19275288 kB) free!
[1572917174] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37835576 kB) free.
[1572917774] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18183956 kB) free!
[1572917894] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37713724 kB) free.
[1572918059] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.7% (22639896 kB) free.
[1572918494] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17878732 kB) free!
[1572918614] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37607072 kB) free.
[1572919214] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19193416 kB) free!
[1572919334] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37612652 kB) free.
[1572919712] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1572919832] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1572919934] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18586284 kB) free!
[1572920054] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36603868 kB) free.
[1572920654] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18654436 kB) free!
[1572920774] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (36005608 kB) free.
[1572921974] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.5% (20466040 kB) free!
[1572922095] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37702644 kB) free.
[1572922694] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19592288 kB) free!
[1572922814] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37393952 kB) free.
[1572923414] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.8% (18234820 kB) free!
[1572923534] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35738024 kB) free.
[1572924734] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19104760 kB) free!
[1572924854] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37149144 kB) free.
[1572925454] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26140184 kB) free!
[1572925574] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34407316 kB) free.
[1572926774] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15606524 kB) free!
[1572926894] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34884256 kB) free.
[1572928094] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15238288 kB) free!
[1572928215] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35274824 kB) free.
[1572928259] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.2% (4000592 kB) free!
[1572928379] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.3% (3763652 kB) free!
[1572928442] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1572928499] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.8% (3662264 kB) free!
[1572928814] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11572628 kB) free!
[1572928934] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.8% (32637540 kB) free.
[1572929099] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1627360 kB) free!
[1572930032] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1572930134] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12269472 kB) free!
[1572930152] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1572930254] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33229032 kB) free.
[1572930854] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10393296 kB) free!
[1572930974] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30839564 kB) free.
[1572931033] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.4% (22562708 kB) free.
[1572931639] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184268 kB) free!
[1572931759] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174580 kB) free!
[1572931791] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.5% (20845888 kB) free.
[1572932174] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9696748 kB) free!
[1572932294] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31387844 kB) free.
[1572932402] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191284 kB) free!
[1572932522] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174076 kB) free!
[1572932642] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1572932642] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173484 kB) free!
[1572932894] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7695304 kB) free!
[1572933014] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28866100 kB) free.
[1572933678] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 14.1% (3313424 kB) free!
[1572933798] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 12.7% (3001720 kB) free!
[1572933918] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 11.1% (2614492 kB) free!
[1572934214] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8905100 kB) free!
[1572934335] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29281332 kB) free.
[1572934456] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.2% (22505948 kB) free.
[1572934518] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (688384 kB) free!
[1572934570] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.3% (4482796 kB) free!
[1572934690] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.7% (4347420 kB) free!
[1572934810] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.1% (4187912 kB) free!
[1572934934] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6670928 kB) free!
[1572935054] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28510052 kB) free.
[1572935163] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1572935283] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1572935399] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.8% (3108552 kB) free!
[1572935403] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1572935519] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 82.8% (18597136 kB) free.
[1572935654] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (4348580 kB) free!
[1572935690] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.1% (3692328 kB) free!
[1572935718] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 96.3% (22686032 kB) free.
[1572935774] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26602388 kB) free.
[1572935810] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.0% (3430612 kB) free!
[1572935930] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.0% (3191028 kB) free!
[1572935996] SERVICE ALERT: sig2-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.91, 6.56, 5.75
[1572936010] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2107604 kB) free!
[1572936116] SERVICE ALERT: sig2-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.79, 5.77, 5.55
[1572936236] SERVICE ALERT: sig2-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.90, 5.38, 5.45
[1572936261] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164288 kB) free!
[1572936331] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.4% (7568924 kB) free!
[1572936374] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.1% (2763604 kB) free!
[1572936381] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180952 kB) free!
[1572936451] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 16.7% (8232640 kB) free!
[1572936494] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.6% (24509676 kB) free!
[1572936501] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199208 kB) free!
[1572936530] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2016608 kB) free!
[1572936571] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.7% (8232816 kB) free!
[1572936603] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 11 users currently logged in
[1572936614] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19684408 kB) free!
[1572936836] SERVICE ALERT: sig2-vm2;CPU Load;OK;HARD;1;OK - load average: 2.93, 3.90, 4.60
[1572937436] SERVICE ALERT: sig2-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.66, 5.43, 5.04
[1572937557] SERVICE ALERT: sig2-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.14, 5.49, 5.10
[1572937677] SERVICE ALERT: sig2-vm2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.58, 5.84, 5.27
[1572937814] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27950032 kB) free.
[1572938277] SERVICE FLAPPING ALERT: sig2-vm2;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1572938277] SERVICE ALERT: sig2-vm2;CPU Load;OK;HARD;1;OK - load average: 2.89, 3.69, 4.42
[1572939014] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9258592 kB) free!
[1572939134] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28571432 kB) free.
[1572939734] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7200312 kB) free!
[1572939854] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27906352 kB) free.
[1572940073] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.7% (22387948 kB) free.
[1572940113] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.6% (22218100 kB) free.
[1572940160] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 64.7% (15864528 kB) free.
[1572940171] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4352380 kB) free!
[1572940229] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.2% (40490828 kB) free.
[1572940455] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5648912 kB) free!
[1572940575] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27120100 kB) free.
[1572941281] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163652 kB) free!
[1572941401] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181596 kB) free!
[1572941521] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170744 kB) free!
[1572941659] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.6% (22369056 kB) free.
[1572941811] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24576812 kB) free!
[1572941931] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.5% (20460056 kB) free!
[1572942051] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41598476 kB) free.
[1572943251] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22445100 kB) free!
[1572943371] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39438712 kB) free.
[1572943459] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.9% (3180452 kB) free!
[1572943566] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9178348 kB) free!
[1572943579] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.4% (3552372 kB) free!
[1572943686] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.2% (8982244 kB) free!
[1572943699] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.2% (3492132 kB) free!
[1572943807] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.8% (8790064 kB) free!
[1572943825] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9276188 kB) free!
[1572943945] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (9152668 kB) free!
[1572943972] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (22338940 kB) free!
[1572944065] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (9026012 kB) free!
[1572944092] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41597252 kB) free.
[1572944299] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (1697160 kB) free!
[1572944692] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21827484 kB) free!
[1572944812] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (41019104 kB) free.
[1572944813] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4794236 kB) free!
[1572944849] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(279) < 2019/11/05 14:32:49 [error] 15233#0: *1409188 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1572944933] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (4678904 kB) free!
[1572944969] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572945053] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (4583596 kB) free!
[1572945412] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21627984 kB) free!
[1572945532] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40524592 kB) free.
[1572945569] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/05 14:39:46 [error] 15233#0: *1410684 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:58/17, bytes from/to upstream:17/58
[1572945689] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572946132] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21491856 kB) free!
[1572946252] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40118976 kB) free.
[1572946852] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20536632 kB) free!
[1572946972] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39665080 kB) free.
[1572947042] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1572947162] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1572947181] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(39) < 2019/11/05 15:16:18 [error] 24131#0: *2760820 recv() failed (104: Connection reset by peer) while proxying connection, client: 121.241.242.109, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1572947283] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1572947301] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(31) < 2019/11/05 15:18:16 [error] 24130#0: *2761107 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572947407] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4263800 kB) free!
[1572947421] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(41) < 2019/11/05 15:20:05 [error] 24129#0: *2761147 recv() failed (104: Connection reset by peer) while proxying connection, client: 121.241.242.109, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1572947572] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.9% (19612620 kB) free!
[1572947692] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39163532 kB) free.
[1572947883] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1572947970] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4650908 kB) free!
[1572948090] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (4495964 kB) free!
[1572948210] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.7% (4346096 kB) free!
[1572948238] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572948289] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.6% (23340312 kB) free.
[1572948292] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18852592 kB) free!
[1572948368] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 94.8% (46694308 kB) free.
[1572948407] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 44.8% (58991428 kB) free.
[1572948464] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.7% (47145244 kB) free.
[1572948516] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 66.6% (16335308 kB) free.
[1572948573] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 94.6% (23192580 kB) free.
[1572948689] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/05 15:33:16 [error] 15225#0: *1419512 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:58/17, bytes from/to upstream:17/58
[1572948809] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572949409] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/05 15:46:47 [error] 15224#0: *1421415 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:58/17, bytes from/to upstream:17/58
[1572949529] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572949677] SERVICE FLAPPING ALERT: sig2-vm2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1572950090] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170048 kB) free!
[1572950210] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181828 kB) free!
[1572950330] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175688 kB) free!
[1572952529] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/05 16:44:47 [error] 15225#0: *1430934 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:58/17, bytes from/to upstream:17/58
[1572952610] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (22068984 kB) free!
[1572952649] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572952730] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.7% (41817484 kB) free.
[1572953320] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.9% (4393184 kB) free!
[1572953330] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21403352 kB) free!
[1572953440] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.2% (4219964 kB) free!
[1572953450] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.3% (41277064 kB) free.
[1572953560] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (3896784 kB) free!
[1572953850] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/05 17:02:47 [error] 15225#0: *1433888 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:58/17, bytes from/to upstream:17/58
[1572953970] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572954029] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9749332 kB) free!
[1572954050] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19537576 kB) free!
[1572954149] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.8% (9250544 kB) free!
[1572954170] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40483784 kB) free.
[1572954269] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (9661716 kB) free!
[1572954760] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1962540 kB) free!
[1572954770] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19808788 kB) free!
[1572954870] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (244864 kB) free!
[1572954890] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39783020 kB) free.
[1572955032] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.1% (40464792 kB) free.
[1572955048] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.0% (20493264 kB) free.
[1572955060] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.8% (22504248 kB) free.
[1572955522] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17905096 kB) free!
[1572955640] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.4% (7101244 kB) free!
[1572955642] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38897168 kB) free.
[1572955652] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166008 kB) free!
[1572955760] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (4599788 kB) free!
[1572955772] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175040 kB) free!
[1572955880] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1572955880] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (2880660 kB) free!
[1572955892] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179684 kB) free!
[1572956242] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17949272 kB) free!
[1572956362] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36522084 kB) free.
[1572956370] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/05 17:48:17 [error] 15229#0: *1440728 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:58/17, bytes from/to upstream:17/58
[1572956468] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (243128 kB) free!
[1572956491] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572956528] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (224368 kB) free!
[1572956588] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 70.4% (17252140 kB) free.
[1572956690] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.1% (40466300 kB) free.
[1572956728] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.3% (23015808 kB) free.
[1572956919] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (4765928 kB) free!
[1572956962] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (15283820 kB) free!
[1572957039] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4643456 kB) free!
[1572957082] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 19.8% (26045264 kB) free!
[1572957090] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/05 18:00:18 [error] 15232#0: *1442268 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: rq-2-1gbps.server, request: "GET /favicon.ico HTTP/1.1", host: "rq-2-1gbps.server:8091", referrer: "http://rq-2-1gbps.server:8091/"
[1572957159] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.9% (4395308 kB) free!
[1572957202] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.4% (21588204 kB) free!
[1572957210] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572957332] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (1675280 kB) free!
[1572957452] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183872 kB) free!
[1572957572] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178276 kB) free!
[1572957802] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26754768 kB) free.
[1572957810] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/05 18:10:17 [error] 15227#0: *1443950 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:58/17, bytes from/to upstream:17/58
[1572957898] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (2126912 kB) free!
[1572957930] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572958018] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (2639900 kB) free!
[1572958138] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (2633188 kB) free!
[1572958402] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7397056 kB) free!
[1572958522] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26819968 kB) free.
[1572959122] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7839692 kB) free!
[1572959242] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (26950564 kB) free.
[1572959253] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.1% (40473216 kB) free.
[1572959271] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.0% (22952836 kB) free.
[1572959290] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 64.0% (15680112 kB) free.
[1572959731] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/05 18:41:48 [error] 15222#0: *1448961 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:58/17, bytes from/to upstream:17/58
[1572959801] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (30.5% change >= 20.0% threshold)
[1572959801] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (35.1% change >= 20.0% threshold)
[1572959842] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8192212 kB) free!
[1572959851] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572959962] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27246824 kB) free.
[1572960240] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (26.1% change >= 20.0% threshold)
[1572960240] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (34.1% change >= 20.0% threshold)
[1572960299] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;NRPE: Command 'check_log_rq1rabbitmqerror' not defined
[1572960419] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;2;NRPE: Command 'check_log_rq1rabbitmqerror' not defined
[1572960539] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;HARD;3;NRPE: Command 'check_log_rq1rabbitmqerror' not defined
[1572960562] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8333152 kB) free!
[1572960634] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1572960634] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (33.2% change >= 20.0% threshold)
[1572960682] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27129588 kB) free.
[1572960696] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1572960696] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (33.2% change >= 20.0% threshold)
[1572960762] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1572960762] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (33.2% change >= 20.0% threshold)
[1572960827] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1572960827] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (33.2% change >= 20.0% threshold)
[1572960847] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1572961682] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185540 kB) free!
[1572961802] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163156 kB) free!
[1572961922] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197308 kB) free!
[1572962263] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (1385336 kB) free!
[1572962273] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (4505024 kB) free!
[1572962383] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (1385488 kB) free!
[1572962393] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4398392 kB) free!
[1572962503] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (250472 kB) free!
[1572962513] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.3% (4239824 kB) free!
[1572964320] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21777460 kB) free!
[1572964440] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40818672 kB) free.
[1572964913] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1848888 kB) free!
[1572965040] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (20995800 kB) free!
[1572965160] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41084424 kB) free.
[1572965320] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.1% (40464424 kB) free.
[1572965335] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.1% (22339968 kB) free.
[1572965347] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.5% (22834472 kB) free.
[1572965760] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23152256 kB) free!
[1572965880] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41675988 kB) free.
[1572965959] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.5% (2583976 kB) free!
[1572966079] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169228 kB) free!
[1572966199] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168272 kB) free!
[1572966480] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22605304 kB) free!
[1572966528] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (243044 kB) free!
[1572966600] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41512080 kB) free.
[1572966648] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (231368 kB) free!
[1572966768] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (242128 kB) free!
[1572967200] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24070820 kB) free!
[1572967320] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.4% (41366472 kB) free.
[1572967920] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23991944 kB) free!
[1572968040] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41141572 kB) free.
[1572968640] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (25129644 kB) free!
[1572968761] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.4% (21672436 kB) free!
[1572968881] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39697364 kB) free.
[1572969481] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.9% (22240108 kB) free!
[1572969601] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40242608 kB) free.
[1572970201] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (23027616 kB) free!
[1572970321] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40755884 kB) free.
[1572970921] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23521436 kB) free!
[1572971041] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (40663892 kB) free.
[1572971641] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (24636124 kB) free!
[1572971761] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (40599240 kB) free.
[1572972361] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25301832 kB) free!
[1572972481] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.7% (21988288 kB) free!
[1572972601] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39267088 kB) free.
[1572973201] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22918188 kB) free!
[1572973321] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40458396 kB) free.
[1572973700] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (4685540 kB) free!
[1572973820] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.9% (4629380 kB) free!
[1572973922] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23940444 kB) free!
[1572973940] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.6% (4560608 kB) free!
[1572974042] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40402736 kB) free.
[1572974642] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24901456 kB) free!
[1572974762] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37328980 kB) free.
[1572975341] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.3% (5705116 kB) free!
[1572975362] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25555544 kB) free!
[1572975461] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (3308804 kB) free!
[1572975482] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.0% (22433816 kB) free!
[1572975581] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (4124388 kB) free!
[1572975602] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39665304 kB) free.
[1572975693] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8173992 kB) free!
[1572975800] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1572975813] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 15.3% (7536696 kB) free!
[1572975933] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.5% (7136208 kB) free!
[1572976202] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24216508 kB) free!
[1572976322] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40270280 kB) free.
[1572976369] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1572976922] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26111956 kB) free!
[1572977042] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.7% (23308780 kB) free!
[1572977162] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39950224 kB) free.
[1572977762] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26231212 kB) free!
[1572977882] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.9% (23644028 kB) free!
[1572978002] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40256148 kB) free.