[1554143461] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (1145872 kB) free!
[1554143543] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (4749580 kB) free!
[1554143664] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (3448940 kB) free!
[1554143784] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (2048116 kB) free!
[1554144061] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.2% (7004572 kB) free!
[1554144384] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.4% (6623480 kB) free!
[1554144662] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3408864 kB) free!
[1554144984] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9952764 kB) free.
[1554145037] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1554145262] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.6% (8697156 kB) free!
[1554145584] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (3684172 kB) free!
[1554145637] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554145704] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.9% (2427860 kB) free!
[1554145757] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554145824] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (381772 kB) free!
[1554145877] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1554146424] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.4% (5124540 kB) free!
[1554146462] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (1287308 kB) free!
[1554146477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1554146597] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554146717] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1554147024] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9939404 kB) free.
[1554147062] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.2% (6498400 kB) free!
[1554147624] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (2712524 kB) free!
[1554147662] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3779700 kB) free!
[1554147744] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.9% (914452 kB) free!
[1554147864] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9946144 kB) free.
[1554147917] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1554148262] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.8% (8786940 kB) free!
[1554148465] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (4203968 kB) free!
[1554148517] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1554148585] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.9% (2920436 kB) free!
[1554148637] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554148705] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (1179572 kB) free!
[1554148757] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1554149305] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.8% (5803144 kB) free!
[1554149357] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554149462] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (3107892 kB) free!
[1554149477] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554149597] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554149905] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9946924 kB) free.
[1554150062] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.9% (8815288 kB) free!
[1554150505] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (3248208 kB) free!
[1554150625] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (1794344 kB) free!
[1554150745] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9947520 kB) free.
[1554150797] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1554151262] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4370024 kB) free!
[1554151345] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4694168 kB) free!
[1554151397] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1554151465] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (3381900 kB) free!
[1554151517] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1554151585] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (1981588 kB) free!
[1554151637] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554151863] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.9% (8832112 kB) free!
[1554152186] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.3% (6539700 kB) free!
[1554152786] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9940576 kB) free.
[1554152837] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1554153386] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (3619852 kB) free!
[1554153437] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 303 processes
[1554153506] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (2366888 kB) free!
[1554153558] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1554153626] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9949436 kB) free.
[1554153678] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1554154226] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.3% (5056748 kB) free!
[1554154263] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (2274064 kB) free!
[1554154278] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1554154346] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (3812796 kB) free!
[1554154398] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554154466] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (2562840 kB) free!
[1554154518] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1554155066] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.5% (7643852 kB) free!
[1554155666] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9934984 kB) free.
[1554155718] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1554156063] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.0% (8854780 kB) free!
[1554156266] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (4373912 kB) free!
[1554156318] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554156387] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (3108220 kB) free!
[1554156439] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554156507] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (1583348 kB) free!
[1554156559] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1554157107] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.8% (6283944 kB) free!
[1554157707] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9930384 kB) free.
[1554157759] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1554158307] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (3468148 kB) free!
[1554158359] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554158427] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (2173912 kB) free!
[1554158479] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554158547] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9928880 kB) free.
[1554158599] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 306 processes
[1554159147] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (4906428 kB) free!
[1554159199] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554159267] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (3628680 kB) free!
[1554159319] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1554159387] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (2366912 kB) free!
[1554159440] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554159987] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.3% (7050784 kB) free!
[1554160587] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9923496 kB) free.
[1554160640] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 304 processes
[1554161188] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (4052596 kB) free!
[1554161240] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 307 processes
[1554161308] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (2774244 kB) free!
[1554161360] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554161428] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (1001596 kB) free!
[1554161480] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1554162028] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.3% (5568324 kB) free!
[1554162080] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554162200] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554162321] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554162628] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9921008 kB) free.
[1554163228] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (3136260 kB) free!
[1554163348] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (1590396 kB) free!
[1554163468] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9928796 kB) free.
[1554163521] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1554163863] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (1657924 kB) free!
[1554164068] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (4560340 kB) free!
[1554164121] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554164188] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (3253856 kB) free!
[1554164241] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554164308] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (1794888 kB) free!
[1554164361] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554164909] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.0% (6409200 kB) free!
[1554165663] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.5% (8634440 kB) free!
[1554166109] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (378060 kB) free!
[1554166162] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1554166762] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 308 processes
[1554166863] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (1898280 kB) free!
[1554166882] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1554167464] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.6% (7678144 kB) free!
[1554167482] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554167602] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554167722] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554167909] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.2% (6025900 kB) free!
[1554168064] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4746344 kB) free!
[1554168206] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 29.92, 35.74, 18.72
[1554168276] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.7% (2856756 kB) free!
[1554168327] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.20, 23.98, 16.48
[1554168447] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.85, 16.14, 14.51
[1554168997] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16.18, 32.84, 18.77
[1554169109] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10099260 kB) free.
[1554169117] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.33, 22.05, 16.53
[1554169237] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.42, 14.79, 14.54
[1554169264] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.2% (7970876 kB) free!
[1554169443] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 327 processes
[1554169467] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.5% (2849984 kB) free!
[1554169647] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.23, 0.52, 4.17
[1554169709] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (2156864 kB) free!
[1554169830] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 18.1% (8922248 kB) free!
[1554169950] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.6% (9171668 kB) free!
[1554170437] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.07, 0.37, 4.09
[1554170464] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (1838688 kB) free!
[1554170550] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (339616 kB) free!
[1554171064] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.7% (7754452 kB) free!
[1554171322] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1554171664] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4203920 kB) free!
[1554171922] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1554172042] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1554172264] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.4% (8104560 kB) free!
[1554172350] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.4% (9571428 kB) free!
[1554172642] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554172763] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554172883] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554172950] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (1184660 kB) free!
[1554173464] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (2697048 kB) free!
[1554174064] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.8% (7802600 kB) free!
[1554174150] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.2% (8462872 kB) free!
[1554174665] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4505384 kB) free!
[1554174750] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10041356 kB) free.
[1554175351] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (2421688 kB) free!
[1554175471] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9945744 kB) free.
[1554175865] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.7% (6272584 kB) free!
[1554176071] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.3% (5092140 kB) free!
[1554176191] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.7% (3792596 kB) free!
[1554176311] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (2412936 kB) free!
[1554176465] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (2583644 kB) free!
[1554176483] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1554177065] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.4% (8071048 kB) free!
[1554177084] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1554177204] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1554177511] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.0% (5426292 kB) free!
[1554177804] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1554177924] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1554178265] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.6% (1788272 kB) free!
[1554178524] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1554178644] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554178711] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10071172 kB) free.
[1554178764] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1554178865] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.7% (7713664 kB) free!
[1554179076] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2426672 kB) free!
[1554179311] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (4007848 kB) free!
[1554179364] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1554179432] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (2744844 kB) free!
[1554179465] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4538156 kB) free!
[1554179485] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1554179552] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (957232 kB) free!
[1554179605] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1554180205] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1554180325] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1554180445] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1554180665] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.0% (6889476 kB) free!
[1554180867] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2472164 kB) free!
[1554181265] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (3572184 kB) free!
[1554181352] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.9% (6838736 kB) free!
[1554181865] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.6% (8169788 kB) free!
[1554182552] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (711880 kB) free!
[1554183065] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (2848260 kB) free!
[1554183665] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.7% (8208764 kB) free!
[1554184352] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.3% (6545728 kB) free!
[1554184865] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (2122420 kB) free!
[1554185465] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.9% (8345444 kB) free!
[1554185553] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (716480 kB) free!
[1554186665] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (1657400 kB) free!
[1554186753] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.3% (5069132 kB) free!
[1554186843] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 335 processes
[1554187265] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.5% (8153460 kB) free!
[1554187443] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 331 processes
[1554187953] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.3% (10006908 kB) free.
[1554188465] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (1525460 kB) free!
[1554188553] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (1648180 kB) free!
[1554188673] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10104220 kB) free.
[1554189065] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.7% (7749564 kB) free!
[1554189273] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (1866004 kB) free!
[1554189393] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.5% (10089180 kB) free.
[1554189445] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1554189666] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4354616 kB) free!
[1554189993] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (4142164 kB) free!
[1554190114] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (4251044 kB) free!
[1554190124] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1554190234] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (3615268 kB) free!
[1554190244] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1554190364] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1554190866] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.2% (7467992 kB) free!
[1554191466] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4242192 kB) free!
[1554192034] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.8% (7269372 kB) free!
[1554192634] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10233404 kB) free.
[1554192666] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.0% (7400192 kB) free!
[1554193234] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (1727472 kB) free!
[1554193266] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4222752 kB) free!
[1554193354] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10371196 kB) free.
[1554193364] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1554193954] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (2456972 kB) free!
[1554194074] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (391128 kB) free!
[1554194195] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.2% (9480964 kB) free!
[1554194466] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.1% (7447016 kB) free!
[1554194707] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 61.17, 34.07, 14.43
[1554194795] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (907304 kB) free!
[1554194827] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 62.23, 43.44, 20.25
[1554194845] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1554194947] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 37.95, 43.81, 23.39
[1554195066] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4344956 kB) free!
[1554195225] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 12.0% (2941256 kB) free!
[1554195445] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1554195547] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.46, 6.43, 12.56
[1554195565] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1554195685] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1554196243] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 334 processes
[1554196266] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.8% (7765684 kB) free!
[1554196595] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.5% (7137460 kB) free!
[1554196866] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4807740 kB) free!
[1554197347] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.09, 0.25, 2.74
[1554197757] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 334 processes
[1554197795] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (1567108 kB) free!
[1554198043] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 336 processes
[1554198067] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.9% (8332336 kB) free!
[1554198685] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 316 processes
[1554198995] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.1% (5468476 kB) free!
[1554199267] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (2123900 kB) free!
[1554199285] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1554199405] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1554199557] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 325 processes
[1554199867] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.5% (8622436 kB) free!
[1554200005] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1554200125] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1554200195] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10256468 kB) free.
[1554200246] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1554200246] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1554200795] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (2210876 kB) free!
[1554200916] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10259916 kB) free.
[1554201067] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (3327296 kB) free!
[1554201515] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (2535368 kB) free!
[1554201635] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8296460 kB) free!
[1554201667] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.2% (7984692 kB) free!
[1554201755] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.0% (9356724 kB) free!
[1554202355] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (787764 kB) free!
[1554202867] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4268692 kB) free!
[1554204067] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 16.8% (8296936 kB) free!
[1554204155] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.3% (6572720 kB) free!
[1554204225] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2173380 kB) free!
[1554205268] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3830108 kB) free!
[1554205355] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (544392 kB) free!
[1554206468] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.7% (8707944 kB) free!
[1554207155] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 16.7% (8242300 kB) free!
[1554207668] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (2466872 kB) free!
[1554207755] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10245344 kB) free.
[1554208268] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.7% (9216476 kB) free!
[1554208355] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (1957904 kB) free!
[1554208475] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10227512 kB) free.
[1554209075] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (2266252 kB) free!
[1554209195] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (257068 kB) free!
[1554209246] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1554209315] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.5% (9631216 kB) free!
[1554209468] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (3979192 kB) free!
[1554209846] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1554209915] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (892600 kB) free!
[1554209966] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1554210086] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1554210668] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.3% (7558736 kB) free!
[1554211268] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 1.9% (946168 kB) free!
[1554211715] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.6% (7200424 kB) free!
[1554212315] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9886776 kB) free.
[1554212468] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 15.3% (7563052 kB) free!
[1554212915] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (1008036 kB) free!
[1554213035] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.8% (10268756 kB) free.
[1554213635] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (1604412 kB) free!
[1554213669] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (2484260 kB) free!
[1554213755] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.9% (10281680 kB) free.
[1554214036] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.9% (3067056 kB) free!
[1554214156] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.7% (3039120 kB) free!
[1554214269] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.3% (9494360 kB) free!
[1554214276] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.5% (3000428 kB) free!
[1554214356] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (2275868 kB) free!
[1554214476] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (233716 kB) free!
[1554214596] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.3% (9506588 kB) free!
[1554215196] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (851564 kB) free!
[1554215469] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4576472 kB) free!
[1554216669] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.8% (9283280 kB) free!
[1554216996] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.3% (7043436 kB) free!
[1554217869] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (3991656 kB) free!
[1554218196] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (1419700 kB) free!
[1554219069] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.2% (8970760 kB) free!
[1554219086] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1554219396] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.3% (5566832 kB) free!
[1554219686] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1554219806] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 310 processes
[1554219926] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1554220269] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3955164 kB) free!
[1554220596] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10038812 kB) free.
[1554220713] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (3231252 kB) free!
[1554220833] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.8% (3221232 kB) free!
[1554220953] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.8% (3218656 kB) free!
[1554221196] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (2823060 kB) free!
[1554221316] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (997668 kB) free!
[1554221436] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 20.0% (9860380 kB) free!
[1554221469] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.4% (9087480 kB) free!
[1554222036] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (1616580 kB) free!
[1554222669] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (3558068 kB) free!
[1554223236] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.4% (5118396 kB) free!
[1554223869] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.0% (8373856 kB) free!
[1554224436] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.4% (10068016 kB) free.
[1554225036] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (1786800 kB) free!
[1554225069] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (2747760 kB) free!
[1554225157] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10208656 kB) free.
[1554225669] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.9% (9307748 kB) free!
[1554225756] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (2220804 kB) free!
[1554225876] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10209156 kB) free.
[1554225927] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1554226476] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (3880432 kB) free!
[1554226527] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554226596] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.4% (2647976 kB) free!
[1554226647] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1554226716] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (695772 kB) free!
[1554226767] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1554227367] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 306 processes
[1554227469] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (2279152 kB) free!
[1554227487] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1554228069] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.0% (9367472 kB) free!
[1554228087] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1554228207] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554228328] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 309 processes
[1554228516] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.6% (6712676 kB) free!
[1554229116] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10198280 kB) free.
[1554229269] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (3651344 kB) free!
[1554229716] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4072932 kB) free!