[1549823794] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2321892 kB) free!
[1549824693] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 50 users currently logged in
[1549824881] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549824980] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1549825001] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549825023] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9578220 kB) free!
[1549825100] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1549825121] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549825143] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.9% (8841456 kB) free!
[1549825263] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11366512 kB) free.
[1549825700] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 405 processes
[1549825820] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 405 processes
[1549825863] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8392904 kB) free!
[1549825940] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 403 processes
[1549825983] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10877168 kB) free.
[1549826921] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1549827183] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9016120 kB) free!
[1549827303] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.1% (10910932 kB) free.
[1549828121] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549828241] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1549828503] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9024524 kB) free!
[1549828623] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.9% (11286480 kB) free.
[1549828841] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549828961] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549829081] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549829223] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8406436 kB) free!
[1549829343] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.0% (10851416 kB) free.
[1549829476] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.67, 11.14, 12.56
[1549829596] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.21, 8.17, 11.31
[1549829681] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1549829681] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1549829716] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.53, 6.29, 10.24
[1549829943] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7863672 kB) free!
[1549830063] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10334244 kB) free.
[1549830281] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549830402] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549830522] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549830663] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7700828 kB) free!
[1549830783] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.7% (10193300 kB) free.
[1549830916] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.40, 2.51, 4.63
[1549831383] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7071560 kB) free!
[1549831504] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.5% (9610772 kB) free!
[1549831624] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.3% (8512032 kB) free!
[1549832322] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1549833140] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 406 processes
[1549833424] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.0% (10326952 kB) free.
[1549833994] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1549834024] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9614556 kB) free!
[1549834123] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549834144] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8223972 kB) free!
[1549834243] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1549834264] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.6% (6715620 kB) free!
[1549834363] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549834963] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1549836064] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1549836064] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.1% (9881528 kB) free.
[1549836664] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.3% (6565220 kB) free!
[1549836784] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.0% (9367672 kB) free!
[1549836905] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 17.4% (8582280 kB) free!
[1549837363] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549837484] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549837604] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549839404] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1549839994] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.3% (3987504 kB) free!
[1549840604] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549840724] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549841924] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549842044] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1549842394] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 24.0% (5890820 kB) free.
[1549842644] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549842764] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1549843964] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549844084] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1549844685] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549844805] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1549846005] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549846125] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1549846725] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549846845] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1549847446] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549847566] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1549848166] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549848286] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1549848305] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1549848886] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549849006] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1549849606] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549849726] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1549850327] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549850447] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1549851047] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1549851647] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549851767] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1549852967] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549853087] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549853208] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1549853807] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549853927] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549854527] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549854647] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549855847] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549855967] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1549856568] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549856688] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549856808] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549859614] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 51.7% (25453880 kB) free.
[1549861594] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4734176 kB) free!
[1549861714] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (4672104 kB) free!
[1549861834] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (4594316 kB) free!
[1549861933] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.1% (22826228 kB) free.
[1549862533] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190664 kB) free!
[1549862653] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171504 kB) free!
[1549862773] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197492 kB) free!
[1549862808] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549862855] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1549862855] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.5% (22908744 kB) free.
[1549863408] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549863528] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549863648] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1549864064] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 14.7% (2385272 kB) free!
[1549864164] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55858 MB (24% inode=99%):
[1549864184] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 13.9% (2262656 kB) free!
[1549864249] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549864304] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 13.8% (2243168 kB) free!
[1549864369] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549864489] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549865244] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (988292 kB) free!
[1549865268] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.0% (3198264 kB) free!
[1549865364] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (976676 kB) free!
[1549865388] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.0% (3175348 kB) free!
[1549865484] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (965116 kB) free!
[1549865504] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (749020 kB) free!
[1549865508] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 12.9% (3152304 kB) free!
[1549867908] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2296044 kB) free!
[1549870022] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.5% (22439124 kB) free.
[1549870072] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.7% (15065464 kB) free.
[1549870118] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 91.9% (14932572 kB) free.
[1549871689] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1549872289] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549872410] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1549872530] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1549872530] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1549873730] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1549873770] SERVICE ALERT: sig4;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1549873773] HOST ALERT: sig4;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.56)
[1549873781] SERVICE ALERT: sig4;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sig4' was out of bounds)
[1549873781] SERVICE ALERT: sig4;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig4' was out of bounds)
[1549873781] SERVICE ALERT: sig4;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig4' was out of bounds)
[1549873837] HOST ALERT: sig4;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1549873890] SERVICE ALERT: sig4;Current Users;OK;HARD;1;USERS OK - 0 users currently logged in
[1549873901] SERVICE ALERT: sig4;CPU Load;OK;HARD;1;OK - load average: 0.29, 0.16, 0.06
[1549873901] SERVICE ALERT: sig4;HDD_Root status;OK;HARD;1;DISK OK - free space: / 37467 MB (73% inode=98%):
[1549873901] SERVICE ALERT: sig4;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 202567 MB (90% inode=99%):
[1549874330] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549874450] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549874570] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549876970] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1549877570] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549877690] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549877810] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549878410] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1549879010] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549879026] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1549879131] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1549879626] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.3% (310144 kB) free!
[1549879730] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549879746] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (223808 kB) free!
[1549879850] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1549879866] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174052 kB) free!
[1549880450] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549880570] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549880690] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549880897] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.8% (23002724 kB) free.
[1549881290] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1549881890] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549881993] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 51 users currently logged in
[1549882011] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549882053] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 51 users currently logged in
[1549882113] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 51 users currently logged in
[1549882131] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549882173] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 51 users currently logged in
[1549882473] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 48 users currently logged in
[1549882517] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.82, 14.16, 9.87
[1549882637] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.52, 9.88, 8.82
[1549882644] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1549882757] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.53, 7.12, 7.94
[1549882764] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1549882884] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1549883357] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.51, 2.44, 5.00
[1549885473] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 52 users currently logged in
[1549885533] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 52 users currently logged in
[1549885593] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 52 users currently logged in
[1549885653] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 52 users currently logged in
[1549885732] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549886313] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195468 kB) free!
[1549886332] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549886433] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185456 kB) free!
[1549886452] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549886553] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190060 kB) free!
[1549886572] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549888653] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 50 users currently logged in
[1549890153] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 54 users currently logged in
[1549890213] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 54 users currently logged in
[1549890273] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 54 users currently logged in
[1549890333] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 54 users currently logged in
[1549890753] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.0% (23049992 kB) free.
[1549891372] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549891972] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549892092] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549892692] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549892812] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1549893413] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1549893533] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1549893753] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.9% (3400356 kB) free!
[1549893874] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191900 kB) free!
[1549893994] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187904 kB) free!
[1549894133] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1549894253] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1549894833] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 49 users currently logged in
[1549894853] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549894973] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1549895223] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1549895223] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.1% (23061892 kB) free.
[1549895573] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549895693] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549895814] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1549895823] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165092 kB) free!
[1549895943] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (211388 kB) free!
[1549896063] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167980 kB) free!
[1549896079] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8754292 kB) free!
[1549896199] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.1% (11891644 kB) free.
[1549896476] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.9% (2912400 kB) free!
[1549896596] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.8% (2890396 kB) free!
[1549896716] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.7% (2870628 kB) free!
[1549896933] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 51 users currently logged in
[1549896993] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 51 users currently logged in
[1549897014] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1549897053] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 52 users currently logged in
[1549897113] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 52 users currently logged in
[1549897614] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549897734] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1549897854] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549897999] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9135956 kB) free!
[1549898119] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.6% (12129144 kB) free.
[1549898454] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1549898719] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.0% (8383604 kB) free!
[1549898839] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.7% (11672468 kB) free.
[1549899054] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549899175] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1549899513] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 41 users currently logged in
[1549900071] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.4% (2020244 kB) free!
[1549900191] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.4% (2022568 kB) free!
[1549900311] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 12.5% (2023720 kB) free!
[1549900639] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.8% (8286400 kB) free!
[1549900759] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.0% (11355976 kB) free.
[1549900911] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 31.3% (5086484 kB) free.
[1549901131] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.7% (23219320 kB) free.
[1549901258] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.6% (15053948 kB) free.
[1549901574] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549901694] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1549901814] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1549903015] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1549903159] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9874888 kB) free!
[1549903279] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.6% (12608120 kB) free.
[1549903879] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9405504 kB) free!
[1549903999] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.0% (11336872 kB) free.
[1549907137] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1549907599] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9309012 kB) free!
[1549907719] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.4% (12506412 kB) free.
[1549907815] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1549907935] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1549908055] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1549908319] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9165008 kB) free!
[1549908439] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.0% (12339984 kB) free.
[1549908655] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1549909039] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9689420 kB) free!
[1549909159] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.4% (12493228 kB) free.
[1549909759] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9059532 kB) free!