[1559241001] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.8% (2910020 kB) free!
[1559241058] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.44, 68.92, 69.00
[1559241223] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.6% (3512192 kB) free.
[1559241267] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.90, 2.07, 4.01
[1559241271] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193536 kB) free!
[1559241578] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1559241601] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180068 kB) free!
[1559241823] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (565512 kB) free!
[1559241943] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (1380520 kB) free!
[1559242003] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4356748 kB) free!
[1559242063] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.1% (2770952 kB) free!
[1559242259] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.11, 1.40, 19.10
[1559242663] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.3% (3459812 kB) free.
[1559242778] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559242800] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 385 processes
[1559242898] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559243018] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559243156] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.0% (23440100 kB) free.
[1559243263] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.7% (3188564 kB) free!
[1559243270] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 95.0% (23283092 kB) free.
[1559243340] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (3424404 kB) free.
[1559243383] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (180628 kB) free!
[1559243400] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 381 processes
[1559243431] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.1% (22818960 kB) free.
[1559243503] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.4% (3470840 kB) free.
[1559243512] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.8% (15080932 kB) free.
[1559243520] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1559243640] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 377 processes
[1559243803] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.4% (9550800 kB) free!
[1559243945] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (332072 kB) free!
[1559244059] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.31, 0.31, 3.01
[1559244065] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.4% (223704 kB) free!
[1559244103] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.9% (2743036 kB) free!
[1559244185] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.3% (366004 kB) free!
[1559244223] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.0% (3401608 kB) free.
[1559244403] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (2818148 kB) free!
[1559244823] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (172280 kB) free!
[1559244840] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 377 processes
[1559244943] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (1370300 kB) free!
[1559245064] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.8% (3700412 kB) free.
[1559245418] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1559245440] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1559245560] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559245603] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.5% (7142500 kB) free!
[1559245663] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (172060 kB) free!
[1559245680] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1559245783] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 11.3% (1836828 kB) free!
[1559245903] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.6% (3503088 kB) free.
[1559246503] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 10.7% (1742424 kB) free!
[1559246618] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559246623] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.5% (3008572 kB) free!
[1559246738] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559246743] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (178216 kB) free!
[1559246880] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 370 processes
[1559247338] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559247343] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.2% (3111096 kB) free!
[1559247403] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4311840 kB) free!
[1559247458] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1559247578] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1559248003] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.9% (6838240 kB) free!
[1559248179] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559248299] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559248419] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1559248543] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (175044 kB) free!
[1559248680] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 383 processes
[1559248800] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 391 processes
[1559248920] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 392 processes
[1559249019] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559249139] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1559249143] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.6% (2039528 kB) free!
[1559249585] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.6% (2526344 kB) free!
[1559249622] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1559249739] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559249803] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3761604 kB) free!
[1559249859] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559249979] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1559250120] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 394 processes
[1559250185] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1559250185] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (150120 kB) free!
[1559250240] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 379 processes
[1559250343] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (752608 kB) free!
[1559250360] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 378 processes
[1559250403] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.0% (6413848 kB) free!
[1559250580] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559250700] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1559250943] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.5% (2672692 kB) free!
[1559250960] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 390 processes
[1559251080] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 389 processes
[1559251300] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559251420] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559251540] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1559251680] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 366 processes
[1559251800] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1559251921] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559252140] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559252203] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (3089860 kB) free!
[1559252260] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1559252743] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (459136 kB) free!
[1559252762] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559252803] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.1% (5966656 kB) free!
[1559252860] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559252981] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559253101] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559253120] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 368 processes
[1559253240] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 367 processes
[1559253701] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559253821] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559253840] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559253941] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1559253960] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1559254260] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 338.81, 225.69, 100.68
[1559254380] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 313.18, 258.04, 127.93
[1559254500] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 42.65, 172.83, 112.48
[1559254541] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559254603] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3726456 kB) free!
[1559254661] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559254782] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1559255160] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 381 processes
[1559255203] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6040156 kB) free!
[1559255280] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 381 processes
[1559255381] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559255400] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 369 processes
[1559255501] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1559256101] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559256221] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1559256300] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.33, 0.76, 16.44
[1559256343] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.8% (1753732 kB) free!
[1559256821] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559256941] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1559257200] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1559257500] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.27, 0.27, 4.74
[1559257541] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559257603] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (3791984 kB) free!
[1559257662] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559257782] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1559257782] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1559258022] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 75.43, 19.13, 6.58
[1559258142] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 157.55, 69.15, 26.43
[1559258143] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (170836 kB) free!
[1559258203] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.2% (6517164 kB) free!
[1559258262] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 168.45, 102.58, 43.86
[1559258382] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559258401] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1559258502] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1559258521] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1559258743] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.3% (3778916 kB) free.
[1559259102] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559259121] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 372 processes
[1559259222] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1559259241] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 371 processes
[1559259342] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1559259462] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.25, 3.90, 19.81
[1559259942] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559260003] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3847168 kB) free!
[1559260063] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1559260183] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1559260500] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 49.50, 58.13, 30.15
[1559260603] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.1% (6459584 kB) free!
[1559260620] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 6.98, 39.00, 26.53
[1559260740] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1559260740] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.33, 26.24, 23.37
[1559260783] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559260903] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1559261041] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 374 processes
[1559261161] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 382 processes
[1559261262] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.35, 0.36, 3.10
[1559261281] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 382 processes
[1559261340] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.38, 3.80, 12.38
[1559261503] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559261585] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559261623] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559261743] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1559261881] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 370 processes
[1559262001] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 376 processes
[1559262121] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 367 processes
[1559262343] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559262403] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3402048 kB) free!
[1559262464] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1559262540] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.18, 0.38, 3.61
[1559262943] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.8% (2396608 kB) free!
[1559263063] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559263063] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 31.6% (5123656 kB) free.
[1559263183] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559263303] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559263921] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 367 processes
[1559264203] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.3% (9530172 kB) free!
[1559264522] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1559264642] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1559264667] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 147.69, 128.53, 71.69
[1559264674] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5263 MB (10% inode=97%):
[1559264787] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 22.69, 88.11, 63.86
[1559264907] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 4.35, 59.44, 56.29
[1559265242] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1559265274] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 7001 MB (13% inode=98%):
[1559265362] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 366 processes
[1559265482] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 365 processes
[1559266107] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.78, 2.74, 16.64
[1559267503] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1559267882] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 397 processes
[1559267907] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.39, 1.55, 3.71
[1559268103] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559268223] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559268343] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559268482] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 379 processes
[1559268602] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 376 processes
[1559268722] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 374 processes
[1559269003] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4326012 kB) free!
[1559269923] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1559270263] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559270523] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 388 processes
[1559270643] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 393 processes
[1559270763] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 388 processes
[1559270803] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.2% (9479216 kB) free!
[1559271343] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1559271343] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559271943] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559271963] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1559272063] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1559272183] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1559272563] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 398 processes
[1559272683] SERVICE ALERT: redis-2;CPU_Procs;OK;HARD;1;CPU OK: 396 processes
[1559273283] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 407 processes
[1559273340] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1559273403] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 408 processes
[1559273523] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 421 processes
[1559275603] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4839608 kB) free!
[1559276381] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (168668 kB) free!
[1559276383] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1559276501] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (180208 kB) free!
[1559276524] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (169312 kB) free!
[1559276621] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (187744 kB) free!
[1559276633] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 370 processes
[1559276644] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (171840 kB) free!
[1559276753] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 367 processes
[1559276764] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166648 kB) free!
[1559276873] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 367 processes
[1559276983] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 352 processes
[1559277103] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559277223] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1559278073] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1559278140] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 356.07, 111.09, 39.19
[1559278260] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 479.56, 239.35, 94.91
[1559278380] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 496.24, 325.21, 143.88
[1559278565] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.8% (1760488 kB) free!
[1559278990] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 18.2% (4464556 kB) free!
[1559279050] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4466876 kB) free!
[1559279110] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 18.2% (4467464 kB) free!
[1559279165] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (967072 kB) free!
[1559279170] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 18.2% (4464476 kB) free!
[1559279203] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.5% (9119072 kB) free!
[1559280221] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 59.4% (9647684 kB) free.
[1559280429] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 144375 MB (64% inode=99%):
[1559280780] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.94, 0.93, 15.55
[1559281388] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.69, 5.86, 5.27
[1559281508] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.84, 6.68, 5.65
[1559281628] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.78, 6.29, 5.63
[1559281663] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (172500 kB) free!
[1559281783] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.6% (3990796 kB) free.
[1559281980] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.16, 0.36, 4.65
[1559282383] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.6% (2849476 kB) free!
[1559282436] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 842 processes
[1559282503] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 16.3% (2644472 kB) free!
[1559282556] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 842 processes
[1559282586] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (1720980 kB) free!
[1559282623] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.2% (2634716 kB) free!
[1559282623] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559282677] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 323 processes
[1559283186] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (183216 kB) free!
[1559283223] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (1510920 kB) free!
[1559283224] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1559283344] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1559283464] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559283596] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 16.5% (1212144 kB) free!
[1559283716] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 13.4% (986716 kB) free!
[1559283823] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.0% (2764976 kB) free!
[1559283836] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 10.8% (793164 kB) free!
[1559284323] SERVICE FLAPPING ALERT: redis-2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1559284421] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 15.3% (2484140 kB) free!
[1559284423] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1559284423] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (158152 kB) free!
[1559284437] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.7% (6821196 kB) free.
[1559284541] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 13.8% (2235804 kB) free!
[1559284603] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4719496 kB) free!
[1559284662] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 12.2% (1988332 kB) free!
[1559285262] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1559285262] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (779796 kB) free!
[1559285623] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.9% (3062476 kB) free!
[1559286824] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (780048 kB) free!
[1559287424] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.4% (3315044 kB) free.
[1559287562] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.0% (4435516 kB) free!
[1559287675] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179692 kB) free!
[1559287682] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.1% (3980256 kB) free!
[1559287795] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185592 kB) free!
[1559287802] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.9% (3914996 kB) free!
[1559287831] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186340 kB) free!
[1559287915] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172780 kB) free!
[1559287951] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173548 kB) free!
[1559288024] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (177700 kB) free!
[1559288071] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192000 kB) free!
[1559288144] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.2% (2467956 kB) free!
[1559288264] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.5% (3005724 kB) free!
[1559288402] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2265892 kB) free!
[1559288864] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1559288864] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (158392 kB) free!
[1559289464] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.9% (3225964 kB) free!
[1559290446] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1559290566] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1559290603] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 18.2% (8970680 kB) free!
[1559290664] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559290664] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (154568 kB) free!
[1559290686] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1559290784] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559290904] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1559292238] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (710260 kB) free!
[1559292358] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (463616 kB) free!
[1559292475] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 45.9% (22625240 kB) free.
[1559292478] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (194412 kB) free!
[1559292513] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.4% (23056716 kB) free.
[1559292562] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.5% (22433828 kB) free.
[1559292614] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 28.1% (4554216 kB) free.
[1559292671] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (1645332 kB) free!
[1559292676] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (986160 kB) free!
[1559292753] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1559292753] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.0% (6843252 kB) free.
[1559292797] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.5% (14873704 kB) free.
[1559292847] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.4% (15023580 kB) free.
[1559292893] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.0% (22806388 kB) free.
[1559293220] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (160580 kB) free!
[1559293340] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (1550328 kB) free!
[1559293447] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 60.1% (14729704 kB) free.
[1559293461] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.9% (3549532 kB) free.
[1559293707] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 44.07, 49.27, 23.15
[1559293827] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 7.42, 33.57, 20.57
[1559293947] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 3.17, 23.27, 18.37
[1559294060] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (148580 kB) free!
[1559294180] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (1204584 kB) free!
[1559294300] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.2% (1813544 kB) free!
[1559294839] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 601 processes
[1559294899] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 613 processes
[1559294900] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (181384 kB) free!
[1559295747] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.67, 2.47, 4.58
[1559297373] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (4444460 kB) free!
[1559297493] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.4% (4266960 kB) free!
[1559297614] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.7% (4100112 kB) free!
[1559297900] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.5% (3159448 kB) free!
[1559298201] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1559298801] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 14.3% (2325276 kB) free!
[1559298815] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164536 kB) free!
[1559298905] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165036 kB) free!
[1559298921] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 12.6% (2051536 kB) free!
[1559299025] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168660 kB) free!
[1559299042] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 10.7% (1740264 kB) free!
[1559299145] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181948 kB) free!
[1559299210] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 16.1% (2624228 kB) free!
[1559299330] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 15.3% (2486996 kB) free!
[1559299450] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 14.6% (2372176 kB) free!
[1559299642] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (388376 kB) free!
[1559299948] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 492.94, 278.51, 116.38
[1559300068] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 500.76, 352.43, 163.13
[1559300188] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 502.06, 402.04, 204.26
[1559300262] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 4.29, 33.65, 21.58
[1559300301] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169120 kB) free!
[1559300382] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.93, 22.61, 18.99
[1559300502] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.41, 15.25, 16.73
[1559300651] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1152108 kB) free!
[1559300901] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.5% (1705960 kB) free!
[1559300924] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (4543976 kB) free!
[1559301044] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (4462980 kB) free!
[1559301104] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559301164] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.7% (4380016 kB) free!
[1559301481] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1559301501] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.7% (3520344 kB) free.
[1559301702] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.29, 0.59, 4.85
[1559301764] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (622276 kB) free!
[1559301884] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (439104 kB) free!
[1559302004] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (244944 kB) free!
[1559302101] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (157900 kB) free!
[1559302221] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.5% (3981384 kB) free.
[1559302821] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (1034344 kB) free!
[1559302941] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (163380 kB) free!
[1559303061] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.2% (2465268 kB) free!
[1559303188] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1559303188] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.02, 2.41, 13.00
[1559303228] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 0.83, 2.49, 4.69
[1559303568] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 95.0% (23284884 kB) free.
[1559303607] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (3927584 kB) free.
[1559303661] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (160672 kB) free!
[1559303663] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.5% (6803744 kB) free.
[1559303700] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.6% (3340244 kB) free.
[1559303759] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1559303759] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.9% (14941204 kB) free.
[1559303802] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 89.7% (14576016 kB) free.
[1559303828] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.83, 5.70, 5.40
[1559303849] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.5% (22661192 kB) free.
[1559303948] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.98, 5.28, 5.28
[1559303991] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1559304068] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.02, 4.96, 5.17
[1559304111] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1559304213] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (2154380 kB) free!
[1559304231] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1559304301] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (160076 kB) free!
[1559304333] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (176392 kB) free!
[1559304421] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.3% (2483216 kB) free!
[1559304453] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (160856 kB) free!
[1559304482] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7709828 kB) free!
[1559304541] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166468 kB) free!
[1559304602] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.2% (6017312 kB) free!
[1559304668] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.27, 3.10, 4.19
[1559304704] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1559304722] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 29.5% (14531916 kB) free.
[1559304764] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1671196 kB) free!
[1559304780] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 496.15, 348.42, 160.08
[1559304900] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 168.07, 300.51, 167.37
[1559305020] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 22.83, 201.15, 147.11
[1559305364] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.6% (4107824 kB) free!
[1559305588] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.07, 2.52, 4.43
[1559305741] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.7% (1728688 kB) free!
[1559305868] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.50, 5.92, 5.28
[1559305989] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.40, 5.78, 5.31
[1559306109] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1559306109] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.22, 5.60, 5.31
[1559306254] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (2167324 kB) free!
[1559306341] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169056 kB) free!
[1559306854] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1559306854] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161892 kB) free!
[1559306941] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.5% (2349376 kB) free!
[1559307164] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1559307164] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191076 kB) free!
[1559307173] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.9% (3896576 kB) free!
[1559307294] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170052 kB) free!
[1559307309] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.73, 4.12, 4.56
[1559307414] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172276 kB) free!
[1559307421] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.77, 1.91, 12.28
[1559307456] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176280 kB) free!
[1559307541] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (157068 kB) free!
[1559307576] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180080 kB) free!
[1559307696] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191000 kB) free!
[1559308141] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.5% (2683196 kB) free!
[1559308304] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1559308322] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.2% (5006504 kB) free!
[1559308424] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1559308442] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.7% (13661092 kB) free.
[1559308544] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1559308621] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.14, 0.99, 4.71
[1559308741] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (1329456 kB) free!
[1559309031] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1559309254] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.8% (2245652 kB) free!
[1559309854] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174116 kB) free!
[1559310310] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.69, 5.32, 5.19
[1559310430] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.31, 5.49, 5.28
[1559310541] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.2% (3608380 kB) free.
[1559310550] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 6.18, 5.73, 5.39
[1559310965] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1559311054] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (1954948 kB) free!
[1559311565] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.9% (2901804 kB) free!
[1559311654] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165808 kB) free!
[1559311685] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.9% (2744672 kB) free!
[1559311741] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (168900 kB) free!
[1559311805] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.9% (2585196 kB) free!
[1559311861] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.8% (4192200 kB) free.
[1559312350] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.10, 4.13, 4.95
[1559312461] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (159792 kB) free!
[1559312581] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (161424 kB) free!
[1559312672] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 13.6% (1003592 kB) free!
[1559312701] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (1466124 kB) free!
[1559312792] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 11.6% (851444 kB) free!
[1559312900] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 592 processes
[1559312912] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (662452 kB) free!
[1559312960] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 591 processes
[1559313005] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (1006520 kB) free!
[1559313454] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.0% (2756944 kB) free!
[1559313513] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 93.0% (6838148 kB) free.
[1559314054] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (191552 kB) free!
[1559314654] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (2169704 kB) free!
[1559315101] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.0% (2917196 kB) free!
[1559315188] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1559315254] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (1388184 kB) free!
[1559315701] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155444 kB) free!
[1559316901] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.8% (3053836 kB) free!
[1559317021] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 43.76, 82.97, 46.65
[1559317141] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 6.45, 55.68, 41.05
[1559317261] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.10, 37.34, 36.10
[1559317364] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.1% (2746916 kB) free!
[1559317501] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (789884 kB) free!
[1559317861] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1559317861] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.30, 5.30, 19.07
[1559318701] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.4% (2817096 kB) free!
[1559319301] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (842344 kB) free!
[1559319842] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (4408496 kB) free!
[1559319962] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12337352 kB) free.
[1559320261] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.17, 1.29, 3.31
[1559320364] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2461152 kB) free!
[1559320501] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.7% (2380152 kB) free!
[1559321212] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3199820 kB) free!
[1559321332] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.6% (3179884 kB) free!
[1559321453] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.4% (3159980 kB) free!
[1559321701] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (161500 kB) free!
[1559322096] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 12.0% (2930816 kB) free!
[1559322102] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 36.96, 189.43, 125.49
[1559322222] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 5.56, 126.95, 110.35
[1559322302] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.3% (1989928 kB) free!
[1559322342] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.90, 85.02, 96.99
[1559322550] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1559322902] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166436 kB) free!
[1559324142] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.45, 0.61, 14.25
[1559324702] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.5% (1873168 kB) free!
[1559324913] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559325302] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165956 kB) free!
[1559325342] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.15, 0.24, 4.11
[1559326502] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 14.6% (2363276 kB) free!
[1559326654] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1559326896] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2457920 kB) free!