[1562524224] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1562524318] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12728640 kB) free.
[1562524344] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562524918] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8192236 kB) free!
[1562524966] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (9989132 kB) free!
[1562525038] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12730520 kB) free.
[1562525086] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.4% (16023628 kB) free.
[1562525544] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562525638] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8306404 kB) free!
[1562525664] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562525758] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12720936 kB) free.
[1562525785] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1562525845] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.2% (39546608 kB) free.
[1562525855] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 91.8% (14922940 kB) free.
[1562526358] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8343476 kB) free!
[1562526384] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562526478] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12728204 kB) free.
[1562526504] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562527011] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 86.94, 25.19, 8.85
[1562527078] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8250316 kB) free!
[1562527104] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562527131] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 130.53, 62.30, 24.43
[1562527198] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12729124 kB) free.
[1562527224] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1562527251] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 136.32, 87.09, 38.10
[1562527486] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (11461192 kB) free!
[1562527606] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (17305744 kB) free.
[1562527799] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8255296 kB) free!
[1562527824] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562527899] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 9 users currently logged in
[1562527919] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12728220 kB) free.
[1562527944] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562528451] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.29, 4.05, 19.24
[1562528519] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8230752 kB) free!
[1562528544] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562528639] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12729864 kB) free.
[1562528665] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562529239] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8263112 kB) free!
[1562529264] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562529359] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12720612 kB) free.
[1562529384] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1562529406] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (10491844 kB) free!
[1562529526] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (16252760 kB) free.
[1562529959] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8235076 kB) free!
[1562529984] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562530079] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12728340 kB) free.
[1562530104] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562530251] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.21, 0.24, 2.95
[1562530679] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8304500 kB) free!
[1562530704] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562530799] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12727672 kB) free.
[1562530824] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562531399] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8345600 kB) free!
[1562531424] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562531440] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (1328932 kB) free!
[1562531519] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12723748 kB) free.
[1562531545] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562531560] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (146908 kB) free!
[1562531680] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (176204 kB) free!
[1562532119] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.0% (8373484 kB) free!
[1562532144] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562532240] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12727560 kB) free.
[1562532264] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562532839] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8456788 kB) free!
[1562532864] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562532947] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 41.84, 27.71, 13.73
[1562532959] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12717040 kB) free.
[1562532984] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1562533055] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1562533067] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 41.52, 32.25, 17.09
[1562533187] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 35.80, 34.15, 19.68
[1562533559] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8283292 kB) free!
[1562533584] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562533679] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12724840 kB) free.
[1562533704] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562533787] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.49, 6.83, 11.54
[1562534279] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8210176 kB) free!
[1562534304] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562534327] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (11276368 kB) free!
[1562534399] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12730164 kB) free.
[1562534424] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562534447] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (17001800 kB) free.
[1562534999] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8152068 kB) free!
[1562535119] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12727596 kB) free.
[1562535587] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.32, 2.99, 4.12
[1562535624] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1562535719] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8150708 kB) free!
[1562535744] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562535839] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12727664 kB) free.
[1562536344] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562536440] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7921520 kB) free!
[1562536464] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562536560] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12708248 kB) free.
[1562537064] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562537160] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.2% (7958968 kB) free!
[1562537184] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562537280] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12716116 kB) free.
[1562537784] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562537815] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1562537880] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8064056 kB) free!
[1562537905] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562537935] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1562538000] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12716744 kB) free.
[1562538504] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562538600] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8104412 kB) free!
[1562538624] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1562538651] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.23, 12.65, 14.24
[1562538720] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12715740 kB) free.
[1562538771] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.26, 8.56, 12.55
[1562538891] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.07, 5.75, 11.04
[1562539224] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562539320] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8129836 kB) free!
[1562539344] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562539440] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12716196 kB) free.
[1562539944] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562540040] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8118024 kB) free!
[1562540064] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562540091] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1562540091] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.19, 0.29, 3.20
[1562540160] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12706788 kB) free.
[1562540664] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562540761] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8175616 kB) free!
[1562540784] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562540881] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12715424 kB) free.
[1562540905] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562541481] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8189976 kB) free!
[1562541504] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562541601] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12717904 kB) free.
[1562541624] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562542201] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8088972 kB) free!
[1562542321] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12714480 kB) free.
[1562542824] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562542921] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8040408 kB) free!
[1562542944] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562543041] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.8% (12714700 kB) free.
[1562543448] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (11249616 kB) free!
[1562543544] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562543568] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (17206680 kB) free.
[1562543664] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562544168] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (12079868 kB) free!
[1562544176] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 73.5% (11947660 kB) free.
[1562544220] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 59.1% (38899848 kB) free.
[1562544265] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562544385] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562544505] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562545441] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8327444 kB) free!
[1562545562] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12463832 kB) free.
[1562546761] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9314948 kB) free!
[1562546881] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.4% (7584268 kB) free!
[1562546905] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562547001] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11882336 kB) free.
[1562547075] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4704936 kB) free!
[1562547195] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (4665688 kB) free!
[1562547315] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (4661188 kB) free!
[1562547505] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562547625] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1562547726] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 461 processes
[1562547746] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562547786] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 460 processes
[1562549546] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1562549546] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562550001] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7544276 kB) free!
[1562550121] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12491224 kB) free.
[1562550891] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562551321] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9130700 kB) free!
[1562551346] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562551442] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12492060 kB) free.
[1562551466] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562551735] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1562551787] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 15005.64, 12100.64, 6157.14
[1562551855] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1562551907] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15009.06, 13063.11, 7230.48
[1562552027] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15011.12, 13707.55, 8173.82
[1562552067] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562552099] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5454 MB (10% inode=97%):
[1562552187] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562552307] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562552699] SERVICE ALERT: pservercluster2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 7075 MB (13% inode=98%):
[1562552715] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 22.6% (5534220 kB) free.
[1562554807] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4934652 kB) free!
[1562554927] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (4909200 kB) free!
[1562555041] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8808188 kB) free!
[1562555047] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.8% (4887208 kB) free!
[1562555161] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12478816 kB) free.
[1562555307] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562555761] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8815820 kB) free!
[1562555881] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12487452 kB) free.
[1562556287] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4819512 kB) free!
[1562556407] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (4798472 kB) free!
[1562556481] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.8% (8784132 kB) free!
[1562556507] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562556527] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.5% (4779564 kB) free!
[1562556602] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12474172 kB) free.
[1562556627] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1562557201] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8742264 kB) free!
[1562557227] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562557321] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12482016 kB) free.
[1562557347] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562557655] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.5% (2028564 kB) free!
[1562557775] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.6% (1726068 kB) free!
[1562557895] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 10.1% (1639052 kB) free!
[1562557921] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.0% (8851324 kB) free!
[1562557947] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562558027] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.73, 2.85, 18.77
[1562558041] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12483716 kB) free.
[1562558067] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562558115] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4812080 kB) free!
[1562558235] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (4788392 kB) free!
[1562558355] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (4756444 kB) free!
[1562558641] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.2% (8947140 kB) free!
[1562558667] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562558761] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12476676 kB) free.
[1562558787] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562558796] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.9% (2906188 kB) free!
[1562558916] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.1% (2775460 kB) free!
[1562559036] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.4% (2657548 kB) free!
[1562559361] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8667608 kB) free!
[1562559388] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562559481] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12479616 kB) free.
[1562559508] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562560081] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8203620 kB) free!
[1562560108] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562560201] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12472004 kB) free.
[1562560228] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562560427] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.58, 3.08, 4.13
[1562560447] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2144076 kB) free!
[1562560732] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.8% (21660348 kB) free.
[1562560801] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7467648 kB) free!
[1562560836] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (953180 kB) free!
[1562560895] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (521272 kB) free!
[1562560921] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11767720 kB) free.
[1562561955] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1893564 kB) free!
[1562562527] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183716 kB) free!
[1562562555] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1562562555] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 14.7% (3593912 kB) free!
[1562562628] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562562749] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562563046] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (245208 kB) free!
[1562563127] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 22.2% (5440876 kB) free.
[1562563166] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (252064 kB) free!
[1562563286] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (246468 kB) free!
[1562563321] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.2% (8489080 kB) free!
[1562563348] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562563441] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12463816 kB) free.
[1562563468] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562564041] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8254368 kB) free!
[1562564068] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562564161] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.3% (12453064 kB) free.
[1562564188] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562564355] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 23.0% (5633892 kB) free.
[1562564761] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8033628 kB) free!
[1562564788] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562564881] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12334552 kB) free.
[1562564908] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562565028] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562565186] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (157224 kB) free!
[1562565307] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (175056 kB) free!
[1562565427] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (176252 kB) free!
[1562567327] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.5% (4524012 kB) free!
[1562567355] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4743188 kB) free!
[1562567437] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 58.5% (9502268 kB) free.
[1562567447] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.2% (4453256 kB) free!
[1562567475] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (4628212 kB) free!
[1562567567] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1562567567] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.1% (4192312 kB) free!
[1562567595] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (4586924 kB) free!
[1562568628] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1562568655] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1562568775] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1562568795] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 21.2% (5197664 kB) free.
[1562569081] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9073872 kB) free!
[1562569198] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.01, 10.30, 7.29
[1562569201] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7843392 kB) free!
[1562569228] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1562569318] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 12.67, 11.46, 8.11
[1562569321] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12279852 kB) free.
[1562569348] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562569395] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4575432 kB) free!
[1562569438] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 15.09, 12.69, 8.95
[1562569469] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562569515] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.1% (4436880 kB) free!
[1562569635] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.6% (4316096 kB) free!
[1562569967] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1764308 kB) free!
[1562570437] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.4% (2825712 kB) free!
[1562570557] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.1% (2451096 kB) free!
[1562570677] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 12.9% (2091036 kB) free!
[1562571075] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 367 processes
[1562571121] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9644940 kB) free!
[1562571195] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 363 processes
[1562571241] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8253520 kB) free!
[1562571277] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1562571277] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (495280 kB) free!
[1562571316] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 367 processes
[1562571362] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12276884 kB) free.
[1562571436] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2407432 kB) free!
[1562571536] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191984 kB) free!
[1562571656] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179680 kB) free!
[1562571699] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1562571776] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184568 kB) free!
[1562571916] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 367 processes
[1562572561] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9476724 kB) free!
[1562572681] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.2% (7990112 kB) free!
[1562572794] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 19.5% (1435088 kB) free!
[1562572801] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.5% (12056548 kB) free.
[1562572914] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 14.4% (1058228 kB) free!
[1562573034] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (691568 kB) free!
[1562575201] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8416384 kB) free!
[1562575321] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12369132 kB) free.
[1562575685] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.04, 5.63, 5.30
[1562575805] SERVICE ALERT: smscapp1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.90, 5.87, 5.43
[1562575925] SERVICE ALERT: smscapp1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.35, 5.30, 5.27
[1562576025] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (13027288 kB) free!
[1562576145] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (20398756 kB) free.
[1562576525] SERVICE ALERT: smscapp1;CPU Load;OK;HARD;1;OK - load average: 2.91, 3.79, 4.54
[1562576575] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1562576695] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1562576745] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (12587556 kB) free!
[1562576865] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (19564440 kB) free.
[1562577122] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.3% (9000720 kB) free!
[1562577242] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12356600 kB) free.
[1562577269] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562577316] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1562577436] SERVICE ALERT: webservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1562577557] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1562577695] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 19.2% (3117448 kB) free!
[1562577699] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 10 users currently logged in
[1562577969] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.37, 6.08, 5.53
[1562578089] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.37, 5.46, 5.37
[1562578209] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.93, 4.91, 5.18
[1562578295] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (152596 kB) free!
[1562578434] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.0% (6769040 kB) free.
[1562578727] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 14906.14, 9336.74, 4033.86
[1562578809] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.04, 1.69, 3.34
[1562578847] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15008.97, 11218.84, 5366.53
[1562578899] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 12 users currently logged in
[1562578967] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 15021.95, 12477.87, 6537.49
[1562579567] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1562580242] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7549120 kB) free!
[1562580291] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 117.69, 68.05, 28.66
[1562580362] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12341656 kB) free.
[1562580411] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 119.39, 85.10, 39.69
[1562580427] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.17, 7.35, 5.33
[1562580531] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 93.30, 90.75, 47.49
[1562580547] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.39, 7.39, 5.59
[1562580667] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.41, 7.00, 5.66
[1562580767] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.4% (2552992 kB) free!
[1562580838] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.49, 3.76, 4.60
[1562580869] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562580962] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7814648 kB) free!
[1562580989] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562581082] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12346852 kB) free.
[1562581109] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562581367] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2327956 kB) free!
[1562581636] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562581682] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.6% (8695660 kB) free!
[1562581709] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562581731] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.30, 2.09, 13.39
[1562581757] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1562581802] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12336404 kB) free.
[1562581829] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1562581867] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 2.53, 2.98, 4.04
[1562582038] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.40, 6.58, 5.44
[1562582159] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.59, 6.18, 5.43
[1562582279] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.20, 6.38, 5.58
[1562582403] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8940680 kB) free!
[1562582430] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562582523] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.4% (7578192 kB) free!
[1562582550] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562582643] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12341400 kB) free.
[1562582677] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562582865] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (12849088 kB) free!
[1562582931] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.68, 0.55, 4.01
[1562582985] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (21238264 kB) free.
[1562583585] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (12635580 kB) free!
[1562583630] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.2% (2961764 kB) free!
[1562583705] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (20984752 kB) free.
[1562583750] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.6% (2864092 kB) free!
[1562583750] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1562583870] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.2% (2789572 kB) free!
[1562584079] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 5.16, 4.67, 4.93
[1562584305] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (12293244 kB) free!
[1562584425] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.7% (20862204 kB) free.
[1562584950] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562584967] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.78, 3.01, 19.71
[1562585025] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (12053344 kB) free!
[1562585070] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562585145] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.3% (20559660 kB) free.
[1562585227] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.5% (1705532 kB) free!
[1562585279] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.95, 7.83, 6.05
[1562585399] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 13.14, 8.99, 6.65
[1562585519] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1562585519] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 12.68, 10.03, 7.31
[1562585519] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.1% (22492360 kB) free.
[1562585570] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 88.3% (21646828 kB) free.
[1562585628] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.44, 5.38, 5.12
[1562585632] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.5% (39172488 kB) free.
[1562585634] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 17.1% (1256876 kB) free!
[1562585744] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.3% (14031812 kB) free.
[1562585745] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (11247652 kB) free!
[1562585748] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.21, 5.30, 5.12
[1562585754] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 13.9% (1022136 kB) free!
[1562585799] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.8% (21026952 kB) free.
[1562585827] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (1569860 kB) free!
[1562585865] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (19959016 kB) free.
[1562585868] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.94, 5.16, 5.10
[1562585873] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.9% (14279884 kB) free.
[1562585874] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 11.5% (844080 kB) free!
[1562585924] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 69.5% (11290560 kB) free.
[1562586125] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.1% (2498512 kB) free!
[1562586243] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8714296 kB) free!
[1562586245] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177600 kB) free!
[1562586270] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1301944 kB) free!
[1562586271] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562586363] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12330016 kB) free.
[1562586365] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197536 kB) free!
[1562586391] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562586465] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (10877324 kB) free!
[1562586474] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (119300 kB) free!
[1562586479] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (176776 kB) free!
[1562586557] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1562586585] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (19338300 kB) free.
[1562586599] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (180712 kB) free!
[1562586719] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1562586719] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (162708 kB) free!
[1562586768] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.51, 1.81, 4.79
[1562586839] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9330500 kB) free!
[1562586959] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.0% (3945864 kB) free!
[1562587079] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (299644 kB) free!
[1562587185] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (10641152 kB) free!
[1562587305] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (19183552 kB) free.
[1562587906] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (10930416 kB) free!
[1562588026] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (19167540 kB) free.
[1562588095] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1562588164] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7432428 kB) free!
[1562588215] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1562588284] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12335280 kB) free.
[1562588519] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.95, 4.48, 4.86
[1562588626] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (10987556 kB) free!
[1562588746] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (19194416 kB) free.
[1562588791] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562588911] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562589120] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.74, 6.32, 5.59
[1562589240] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.55, 6.08, 5.60
[1562589271] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 83.3% (13538608 kB) free.
[1562589346] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (10082228 kB) free!
[1562589360] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.42, 5.36, 5.39
[1562589466] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (18692132 kB) free.
[1562589474] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1562589474] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.7% (6815640 kB) free.
[1562589511] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562589632] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562589752] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1562589960] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.92, 3.62, 4.46
[1562590066] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (10145868 kB) free!
[1562590084] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9783700 kB) free!
[1562590186] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (18059152 kB) free.
[1562590204] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.1% (8403756 kB) free!
[1562590324] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12316888 kB) free.
[1562590560] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.35, 7.94, 6.29
[1562590680] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.99, 8.26, 6.62
[1562590800] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.32, 6.80, 6.28
[1562591524] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9307416 kB) free!
[1562591644] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7668408 kB) free!
[1562591649] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.6% (21856144 kB) free.
[1562591712] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.1% (38510076 kB) free.
[1562591764] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11872240 kB) free.
[1562591771] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 84.1% (13664248 kB) free.
[1562591957] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1562592255] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171816 kB) free!
[1562592318] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.6% (8194188 kB) free!
[1562592344] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.8% (2734540 kB) free!
[1562592375] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188972 kB) free!
[1562592378] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (168896 kB) free!
[1562592438] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (3432912 kB) free!
[1562592465] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 15.1% (2457524 kB) free!
[1562592495] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1562592495] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196432 kB) free!
[1562592499] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (174316 kB) free!
[1562592558] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1562592558] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (221948 kB) free!
[1562592585] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 14.0% (2281812 kB) free!
[1562592619] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (170120 kB) free!
[1562592752] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562592776] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168188 kB) free!
[1562592896] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170492 kB) free!
[1562593006] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173760 kB) free!
[1562593016] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189480 kB) free!
[1562593126] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190588 kB) free!
[1562593246] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165180 kB) free!
[1562593352] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562593473] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562593564] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8430136 kB) free!
[1562593593] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562593684] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12289028 kB) free.
[1562593785] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (788204 kB) free!
[1562594884] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9429348 kB) free!
[1562595004] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.0% (7894532 kB) free!
[1562595124] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12009712 kB) free.
[1562595240] SERVICE ALERT: sig2-vm1;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 13180 MB (10% inode=96%):
[1562595993] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1562595993] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562596324] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9092128 kB) free!
[1562596444] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.0% (7408688 kB) free!
[1562596564] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.7% (11698160 kB) free.
[1562596593] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562596713] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562596834] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562598064] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.2% (21016760 kB) free.
[1562598124] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 82.2% (20150536 kB) free.
[1562598184] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.7% (39766748 kB) free.
[1562598242] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.8% (14109008 kB) free.
[1562598310] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.6% (21719880 kB) free.
[1562598372] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.7% (14416116 kB) free.
[1562598474] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1562598634] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562598664] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.5% (4072924 kB) free!
[1562598784] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (1588848 kB) free!
[1562598904] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170948 kB) free!
[1562598964] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9320568 kB) free!
[1562598979] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (152304 kB) free!
[1562599084] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8200108 kB) free!
[1562599099] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (157344 kB) free!
[1562599204] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12207804 kB) free.
[1562599219] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165388 kB) free!
[1562599668] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.71, 4.81, 4.95
[1562599834] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562599954] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562600074] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562601586] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (11626184 kB) free!
[1562601600] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.74, 4.19, 4.56
[1562601706] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (19277508 kB) free.
[1562602157] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1562602209] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 82.8% (20425824 kB) free.
[1562602281] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 89.1% (14475004 kB) free.
[1562602306] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (12720524 kB) free!
[1562602426] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (19550764 kB) free.
[1562602757] SERVICE ALERT: webservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 13 users currently logged in
[1562602889] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (162008 kB) free!
[1562603009] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (162932 kB) free!
[1562603129] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (180772 kB) free!
[1562603268] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.33, 5.16, 5.02
[1562603388] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.50, 5.18, 5.05
[1562603404] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8277320 kB) free!
[1562603409] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.27, 7.41, 5.67
[1562603508] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.07, 5.12, 5.04
[1562603524] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12208156 kB) free.
[1562603529] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.46, 5.13, 5.05
[1562603590] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1562603649] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.43, 3.53, 4.47
[1562603674] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562604015] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196716 kB) free!
[1562604108] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.55, 4.85, 4.92
[1562604124] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.5% (8140772 kB) free!
[1562604135] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198468 kB) free!
[1562604226] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (12466460 kB) free!
[1562604244] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12210612 kB) free.
[1562604255] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.0% (486728 kB) free!
[1562604346] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (19248160 kB) free.
[1562604708] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 21.94, 34.56, 18.69
[1562604790] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.1% (7935040 kB) free!
[1562604828] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.39, 24.73, 17.01
[1562604844] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8317520 kB) free!
[1562604874] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562604911] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.7% (7722004 kB) free!
[1562604948] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1562604948] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.11, 18.13, 15.53
[1562604964] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12203312 kB) free.
[1562604994] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562605031] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.6% (7708228 kB) free!
[1562605449] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.57, 6.01, 5.33
[1562605564] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.1% (7947048 kB) free!
[1562605569] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.70, 4.11, 4.71
[1562605595] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562605684] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12210800 kB) free.
[1562605715] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562605757] SERVICE ALERT: webservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 12 users currently logged in
[1562606231] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (243936 kB) free!
[1562606284] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7495352 kB) free!
[1562606330] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (151900 kB) free!
[1562606404] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12197360 kB) free.
[1562606450] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (178564 kB) free!
[1562606570] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (390268 kB) free!
[1562606915] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562607035] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562607600] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1562607605] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.7% (8741992 kB) free!
[1562607635] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562607725] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12207952 kB) free.
[1562607755] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562607946] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (12953712 kB) free!
[1562608066] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (19119588 kB) free.
[1562608325] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8518596 kB) free!
[1562608355] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562608445] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12211164 kB) free.
[1562608476] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562608666] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (11548936 kB) free!
[1562608786] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.0% (23643756 kB) free.
[1562609045] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8165960 kB) free!
[1562609075] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562609165] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12197360 kB) free.
[1562609195] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562609748] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 5.43, 4.95, 4.94
[1562609765] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.6% (7711976 kB) free!
[1562609885] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12207088 kB) free.
[1562610395] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562610485] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7345972 kB) free!
[1562610515] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes