[1555698709] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 501.89, 328.74, 145.39
[1555698829] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 503.91, 386.81, 188.90
[1555700727] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.9% (12757840 kB) free.
[1555701229] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.48, 0.57, 16.27
[1555701237] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.99, 2.00, 13.72
[1555702400] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.9% (5381244 kB) free!
[1555702429] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.29, 0.37, 4.74
[1555703000] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235736 kB) free!
[1555703037] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.66, 1.79, 3.56
[1555703727] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.6% (8190584 kB) free!
[1555703848] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.9% (5382052 kB) free!
[1555703968] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (2884500 kB) free!
[1555705531] SERVICE ALERT: db2;Mariadb status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1555705534] SERVICE ALERT: db2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1555705627] SERVICE ALERT: db2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1555705661] SERVICE ALERT: db2;Mariadb status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1555705664] SERVICE ALERT: db2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1555705735] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 2.57, 60.54, 45.63
[1555705757] SERVICE ALERT: db2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1555705768] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.9% (12749808 kB) free.
[1555705791] SERVICE ALERT: db2;Mariadb status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1555705794] SERVICE ALERT: db2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1555705855] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.97, 40.71, 40.17
[1555705874] SERVICE ALERT: db2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.53, 8.01, 5.14
[1555705877] SERVICE ALERT: db2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 44510 MB (86% inode=99%):
[1555705975] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.75, 27.43, 35.37
[1555705994] SERVICE ALERT: db2;CPU Load;OK;HARD;1;OK - load average: 2.63, 6.01, 4.75
[1555706000] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.8% (5804828 kB) free!
[1555706391] SERVICE ALERT: db2;Mariadb status;OK;HARD;1;PROCS OK: 1 process with command name 'mysqld'
[1555706394] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 55775 MB (24% inode=99%):
[1555706575] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.65, 4.17, 18.80
[1555706601] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1555706601] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (432476 kB) free!
[1555707724] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1555708168] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.9% (8325852 kB) free!
[1555708288] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.1% (5482156 kB) free!
[1555708375] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.33, 0.54, 3.18
[1555708408] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (2703864 kB) free!
[1555709008] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.9% (12787800 kB) free.
[1555710808] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9096884 kB) free!
[1555710928] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.1% (6441132 kB) free!
[1555711048] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3672720 kB) free!
[1555711648] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.1% (12879640 kB) free.
[1555711924] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 316 processes
[1555712001] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 26.4% (13028588 kB) free.
[1555712848] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9866308 kB) free!
[1555712968] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 14.2% (7001868 kB) free!
[1555713088] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4180612 kB) free!
[1555713688] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.2% (12904748 kB) free.
[1555714888] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.7% (7240068 kB) free!
[1555715008] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (4561144 kB) free!
[1555715129] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.3% (12974740 kB) free.
[1555716328] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.2% (8452060 kB) free!
[1555716448] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.4% (5603228 kB) free!
[1555716568] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.6% (13090100 kB) free.
[1555717768] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.9% (8333016 kB) free!
[1555717888] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.9% (5364068 kB) free!
[1555718008] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.8% (13210548 kB) free.
[1555718601] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1555719208] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.4% (7587344 kB) free!
[1555719229] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 503.81, 330.23, 146.11
[1555719328] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (4801320 kB) free!
[1555719349] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 506.51, 388.61, 189.83
[1555719448] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.9% (13250652 kB) free.
[1555719469] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 69.05, 260.51, 167.00
[1555719724] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 319 processes
[1555719845] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1555719965] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 324 processes
[1555720049] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.2% (8006932 kB) free!
[1555720169] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.7% (5269140 kB) free!
[1555720289] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (2557176 kB) free!
[1555720889] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.1% (11888228 kB) free.
[1555721324] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4908624 kB) free!
[1555721444] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (4903924 kB) free!
[1555721489] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.0% (8377384 kB) free!
[1555721564] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.8% (4897372 kB) free!
[1555721609] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.7% (5293960 kB) free!
[1555721729] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.2% (13412400 kB) free.
[1555721869] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.60, 0.55, 13.04
[1555722365] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 319 processes
[1555722929] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8821632 kB) free!
[1555722965] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1555723049] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.6% (5700132 kB) free!
[1555723069] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.54, 0.50, 3.92
[1555723085] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 321 processes
[1555723169] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 26.5% (13063572 kB) free.
[1555723206] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1555723206] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 326 processes
[1555724369] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.5% (8152384 kB) free!
[1555724489] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 10.8% (5344064 kB) free!
[1555724610] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 25.7% (12642008 kB) free.
[1555725594] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1555725809] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.9% (7341044 kB) free!
[1555725929] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (4448664 kB) free!
[1555726049] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.4% (12025748 kB) free.
[1555726286] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1555726406] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1555726526] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1555726649] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9696876 kB) free!
[1555726650] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 322 processes
[1555726769] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 14.1% (6930376 kB) free!
[1555726889] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4197988 kB) free!
[1555727126] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1555727406] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1555727489] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.1% (10915448 kB) free.
[1555727726] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1555727846] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1555727966] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1555728007] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1555728089] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.7% (8226948 kB) free!
[1555728127] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 320 processes
[1555728209] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.1% (5488860 kB) free!
[1555728247] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 320 processes
[1555728329] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 24.5% (12087672 kB) free.
[1555729530] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.3% (8513920 kB) free!
[1555729650] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.9% (5859324 kB) free!
[1555729770] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (3218264 kB) free!
[1555730370] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.0% (10819608 kB) free.
[1555730762] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555730882] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555730966] SERVICE FLAPPING ALERT: db2;Current Users;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1555730966] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1555730970] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 12.1% (5976952 kB) free!
[1555730994] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1555731090] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (3122860 kB) free!
[1555731210] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (392040 kB) free!
[1555731247] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 320 processes
[1555731688] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.75, 13.62, 8.46
[1555731764] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2451428 kB) free!
[1555731808] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.43, 9.21, 7.47
[1555731810] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.6% (9660548 kB) free!
[1555731837] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.19, 12.20, 10.42
[1555731848] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1555731928] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.48, 6.34, 6.63
[1555731957] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.91, 8.77, 9.37
[1555731968] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 323 processes
[1555732077] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.79, 6.48, 8.46
[1555732088] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1555732364] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.5% (2593612 kB) free!
[1555732376] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.75, 21.96, 16.16
[1555732496] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.94, 14.98, 14.31
[1555732528] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.08, 1.06, 3.60
[1555732616] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.79, 10.25, 12.66
[1555732650] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 323 processes
[1555733010] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (3361216 kB) free!
[1555733277] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.84, 1.82, 3.59
[1555733564] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2443972 kB) free!
[1555733610] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.7% (10202336 kB) free.
[1555733816] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.73, 0.90, 4.01
[1555734210] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 15.1% (7465980 kB) free!
[1555734330] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (4775116 kB) free!
[1555734451] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (2037444 kB) free!
[1555735051] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.7% (9718756 kB) free!
[1555735194] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1555735616] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 475.81, 289.53, 124.26
[1555735714] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1555735714] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.2% (22023296 kB) free.
[1555735736] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 479.90, 352.72, 167.45
[1555735776] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.8% (22497076 kB) free.
[1555735856] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 98.87, 264.30, 158.75
[1555736251] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3903320 kB) free!
[1555736289] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 322 processes
[1555736889] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555737009] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 324 processes
[1555737129] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 319 processes
[1555737451] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 19.9% (9811788 kB) free!
[1555737477] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 481.87, 352.53, 167.12
[1555737598] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 483.08, 395.76, 205.44
[1555737718] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 151.39, 324.11, 204.16
[1555738069] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 433.56, 193.40, 74.20
[1555738082] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555738189] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 99.50, 162.38, 78.65
[1555738202] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555738256] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1555738256] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.35, 0.86, 12.65
[1555738309] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 14.08, 108.90, 69.21
[1555738322] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555739251] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (2319660 kB) free!
[1555739456] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.76, 1.14, 4.15
[1555739509] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.83, 2.70, 19.51
[1555739966] SERVICE FLAPPING ALERT: db2;Current Users;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1555740118] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1555740118] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.67, 1.99, 17.23
[1555740130] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1555740451] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.0% (11325652 kB) free.
[1555740581] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.0% (3916992 kB) free!
[1555740701] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177784 kB) free!
[1555740730] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1555740821] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.4% (3147496 kB) free!
[1555740821] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196416 kB) free!
[1555740850] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 319 processes
[1555740941] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.4% (2991336 kB) free!
[1555740970] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 324 processes
[1555741061] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.5% (2841760 kB) free!
[1555741309] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.33, 0.87, 3.42
[1555741545] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 281 processes
[1555741656] SERVICE ALERT: sig2-vm1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 285 processes
[1555741665] SERVICE ALERT: sig5-vm2;CPU_Procs;OK;HARD;1;CPU OK: 276 processes
[1555741686] SERVICE ALERT: sig3-vm1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 281 processes
[1555741776] SERVICE ALERT: sig2-vm1;CPU_Procs;OK;HARD;1;CPU OK: 280 processes
[1555741806] SERVICE ALERT: sig3-vm1;CPU_Procs;OK;HARD;1;CPU OK: 276 processes
[1555741918] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.06, 2.63, 4.73
[1555742021] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 81.2% (19907304 kB) free.
[1555742251] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.7% (8717972 kB) free!
[1555742262] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 80.0% (13008084 kB) free.
[1555742371] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 11.6% (5723948 kB) free!
[1555742492] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (2877012 kB) free!
[1555743779] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555743899] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555744019] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555744124] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1555744922] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555745456] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.54, 8.07, 7.60
[1555745576] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.40, 5.53, 6.73
[1555745696] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.35, 3.83, 5.96
[1555745819] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1555745994] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1555746296] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.91, 1.54, 3.59
[1555747803] HOST ALERT: webservercluster2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1555748402] HOST ALERT: webservercluster2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.52)
[1555748715] HOST ALERT: webservercluster2;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1555749594] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1555750570] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 321 processes
[1555751170] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 319 processes
[1555751290] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1555751410] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 320 processes
[1555752092] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.3% (5579188 kB) free!
[1555752118] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1555752525] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (185008 kB) free!
[1555752645] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183676 kB) free!
[1555752765] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174944 kB) free!
[1555753062] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.3% (3138176 kB) free!
[1555753182] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.4% (2996556 kB) free!
[1555753302] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.5% (2840084 kB) free!
[1555753794] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1555753810] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 322 processes
[1555753892] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.0% (11327612 kB) free.
[1555754219] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 8 users currently logged in
[1555754339] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 8 users currently logged in
[1555754411] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 323 processes
[1555754459] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1555754503] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1243360 kB) free!
[1555754531] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1555754651] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1555754696] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.20, 44.70, 47.64
[1555754816] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.77, 30.14, 41.95
[1555754936] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.01, 20.59, 37.02
[1555755093] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (248776 kB) free!
[1555755213] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.4% (10567544 kB) free.
[1555755536] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.82, 3.37, 19.73
[1555755659] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1555755659] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1555756413] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (218848 kB) free!
[1555756533] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 21.6% (10663596 kB) free.
[1555757001] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8181428 kB) free!
[1555757051] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1555757121] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11711660 kB) free.
[1555757336] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.90, 0.86, 3.53
[1555757651] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1555757733] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (213968 kB) free!
[1555757772] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 322 processes
[1555757853] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 23.1% (11363424 kB) free.
[1555757892] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1555758453] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (234328 kB) free!
[1555758574] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 20.0% (9841916 kB) free!
[1555758694] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.5% (6181084 kB) free!
[1555759521] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8541832 kB) free!
[1555759641] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.4% (12047652 kB) free.
[1555760494] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4744308 kB) free!
[1555760841] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7722272 kB) free!
[1555760892] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 316 processes
[1555760962] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11384140 kB) free.
[1555761492] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555761561] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9267984 kB) free!
[1555761613] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 322 processes
[1555761681] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.6% (7699008 kB) free!
[1555761733] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 315 processes
[1555761801] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.0% (11349560 kB) free.
[1555762333] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1555762401] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9510660 kB) free!
[1555762453] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1555762521] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8175592 kB) free!
[1555762573] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 320 processes
[1555762641] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11753088 kB) free.
[1555762894] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 20.6% (10141060 kB) free.
[1555763128] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 45.29, 230.42, 152.02
[1555763173] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 322 processes
[1555763241] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9592244 kB) free!
[1555763248] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 6.31, 154.29, 133.62
[1555763293] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1555763326] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.8% (2637540 kB) free!
[1555763361] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.6% (8157468 kB) free!
[1555763368] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.43, 103.45, 117.50
[1555763413] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 323 processes
[1555763481] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11707820 kB) free.
[1555763494] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (243236 kB) free!
[1555763503] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 77.6% (12614920 kB) free.
[1555763614] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (232292 kB) free!
[1555763734] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (248720 kB) free!
[1555763926] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2290248 kB) free!
[1555764014] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1555764081] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9696168 kB) free!
[1555764134] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555764201] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.9% (8328804 kB) free!
[1555764254] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1555764322] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11954068 kB) free.
[1555765168] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.46, 0.74, 17.31
[1555765521] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7802020 kB) free!
[1555765641] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.1% (11363896 kB) free.
[1555766241] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.7% (9199812 kB) free!
[1555766361] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.7% (7760152 kB) free!
[1555766459] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1555766481] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.7% (11188172 kB) free.
[1555766968] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.47, 0.50, 2.91
[1555767081] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.1% (8900524 kB) free!
[1555767201] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 14.7% (7224112 kB) free!
[1555767254] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 314 processes
[1555767321] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.2% (10916404 kB) free.
[1555767536] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1555767855] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555767921] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (2821300 kB) free!
[1555767975] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 319 processes
[1555768042] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (4414608 kB) free!
[1555768095] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1555768162] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.9% (12261480 kB) free.
[1555768762] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3090244 kB) free!
[1555768882] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (1760352 kB) free!
[1555769002] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (234568 kB) free!
[1555769134] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 13.1% (6441620 kB) free!
[1555769326] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 47.6% (11668568 kB) free.
[1555769459] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555769503] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (172800 kB) free!
[1555769579] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555769602] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.8% (7286564 kB) free!
[1555769624] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (169000 kB) free!
[1555769699] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555769734] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (239756 kB) free!
[1555769744] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (180244 kB) free!
[1555769926] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190656 kB) free!
[1555770046] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193664 kB) free!
[1555770166] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1555770166] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176744 kB) free!
[1555770802] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (4298132 kB) free!
[1555770944] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 71.0% (11547360 kB) free.
[1555771402] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1555771402] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.2% (9457776 kB) free!
[1555771544] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (167456 kB) free!
[1555771664] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (152272 kB) free!
[1555771695] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1555771784] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1555771784] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (175564 kB) free!
[1555772021] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4726028 kB) free!
[1555772134] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.1% (6937140 kB) free!
[1555772141] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (4676836 kB) free!
[1555772262] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (4633628 kB) free!
[1555772295] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1555772415] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 321 processes
[1555772535] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1555772602] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 28.0% (13791072 kB) free.
[1555772734] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (218712 kB) free!
[1555772984] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 76.6% (12444648 kB) free.
[1555773802] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9815388 kB) free!
[1555773900] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555773922] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 27.3% (13463352 kB) free.
[1555773935] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 18.9% (9319388 kB) free!
[1555774535] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240604 kB) free!
[1555775535] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1555775735] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 14.1% (6960940 kB) free!
[1555776110] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 59.11, 235.57, 151.10
[1555776135] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1555776230] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 8.74, 157.95, 132.90
[1555776255] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 314 processes
[1555776335] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235788 kB) free!
[1555776350] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.92, 105.94, 116.88
[1555776375] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555777535] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 22.9% (11274572 kB) free.
[1555778135] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (3480204 kB) free!
[1555778150] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.74, 1.01, 17.41
[1555778255] SERVICE ALERT: db1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (1350052 kB) free!
[1555778375] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (220908 kB) free!
[1555779575] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 12.4% (6095028 kB) free!
[1555779950] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.50, 0.56, 3.06
[1555780176] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (244204 kB) free!
[1555780500] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555780620] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555780740] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555780794] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1555781376] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 17.5% (8607152 kB) free!
[1555781567] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1555782576] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (1996008 kB) free!
[1555783522] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1555783784] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1555784376] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 11.2% (5508384 kB) free!
[1555784850] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 330 processes