[1586370734] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20330968 kB) free!
[1586371334] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37260424 kB) free.
[1586371934] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16714788 kB) free!
[1586372054] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12924220 kB) free!
[1586372174] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (9186276 kB) free!
[1586372478] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 7.5% (1834584 kB) free!
[1586372774] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.1% (25150248 kB) free!
[1586373048] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (604896 kB) free!
[1586373169] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (236832 kB) free!
[1586373289] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.4% (2336248 kB) free!
[1586373375] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.7% (6244392 kB) free!
[1586373975] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20796012 kB) free!
[1586373979] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 83.6% (20496308 kB) free.
[1586374575] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37074408 kB) free.
[1586374992] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.93, 4.21, 4.86
[1586375089] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (1151652 kB) free!
[1586375175] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16247936 kB) free!
[1586375295] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (12456832 kB) free!
[1586375415] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (8866272 kB) free!
[1586375551] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 11.3% (14931468 kB) free!
[1586375671] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (8177920 kB) free!
[1586375791] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (6263356 kB) free!
[1586375989] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (217756 kB) free!
[1586376015] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.0% (23769428 kB) free!
[1586376109] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (219616 kB) free!
[1586376192] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.91, 5.55, 5.03
[1586376229] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (218204 kB) free!
[1586376313] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.86, 6.08, 5.29
[1586376433] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.87, 6.37, 5.49
[1586376615] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (5101540 kB) free!
[1586376671] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(118) < 92.42.104.166 [09/Apr/2020:01:35:16 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1586376791] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586376889] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (3118936 kB) free!
[1586377215] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.1% (18528364 kB) free!
[1586377815] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34174856 kB) free.
[1586378029] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 21.9% (10786608 kB) free.
[1586378415] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11436600 kB) free!
[1586378535] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7518940 kB) free!
[1586378655] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (3855096 kB) free!
[1586379255] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.5% (17775240 kB) free!
[1586379433] SERVICE ALERT: sigqscluster1;CPU Load;OK;HARD;1;OK - load average: 3.96, 4.37, 4.87
[1586379855] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36640692 kB) free.
[1586379951] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 72.8% (35900864 kB) free.
[1586380456] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12691472 kB) free!
[1586380489] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.0% (3735816 kB) free.
[1586380576] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (8655132 kB) free!
[1586380591] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 15.2% (20040356 kB) free!
[1586380609] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.6% (45139880 kB) free.
[1586380696] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (4926756 kB) free!
[1586381296] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18364944 kB) free!
[1586381569] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(273) < 103.16.101.72 [09/Apr/2020:03:00:23 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1586381592] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3266) < 203.212.70.170 [09/Apr/2020:03:03:07 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1586381617] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.2% (2504832 kB) free!
[1586381689] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586381712] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(187) < 49.50.86.152 [09/Apr/2020:03:04:59 +0530] TCP 502 0 0 120.001 "10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.001, 0.000, 0.000"
[1586381772] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.8% (2657736 kB) free!
[1586381832] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(21) < 172.105.59.173 [09/Apr/2020:03:07:12 +0530] TCP 502 0 0 120.002 "10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.142:3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.001, 0.000, 0.001"
[1586381863] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 44.6% (10943232 kB) free.
[1586381896] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33601220 kB) free.
[1586381950] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(342) < 2020/04/09 03:01:32 [error] 22134#0: *82451571 upstream timed out (110: Connection timed out) while connecting to upstream, client: 182.18.144.246, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1586382071] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586382217] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183024 kB) free!
[1586382432] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586382496] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11297200 kB) free!
[1586382616] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7481920 kB) free!
[1586382736] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (4043416 kB) free!
[1586383336] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (17028920 kB) free!
[1586383936] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28656732 kB) free.
[1586384172] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2441068 kB) free!
[1586384536] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9917872 kB) free!
[1586384656] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36559924 kB) free.
[1586385256] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15629152 kB) free!
[1586385376] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (11951140 kB) free!
[1586385496] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (8518776 kB) free!
[1586386096] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17940356 kB) free!
[1586387296] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (7510800 kB) free!
[1586387791] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1586387896] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17546600 kB) free!
[1586388890] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1586389064] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (289000 kB) free!
[1586389096] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36106132 kB) free.
[1586389184] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188552 kB) free!
[1586389304] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191660 kB) free!
[1586389696] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12252800 kB) free!
[1586389816] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8350664 kB) free!
[1586389936] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36149800 kB) free.
[1586390017] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 45.4% (11206108 kB) free.
[1586390172] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.4% (20189768 kB) free.
[1586390536] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14704172 kB) free!
[1586390628] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 73.0% (17897760 kB) free.
[1586390656] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (11334464 kB) free!
[1586390776] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33664532 kB) free.
[1586391376] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20113992 kB) free!
[1586391496] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.5% (16487356 kB) free!
[1586391616] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35383900 kB) free.
[1586392216] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20708708 kB) free!
[1586392337] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.9% (17053012 kB) free!
[1586392457] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37143452 kB) free.
[1586393017] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.9% (1693080 kB) free!
[1586393057] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.8% (20782608 kB) free!
[1586393137] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (1441400 kB) free!
[1586393177] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.1% (17320056 kB) free!
[1586393191] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (7412480 kB) free!
[1586393258] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (1440008 kB) free!
[1586393297] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38186172 kB) free.
[1586393791] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 29.4% (38818512 kB) free.
[1586393897] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20636292 kB) free!
[1586394017] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.9% (16961288 kB) free!
[1586394137] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38033756 kB) free.
[1586394737] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19779844 kB) free!
[1586394857] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.2% (16021860 kB) free!
[1586394977] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.8% (37917624 kB) free.
[1586395577] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18482212 kB) free!
[1586395697] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.2% (14744032 kB) free!
[1586395817] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37225228 kB) free.
[1586396417] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17637540 kB) free!
[1586396538] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (13991532 kB) free!
[1586396658] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36193908 kB) free.
[1586397258] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (17107288 kB) free!
[1586397378] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37446312 kB) free.
[1586397978] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19285724 kB) free!
[1586398098] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (15633616 kB) free!
[1586398218] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.3% (37325032 kB) free.
[1586398818] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18350040 kB) free!
[1586398829] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16305.29, 12118.74, 5773.06
[1586398938] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.3% (16225924 kB) free!
[1586398949] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16321.62, 13510.15, 7052.35
[1586399058] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35482460 kB) free.
[1586399069] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16323.60, 14441.13, 8176.51
[1586399230] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16354.42, 13080.85, 6605.82
[1586399232] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(855) < 203.212.70.170 [09/Apr/2020:07:57:10 +0530] TCP 502 0 0 120.002 "10.147.212.154:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.001, 0.000"
[1586399350] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16359.02, 14165.74, 7788.47
[1586399352] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(414) < 203.212.70.170 [09/Apr/2020:07:59:11 +0530] TCP 502 0 0 120.002 "10.147.212.143:3001, 10.147.212.155:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.001, 0.000, 60.001, 0.000"
[1586399470] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16361.05, 14892.18, 8827.91
[1586399472] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(318) < 203.212.70.170 [09/Apr/2020:08:01:11 +0530] TCP 502 0 0 120.001 "10.147.212.143:3001, 10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.001, 60.000, 0.000, 0.000"
[1586399658] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17347856 kB) free!
[1586399778] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37031948 kB) free.
[1586400378] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17767744 kB) free!
[1586400498] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.5% (13876032 kB) free!
[1586400618] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36473312 kB) free.
[1586400673] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586401034] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.80, 6.34, 5.28
[1586401154] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.13, 6.20, 5.36
[1586401219] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16515568 kB) free!
[1586401274] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.33, 6.43, 5.53
[1586401339] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (12539020 kB) free!
[1586401459] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33279764 kB) free.
[1586402059] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12729056 kB) free!
[1586402070] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.34, 5.53, 583.94
[1586402179] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29204828 kB) free.
[1586402471] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.24, 3.14, 366.91
[1586402779] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.6% (8700660 kB) free!
[1586402899] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4825308 kB) free!
[1586403019] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27828052 kB) free.
[1586403619] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.8% (3638416 kB) free!
[1586403739] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (405488 kB) free!
[1586403859] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1586403859] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28759844 kB) free.
[1586404459] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6506156 kB) free!
[1586404579] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2186056 kB) free!
[1586404699] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27756132 kB) free.
[1586405028] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (970784 kB) free!
[1586405148] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (626704 kB) free!
[1586405268] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (320116 kB) free!
[1586405299] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.2% (4248888 kB) free!
[1586405377] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 9.7% (2380752 kB) free!
[1586405420] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (374288 kB) free!
[1586405540] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30001268 kB) free.
[1586406140] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7198580 kB) free!
[1586406260] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.2% (2875028 kB) free!
[1586406380] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (376848 kB) free!
[1586407580] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.2% (18756284 kB) free!
[1586407591] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.0% (22418148 kB) free!
[1586407711] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.8% (24747636 kB) free!
[1586407831] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.2% (27977052 kB) free.
[1586408180] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (397824 kB) free!
[1586408780] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1586408780] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15731540 kB) free!
[1586409031] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 15.7% (20689976 kB) free!
[1586409151] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.5% (23124832 kB) free!
[1586409271] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.5% (25688556 kB) free!
[1586409380] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30752584 kB) free.
[1586409750] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 360 processes
[1586409870] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 360 processes
[1586409980] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5191264 kB) free!
[1586409990] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 362 processes
[1586410100] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (642724 kB) free!
[1586410220] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (404452 kB) free!
[1586410314] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 11.6% (2602720 kB) free!
[1586410435] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (2008828 kB) free!
[1586410555] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (1401396 kB) free!
[1586411089] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (3049604 kB) free!
[1586411209] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.7% (3846496 kB) free.
[1586411420] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28812204 kB) free.
[1586412021] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4543084 kB) free!
[1586412141] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (382956 kB) free!
[1586412261] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389312 kB) free!
[1586413279] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 15.1% (5274268 kB) free!
[1586413399] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 14.0% (4896272 kB) free!
[1586413461] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.9% (26248588 kB) free!
[1586413519] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 12.8% (4485212 kB) free!
[1586414061] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2385864 kB) free!
[1586414071] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (11884020 kB) free!
[1586414120] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 90.3% (31633956 kB) free.
[1586414155] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.2% (16220104 kB) free.
[1586414209] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (3088628 kB) free!
[1586414329] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.2% (3763464 kB) free.
[1586414671] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.0% (25111480 kB) free!
[1586415529] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.5% (2189920 kB) free!
[1586415649] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3357516 kB) free.
[1586415861] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.2% (13432236 kB) free!
[1586416249] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.4% (2021920 kB) free!
[1586416369] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (3184768 kB) free!
[1586416461] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29071188 kB) free.
[1586416489] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.4% (3155028 kB) free!
[1586416551] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9419504 kB) free!
[1586416671] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (9328996 kB) free!
[1586416791] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (9211404 kB) free!
[1586417061] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (5593044 kB) free!
[1586417181] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (1515060 kB) free!
[1586417301] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33373864 kB) free.
[1586417670] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16052.10, 9352.21, 3914.02
[1586417791] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16250.65, 11644.92, 5413.56
[1586417901] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9936224 kB) free!
[1586417911] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16277.14, 13179.05, 6731.29
[1586418022] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.4% (5757148 kB) free!
[1586418071] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16191.84, 11128.39, 5044.16
[1586418073] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1061) < 103.13.97.241 [09/Apr/2020:13:11:13 +0530] TCP 502 0 0 120.001 "10.147.212.154:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1586418142] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (1711364 kB) free!
[1586418192] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16236.54, 12820.90, 6402.10
[1586418193] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(522) < 203.212.70.170 [09/Apr/2020:13:13:13 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1586418312] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16242.24, 13953.24, 7595.32
[1586418314] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(761) < 203.212.70.170 [09/Apr/2020:13:15:13 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1586418390] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 843 processes
[1586418858] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 338 processes
[1586418978] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1586419342] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.0% (34255016 kB) free.
[1586419514] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586419609] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9336336 kB) free!
[1586419729] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.5% (9127256 kB) free!
[1586419850] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.2% (8959736 kB) free!
[1586419942] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9847892 kB) free!
[1586420062] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5368428 kB) free!
[1586420182] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (987252 kB) free!
[1586420782] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16171712 kB) free!
[1586420911] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.88, 4.32, 562.04
[1586420991] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4544924 kB) free!
[1586421312] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.90, 3.63, 410.01
[1586421382] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33167544 kB) free.
[1586421889] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (1488556 kB) free!
[1586421982] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (7761084 kB) free!
[1586422102] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (3227868 kB) free!
[1586422222] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (382260 kB) free!
[1586422489] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.9% (2744520 kB) free!
[1586422822] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (13933888 kB) free!
[1586423089] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1586423089] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (1591120 kB) free!
[1586423156] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.2% (3646216 kB) free!
[1586423276] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.1% (3163084 kB) free!
[1586423396] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.0% (2924744 kB) free!
[1586423422] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34498520 kB) free.
[1586423450] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4633588 kB) free!
[1586423689] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.6% (2537180 kB) free!
[1586423835] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 336 processes
[1586423955] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 4 warn out of 336 processes
[1586423996] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (1805084 kB) free!
[1586424022] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8298228 kB) free!
[1586424075] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1586424143] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (3581812 kB) free!
[1586424263] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (403472 kB) free!
[1586424373] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193224 kB) free!
[1586424378] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1586424493] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190388 kB) free!
[1586424498] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1586424508] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1586424596] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1586424596] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.1% (15748272 kB) free.
[1586424613] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173736 kB) free!
[1586424618] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1586424629] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1586424749] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1586424863] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.7% (15416124 kB) free!
[1586424914] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(158) < 92.42.104.166 [09/Apr/2020:15:00:12 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1586425034] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1586425218] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1586425349] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1586425463] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37035324 kB) free.
[1586425471] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (12209504 kB) free!
[1586426063] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9586660 kB) free!
[1586426072] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.6% (21892852 kB) free!
[1586426089] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (1379932 kB) free!
[1586426183] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.6% (4797124 kB) free!
[1586426303] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389204 kB) free!
[1586426903] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19558116 kB) free!
[1586427272] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (10485052 kB) free!
[1586427503] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (375168 kB) free!
[1586427872] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1586427872] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 12.0% (15754068 kB) free!
[1586428103] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.2% (14812976 kB) free!
[1586428150] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 72.1% (17665288 kB) free.
[1586428490] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.7% (2068024 kB) free!
[1586428668] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 39.2% (9609868 kB) free.
[1586428703] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37705736 kB) free.
[1586428875] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 330 processes
[1586428995] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1586429072] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (9387308 kB) free!
[1586429115] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1586429303] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (374668 kB) free!
[1586429423] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (406588 kB) free!
[1586429543] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401868 kB) free!
[1586430143] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18360120 kB) free!
[1586430744] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (401472 kB) free!
[1586430872] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.8% (18225520 kB) free!
[1586431515] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1586431635] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1586432072] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (11687324 kB) free!
[1586432268] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.9% (4393632 kB) free!
[1586432388] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.3% (4243840 kB) free!
[1586432396] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.5% (4161704 kB) free!
[1586432508] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.7% (4093988 kB) free!
[1586432516] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.2% (3857932 kB) free!
[1586432637] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.0% (3581840 kB) free!
[1586432690] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (1502628 kB) free!
[1586433237] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (858184 kB) free!
[1586433290] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 16.0% (2602904 kB) free!
[1586433708] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (1626480 kB) free!
[1586433744] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20479456 kB) free!
[1586433872] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.1% (22485396 kB) free!
[1586434344] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402960 kB) free!
[1586434490] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.3% (3783328 kB) free.
[1586434813] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 29.2% (7159832 kB) free.
[1586435090] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (2978620 kB) free!
[1586435210] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3322012 kB) free.
[1586435474] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 31.39, 28.83, 20.69
[1586435637] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.7% (16093552 kB) free.
[1586435810] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.4% (2009136 kB) free!
[1586435930] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.1% (2944928 kB) free!
[1586435950] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.4% (4520972 kB) free!
[1586436013] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4882144 kB) free!
[1586436051] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.1% (2941428 kB) free!
[1586436070] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.2% (4218904 kB) free!
[1586436133] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (4717132 kB) free!
[1586436144] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17128136 kB) free!
[1586436190] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 15.7% (3853164 kB) free!
[1586436253] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (4507764 kB) free!
[1586436272] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (500820 kB) free!
[1586436744] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (435868 kB) free!
[1586436790] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2313736 kB) free!
[1586436902] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 344 processes
[1586437022] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1586437142] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1586437251] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.5% (4796980 kB) free.
[1586437274] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.90, 10.36, 16.92
[1586437344] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.2% (17394492 kB) free!
[1586437472] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.1% (14567964 kB) free!
[1586437742] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1586437944] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (454556 kB) free!
[1586438342] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1586438462] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1586438544] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.4% (15015164 kB) free!
[1586438653] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2156836 kB) free!
[1586439144] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (388272 kB) free!
[1586439618] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1586439738] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1586440251] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.4% (2672564 kB) free!
[1586440371] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.9% (3554808 kB) free.
[1586440971] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.1% (2776704 kB) free!
[1586441091] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3319028 kB) free.
[1586441691] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.6% (2697656 kB) free!
[1586441811] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (3085300 kB) free!
[1586441931] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (3070996 kB) free!
[1586442272] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (7326072 kB) free!
[1586442531] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 32.5% (5289644 kB) free.
[1586442872] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.2% (27935552 kB) free.
[1586444072] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.9% (26192764 kB) free!
[1586444192] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.8% (18208944 kB) free!
[1586444312] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.4% (28233232 kB) free.
[1586444931] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (2840600 kB) free!
[1586445051] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.4% (3644748 kB) free.
[1586445144] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36195960 kB) free.
[1586445237] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4463788 kB) free!
[1586445357] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 19.2% (4305636 kB) free!
[1586445477] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 18.2% (4095788 kB) free!
[1586445512] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 20.0% (26325576 kB) free!
[1586445632] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.7% (18087728 kB) free!
[1586445651] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.0% (2602124 kB) free!
[1586445744] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.0% (5274160 kB) free!
[1586445752] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (11395584 kB) free!
[1586445771] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.7% (3361576 kB) free.
[1586445864] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (598016 kB) free!
[1586445984] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (391316 kB) free!
[1586446352] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.0% (25111316 kB) free!
[1586446371] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.2% (2463804 kB) free!
[1586446491] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (3154492 kB) free!
[1586446584] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.7% (20715952 kB) free!
[1586446611] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.9% (3882148 kB) free.
[1586446678] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 72.1% (16189188 kB) free.
[1586447185] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (467876 kB) free!
[1586447211] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.9% (2743172 kB) free!
[1586447331] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.5% (3494380 kB) free.
[1586447785] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16195200 kB) free!
[1586447931] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.5% (2363404 kB) free!
[1586448051] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (3029900 kB) free!
[1586448171] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.9% (3729740 kB) free.
[1586448385] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (404560 kB) free!
[1586448752] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.2% (27956456 kB) free.
[1586448771] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.7% (2546716 kB) free!
[1586448891] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3342812 kB) free.
[1586449492] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.6% (2043580 kB) free!
[1586449612] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.8% (2731380 kB) free!
[1586449732] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3299920 kB) free.
[1586449952] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.8% (23478440 kB) free!
[1586450073] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 11.9% (15711448 kB) free!
[1586450193] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.1% (25152952 kB) free!
[1586450932] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3219892 kB) free!
[1586451052] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.5% (3983556 kB) free.
[1586451393] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.8% (31313444 kB) free.
[1586451674] SERVICE ALERT: sigqscluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 58.70, 30.55, 15.79
[1586451993] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.1% (21221420 kB) free!
[1586452113] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.7% (29915448 kB) free.
[1586452274] SERVICE ALERT: sigqscluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 14.64, 19.98, 18.76
[1586453313] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.7% (16747636 kB) free!
[1586453433] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.9% (27523876 kB) free.
[1586453452] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.6% (2533960 kB) free!
[1586453572] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (3092036 kB) free!
[1586453692] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.4% (3809748 kB) free.
[1586454033] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.9% (22274800 kB) free!
[1586454153] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.4% (29552984 kB) free.
[1586454292] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.6% (2218552 kB) free!
[1586454385] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16469704 kB) free!
[1586454412] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.4% (2833540 kB) free!
[1586454532] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3377432 kB) free.
[1586454985] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (10465152 kB) free!
[1586455132] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.5% (1713064 kB) free!
[1586455253] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.4% (2174984 kB) free!
[1586455373] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.9% (2590088 kB) free!
[1586455973] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.2% (4100732 kB) free.
[1586456785] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16847252 kB) free!