[1581964313] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/02/18 00:00:59 [error] 12257#0: *55562339 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581964433] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581964555] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 58.6% (28858724 kB) free.
[1581964686] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24270028 kB) free!
[1581964788] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189272 kB) free!
[1581964806] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.7% (20702792 kB) free!
[1581964809] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(19) < 203.212.70.170 [18/Feb/2020:00:02:58 +0530] TCP 502 0 0 120.000 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1581964876] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175344 kB) free!
[1581964909] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199192 kB) free!
[1581964927] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.7% (47062240 kB) free.
[1581964929] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581964996] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180836 kB) free!
[1581965034] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169912 kB) free!
[1581965105] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192464 kB) free!
[1581965116] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176056 kB) free!
[1581965226] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196544 kB) free!
[1581965346] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183000 kB) free!
[1581965526] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23726844 kB) free!
[1581965646] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.3% (20211248 kB) free!
[1581965766] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.4% (45410380 kB) free.
[1581965937] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.8% (14202916 kB) free!
[1581966160] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1581966280] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1581966336] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185004 kB) free!
[1581966366] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23268284 kB) free!
[1581966396] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192224 kB) free!
[1581966456] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 66.2% (16219360 kB) free.
[1581966486] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.0% (48818040 kB) free.
[1581966730] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 103.16.101.72 [18/Feb/2020:00:36:53 +0530] TCP 502 0 0 60.003 "10.147.212.142:13000, 10.147.212.155:13000, 10.147.212.143:13000, 10.147.212.154:13000" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 0.000, 0.001, 60.002"
[1581966850] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581967289] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1581967409] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1581967434] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 44.7% (11034340 kB) free.
[1581967737] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (12497348 kB) free!
[1581968634] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165520 kB) free!
[1581968754] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197080 kB) free!
[1581968874] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190236 kB) free!
[1581968917] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.1% (1011424 kB) free!
[1581968937] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1581968937] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.8% (14272588 kB) free!
[1581971886] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24161116 kB) free!
[1581972007] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.4% (46658528 kB) free.
[1581972461] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 705 processes
[1581972546] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.7% (9980804 kB) free.
[1581972581] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 705 processes
[1581972606] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25452848 kB) free!
[1581972701] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1581972726] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.9% (48621000 kB) free.
[1581972886] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (219836 kB) free!
[1581973006] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (240184 kB) free!
[1581973126] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (216320 kB) free!
[1581973238] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (224232 kB) free!
[1581973326] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26190280 kB) free!
[1581973358] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (238120 kB) free!
[1581973446] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.9% (48583316 kB) free.
[1581973478] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (219860 kB) free!
[1581973726] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 80.1% (39491208 kB) free.
[1581974046] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1581974078] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 86.3% (42509660 kB) free.
[1581974501] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1581975101] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 705 processes
[1581975221] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1581975233] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(300) < 2020/02/18 03:01:31 [error] 12251#0: *55599149 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1581975250] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2434) < 203.212.70.170 [18/Feb/2020:03:02:31 +0530] TCP 502 0 0 120.001 "10.147.212.142:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.000, 0.000"
[1581975281] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(237) < 103.16.101.72 [18/Feb/2020:03:00:31 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1581975341] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1581975353] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581975371] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581975403] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581977646] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21921356 kB) free!
[1581977766] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.4% (47968116 kB) free.
[1581978366] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24200644 kB) free!
[1581978486] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.4% (48011224 kB) free.
[1581979086] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24824480 kB) free!
[1581979207] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.3% (47856876 kB) free.
[1581979737] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1581982746] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169180 kB) free!
[1581982866] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 40.5% (9937280 kB) free.
[1581983467] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186580 kB) free!
[1581983587] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192784 kB) free!
[1581983707] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193548 kB) free!
[1581983717] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.6% (21221832 kB) free.
[1581983741] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1581983862] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1581983874] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 45.2% (11163620 kB) free.
[1581986591] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.4% (39644244 kB) free.
[1581986623] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.5% (21939868 kB) free.
[1581986937] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 40.8% (53839628 kB) free.
[1581987953] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(98) < 2020/02/18 06:35:09 [error] 12245#0: *55643177 access forbidden by rule while initializing session, client: 124.153.112.226, server: 0.0.0.0:4000
[1581988073] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(50) < 2020/02/18 06:37:09 [error] 12255#0: *55643956 access forbidden by rule while initializing session, client: 124.153.112.226, server: 0.0.0.0:4000
[1581988074] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191940 kB) free!
[1581988194] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;HARD;3;(50) < 2020/02/18 06:39:10 [error] 12248#0: *55644707 access forbidden by rule while initializing session, client: 124.153.112.226, server: 0.0.0.0:4000
[1581988194] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (200140 kB) free!
[1581988314] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (203216 kB) free!
[1581989698] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.8% (21923832 kB) free.
[1581994002] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3) < 103.16.101.72 [18/Feb/2020:08:16:06 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1581994122] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1) < 103.16.101.72 [18/Feb/2020:08:16:56 +0530] TCP 502 0 0 61.029 "10.147.212.49:4000, 10.147.212.61:4000, 10.147.212.38:4000, 10.147.212.37:4000" "0, 0, 0, 0" "0, 0, 0, 0" "1.003, 0.000, 0.001, 60.025"
[1581994242] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581994570] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(12) < 203.212.70.170 [18/Feb/2020:08:18:05 +0530] TCP 502 0 0 120.001 "10.147.212.142:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1581994690] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581995291] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1075) < 180.179.218.150 [18/Feb/2020:08:33:35 +0530] TCP 502 0 0 0.000 "backend_13000" "0" "0" "0.000"
[1581995411] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581995442] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(185) < 103.16.101.72 [18/Feb/2020:08:31:30 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1581995562] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1581998557] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 20.0% (4890452 kB) free!
[1581998677] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.5% (4771208 kB) free!
[1581998797] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.0% (4657032 kB) free!
[1582000597] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2222476 kB) free!
[1582000718] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.2% (21609776 kB) free.
[1582001080] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1582001107] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4684840 kB) free!
[1582001200] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 332 processes
[1582001227] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.4% (4535324 kB) free!
[1582001320] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 332 processes
[1582001347] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.0% (4448312 kB) free!
[1582001995] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582002520] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 335 processes
[1582002547] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (381524 kB) free!
[1582002778] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.2% (22005948 kB) free.
[1582002829] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.7% (14090056 kB) free.
[1582002917] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (202672 kB) free!
[1582003038] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180196 kB) free!
[1582003120] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 332 processes
[1582003130] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191488 kB) free!
[1582003158] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163868 kB) free!
[1582003250] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169364 kB) free!
[1582003370] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1582003370] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183732 kB) free!
[1582003720] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1582003720] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 339 processes
[1582003806] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.1% (17330080 kB) free!
[1582003926] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38835452 kB) free.
[1582003971] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 81.1% (19876964 kB) free.
[1582004261] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 704 processes
[1582004381] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 703 processes
[1582004501] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 701 processes
[1582004526] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14091480 kB) free!
[1582004550] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.9% (21544648 kB) free.
[1582004646] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36109696 kB) free.
[1582004827] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.7% (3068696 kB) free!
[1582004902] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1582004948] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 11.9% (2677400 kB) free!
[1582005012] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(517) < 103.255.146.164 [18/Feb/2020:11:12:15 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1582005022] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1582005068] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2247844 kB) free!
[1582005101] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 697 processes
[1582005132] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582005142] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1582005247] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10810332 kB) free!
[1582005367] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.1% (46232632 kB) free.
[1582005756] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181872 kB) free!
[1582005786] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (1776508 kB) free!
[1582005792] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3087380 kB) free!
[1582005877] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191256 kB) free!
[1582005907] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (1574540 kB) free!
[1582005912] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (3082668 kB) free!
[1582005956] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.3% (22038616 kB) free.
[1582005967] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22478588 kB) free!
[1582005997] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185644 kB) free!
[1582006032] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (3076520 kB) free!
[1582006060] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.7% (39272544 kB) free.
[1582006087] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.8% (45863688 kB) free.
[1582006095] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1582006095] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.2% (21374052 kB) free.
[1582006342] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1582006376] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.7% (15864628 kB) free.
[1582006556] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7462000 kB) free!
[1582006676] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 13.7% (6759588 kB) free!
[1582006687] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (21773664 kB) free!
[1582006748] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4907336 kB) free!
[1582006796] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 12.3% (6049144 kB) free!
[1582006807] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (17998116 kB) free!
[1582006868] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 17.8% (4385624 kB) free!
[1582006927] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42168556 kB) free.
[1582006988] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 15.9% (3914100 kB) free!
[1582007396] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (3226812 kB) free!
[1582007527] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18591876 kB) free!
[1582007588] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1939016 kB) free!
[1582007647] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.7% (41814064 kB) free.
[1582007920] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1582007995] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(71) < 2020/02/18 12:01:51 [error] 12255#0: *55793184 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582007996] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.9% (30001740 kB) free.
[1582008115] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582008133] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(12) < 172.105.59.173 [18/Feb/2020:12:02:59 +0530] TCP 502 0 0 120.000 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1582008247] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (18945144 kB) free!
[1582008253] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582008366] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196232 kB) free!
[1582008367] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.3% (14936132 kB) free!
[1582008486] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181436 kB) free!
[1582008487] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39166924 kB) free.
[1582008606] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165164 kB) free!
[1582008788] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.7% (22875944 kB) free.
[1582009087] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.7% (15358040 kB) free!
[1582009097] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4683620 kB) free!
[1582009131] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1582009207] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38684320 kB) free.
[1582009217] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4518664 kB) free!
[1582009251] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1582009337] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (4353276 kB) free!
[1582009426] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4797180 kB) free!
[1582009546] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.0% (4660268 kB) free!
[1582009666] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (4515584 kB) free!
[1582009807] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (16980556 kB) free!
[1582009928] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.8% (12952572 kB) free!
[1582010009] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1582010048] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37172708 kB) free.
[1582010078] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.5% (8648356 kB) free!
[1582010129] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1582010198] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.0% (8362220 kB) free!
[1582010318] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.5% (8114916 kB) free!
[1582010648] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.6% (13917632 kB) free!
[1582010768] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9813384 kB) free!
[1582010815] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.0% (21972856 kB) free.
[1582010863] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 59.7% (78663412 kB) free.
[1582010951] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.0% (42893760 kB) free.
[1582010993] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.0% (21821300 kB) free.
[1582011026] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.9% (21547636 kB) free.
[1582012971] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.5% (4779716 kB) free!
[1582013091] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.7% (4589564 kB) free!
[1582013211] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.1% (4440796 kB) free!
[1582013269] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9827328 kB) free!
[1582013389] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (9827212 kB) free!
[1582013510] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.9% (9824072 kB) free!
[1582014411] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2357680 kB) free!
[1582014419] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (4496984 kB) free!
[1582014539] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.5% (4325436 kB) free!
[1582014659] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 16.7% (4125376 kB) free!
[1582014788] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.7% (3741920 kB) free!
[1582014908] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.6% (3278088 kB) free!
[1582015028] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.0% (2923832 kB) free!
[1582015066] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24836056 kB) free!
[1582015186] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35723860 kB) free.
[1582015226] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.6% (21714808 kB) free.
[1582015259] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1582015259] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (1001608 kB) free!
[1582015548] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.6% (21858416 kB) free.
[1582015628] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (1118480 kB) free!
[1582015720] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1582015726] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9768320 kB) free!
[1582015786] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12938524 kB) free!
[1582015840] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 337 processes
[1582015846] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (9585920 kB) free!
[1582015851] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1582015906] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27662752 kB) free.
[1582015960] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1582015960] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1582015966] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (9398736 kB) free!
[1582015971] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1582016507] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16911712 kB) free!
[1582016509] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4225672 kB) free!
[1582016561] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 336 processes
[1582016627] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (12602412 kB) free!
[1582016681] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 332 processes
[1582016747] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (8644696 kB) free!
[1582016801] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 332 processes
[1582016996] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.5% (9125584 kB) free!
[1582017116] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.5% (8628784 kB) free!
[1582017189] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4687628 kB) free!
[1582017236] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.7% (8230180 kB) free!
[1582017309] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 18.0% (4435276 kB) free!
[1582017429] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 16.9% (4168020 kB) free!
[1582017947] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28686156 kB) free.
[1582017955] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180276 kB) free!
[1582018001] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1582018075] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (168448 kB) free!
[1582018195] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (201096 kB) free!
[1582018195] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177528 kB) free!
[1582018236] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (431920 kB) free!
[1582018315] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195648 kB) free!
[1582018315] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/02/18 14:52:59 [error] 12257#0: *55909993 open() "/usr/share/nginx/html/js/tmpl/404.ejs" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /js/tmpl/404.ejs?0.5846091757790712 HTTP/1.1", host: "qservercluster1.com:8091", referrer: "http://qservercluster1.com:8091/"
[1582018356] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179576 kB) free!
[1582018435] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184120 kB) free!
[1582018436] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582018437] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1582018437] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.9% (29504104 kB) free.
[1582018454] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(110) < 49.50.67.175 [18/Feb/2020:15:00:11 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1582018476] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194132 kB) free!
[1582018547] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.8% (2325832 kB) free!
[1582018574] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582018628] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.6% (16076916 kB) free.
[1582018629] SERVICE FLAPPING ALERT: db1-vm2;Memory;STARTED; Service appears to have started flapping (24.0% change >= 20.0% threshold)
[1582018629] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (2010316 kB) free!
[1582018667] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (378832 kB) free!
[1582018787] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37156400 kB) free.
[1582018789] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.3% (1998280 kB) free!
[1582018909] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.6% (2045828 kB) free!
[1582019029] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.9% (4705020 kB) free.
[1582019229] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.7% (22877208 kB) free.
[1582019249] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.5% (21344956 kB) free.
[1582019315] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.4% (39616048 kB) free.
[1582019337] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1582019337] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.8% (22013320 kB) free.
[1582019369] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.7% (21734152 kB) free.
[1582019387] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6710888 kB) free!
[1582019419] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 84.6% (41678760 kB) free.
[1582019455] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 40.6% (53563496 kB) free.
[1582019801] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1582019821] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582019922] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1582019942] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 349 processes
[1582019975] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193228 kB) free!
[1582020037] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187432 kB) free!
[1582020042] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1582020042] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1582020062] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1582020095] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196664 kB) free!
[1582020157] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192804 kB) free!
[1582020205] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;1;CRITICAL - 0.8% (187048 kB) free!
[1582020215] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190928 kB) free!
[1582020229] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.7% (2721716 kB) free!
[1582020277] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181704 kB) free!
[1582020282] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1582020325] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;2;CRITICAL - 0.8% (199068 kB) free!
[1582020349] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.8% (2892292 kB) free!
[1582020361] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1582020402] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1582020445] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 0.8% (193164 kB) free!
[1582020456] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169600 kB) free!
[1582020469] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (3048956 kB) free!
[1582020481] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1582020522] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1582020576] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178692 kB) free!
[1582020601] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1582020642] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 343 processes
[1582020662] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1582020696] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182292 kB) free!
[1582020877] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 83.7% (20512096 kB) free.
[1582021045] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 89.1% (21835872 kB) free.
[1582021069] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.1% (3922836 kB) free.
[1582021201] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1582021242] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1582021257] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21113984 kB) free!
[1582021377] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (16915072 kB) free!
[1582021497] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13307136 kB) free!
[1582021842] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 346 processes
[1582021963] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 49.50.86.152 [18/Feb/2020:16:01:00 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1582022015] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.7% (4091692 kB) free!
[1582022035] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(194) < 2020/02/18 16:02:04 [error] 12257#0: *55952813 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582022084] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582022156] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582022174] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 172.105.59.173 [18/Feb/2020:16:01:55 +0530] TCP 502 0 0 60.001 "10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 0.000, 60.000, 0.000"
[1582022269] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.7% (2388052 kB) free!
[1582022294] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582022389] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.7% (2708456 kB) free!
[1582022442] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1582022509] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (2991204 kB) free!
[1582022615] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.6% (1129336 kB) free!
[1582022697] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1582022697] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42228788 kB) free.
[1582023042] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1582023096] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.3% (21535808 kB) free.
[1582023128] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.8% (21773896 kB) free.
[1582023142] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1582023162] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1582023262] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1582023282] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1582023284] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1582023284] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 43.7% (7108780 kB) free.
[1582023297] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18749136 kB) free!
[1582023418] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.4% (50647272 kB) free.
[1582023517] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9762820 kB) free!
[1582023637] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 13.9% (6830664 kB) free!
[1582023705] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193744 kB) free!
[1582023757] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (4035244 kB) free!
[1582023825] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179712 kB) free!
[1582023945] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165776 kB) free!
[1582024017] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24386920 kB) free!
[1582024137] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.5% (20378112 kB) free!
[1582024154] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198160 kB) free!
[1582024257] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16578388 kB) free!
[1582024274] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167456 kB) free!
[1582024394] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188448 kB) free!
[1582024857] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27755832 kB) free.
[1582024938] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191900 kB) free!
[1582025059] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171620 kB) free!
[1582025062] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1582025179] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (209864 kB) free!
[1582025182] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1582025302] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1582026057] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15848068 kB) free!
[1582026177] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (11827788 kB) free!
[1582026297] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.7% (44413424 kB) free.
[1582026428] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.8% (21662712 kB) free.
[1582026463] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.4% (39610796 kB) free.
[1582026493] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.9% (22279108 kB) free.
[1582026502] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1582026523] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.2% (21856288 kB) free.
[1582026829] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.4% (2832144 kB) free!
[1582026887] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.3% (2981948 kB) free!
[1582026897] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18047976 kB) free!
[1582026949] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.4% (2835576 kB) free!
[1582027007] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.9% (3552068 kB) free.
[1582027017] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.6% (13929684 kB) free!
[1582027069] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 17.5% (2838340 kB) free!
[1582027102] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1582027137] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29921848 kB) free.
[1582027156] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9088108 kB) free!
[1582027222] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1582027276] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8877528 kB) free!
[1582027342] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1582027396] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.5% (8620448 kB) free!
[1582027555] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/18 17:33:30 [error] 12257#0: *56015492 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1582027628] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.9% (3335516 kB) free!
[1582027636] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175344 kB) free!
[1582027675] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582027726] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170436 kB) free!
[1582027737] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18781824 kB) free!
[1582027749] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.4% (3000844 kB) free!
[1582027756] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175096 kB) free!
[1582027846] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (217520 kB) free!
[1582027857] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.8% (14289844 kB) free!
[1582027869] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.1% (2713332 kB) free!
[1582027876] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182552 kB) free!
[1582027966] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181108 kB) free!
[1582027977] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (10268316 kB) free!
[1582028301] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186232 kB) free!
[1582028421] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184460 kB) free!
[1582028469] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (1136904 kB) free!
[1582028541] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187248 kB) free!
[1582028577] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.2% (21358320 kB) free!
[1582029069] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1582029069] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.6% (16085208 kB) free.
[1582029178] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34388340 kB) free.
[1582029196] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4501664 kB) free!
[1582029238] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1582029407] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.8% (1756356 kB) free!
[1582029527] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.2% (2308344 kB) free!
[1582029647] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (2967988 kB) free!
[1582029778] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9453076 kB) free!
[1582029898] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4509364 kB) free!
[1582030018] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (650224 kB) free!
[1582030247] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.2% (853192 kB) free!
[1582030630] SERVICE FLAPPING ALERT: db1-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582030665] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (4272880 kB) free!
[1582030785] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (270528 kB) free!
[1582030847] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3166892 kB) free!
[1582030905] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (255988 kB) free!
[1582031137] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.5% (21588392 kB) free.
[1582031189] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 53.2% (70129152 kB) free.
[1582031217] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1582031217] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 60.9% (30008476 kB) free.
[1582031248] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.9% (22288520 kB) free.
[1582031278] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 47.6% (23473596 kB) free.
[1582031306] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.5% (21948744 kB) free.
[1582031335] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.8% (14112624 kB) free.
[1582031356] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 51.6% (8383128 kB) free.
[1582031747] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171576 kB) free!
[1582031867] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169540 kB) free!
[1582031987] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173460 kB) free!
[1582032277] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582032427] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (237088 kB) free!
[1582032547] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (228340 kB) free!
[1582032568] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.3% (2819616 kB) free!
[1582032667] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237208 kB) free!
[1582032688] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3316988 kB) free.
[1582033596] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (24991884 kB) free!
[1582033716] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.9% (20958628 kB) free!
[1582033836] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16987552 kB) free!
[1582033888] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.1% (2936860 kB) free!
[1582034008] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.3% (3305740 kB) free.
[1582034082] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582034436] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28109076 kB) free.
[1582034542] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1582034677] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.5% (4777864 kB) free!
[1582034797] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.2% (4699936 kB) free!
[1582034917] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.9% (4622100 kB) free!
[1582035208] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.2% (2476988 kB) free!
[1582035328] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.6% (2701328 kB) free!
[1582035448] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.0% (2925724 kB) free!
[1582035537] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 76.6% (18900204 kB) free.
[1582035581] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.1% (38475124 kB) free.
[1582035636] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.0% (21149624 kB) free!
[1582035756] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.9% (17017504 kB) free!
[1582035876] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.3% (50441564 kB) free.
[1582036048] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.3% (3793520 kB) free.
[1582036145] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194712 kB) free!
[1582036186] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (245752 kB) free!
[1582036265] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (300844 kB) free!
[1582036306] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (226160 kB) free!
[1582036385] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190740 kB) free!
[1582036426] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (219364 kB) free!
[1582036476] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26227436 kB) free!
[1582036596] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.6% (20550872 kB) free!
[1582036675] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(30) < 2020/02/18 20:01:02 [error] 12250#0: *56098270 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582036694] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(33) < 203.212.70.170 [18/Feb/2020:20:03:57 +0530] TCP 502 0 0 183.007 "10.147.212.143:3001, 10.147.212.142:3001, 10.147.212.154:3001, 10.147.212.155:3001" "0, 0, 0, 0" "0, 0, 0, 0" "3.006, 60.000, 60.000, 60.001"
[1582036716] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.1% (15915032 kB) free!
[1582036795] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582036814] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582037316] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1582037316] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31266668 kB) free.
[1582038448] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.0% (1953676 kB) free!
[1582038455] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1582038517] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (18629408 kB) free!
[1582038568] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 35.3% (5739296 kB) free.
[1582038637] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.1% (14679776 kB) free!
[1582038757] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.3% (49109268 kB) free.
[1582039357] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24540964 kB) free!
[1582039477] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.4% (20252252 kB) free!
[1582039597] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16365676 kB) free!
[1582040197] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (29993472 kB) free.
[1582040317] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2386812 kB) free!
[1582040469] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582041397] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24251440 kB) free!
[1582041517] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.4% (20296788 kB) free!
[1582041637] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.6% (16611676 kB) free!
[1582042237] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28044188 kB) free.
[1582042791] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.5% (21436736 kB) free.
[1582042821] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 68.7% (16945580 kB) free.
[1582042853] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 47.1% (23206516 kB) free.
[1582043437] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19553668 kB) free!
[1582043557] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (15901800 kB) free!
[1582043677] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.2% (47677380 kB) free.
[1582044278] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21581068 kB) free!
[1582044398] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.6% (17873724 kB) free!
[1582044518] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.9% (14335760 kB) free!
[1582044626] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173236 kB) free!
[1582044746] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195552 kB) free!
[1582044866] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192528 kB) free!
[1582045118] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26689468 kB) free.
[1582046318] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19051120 kB) free!
[1582046438] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.8% (15610188 kB) free!
[1582046558] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.1% (46224504 kB) free.
[1582047158] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (22726000 kB) free!
[1582047278] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.6% (19226636 kB) free!
[1582047398] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.0% (15878828 kB) free!
[1582047568] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582048256] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1582048598] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36566672 kB) free.
[1582049198] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17448676 kB) free!
[1582049318] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.5% (45428728 kB) free.
[1582049637] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (234752 kB) free!
[1582049757] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (216608 kB) free!
[1582049877] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (214300 kB) free!
[1582049919] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20527892 kB) free!
[1582049937] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.8% (18222896 kB) free!
[1582050039] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (16917472 kB) free!
[1582050057] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (9179868 kB) free!
[1582050159] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.5% (46779068 kB) free.
[1582050177] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (7330124 kB) free!