[1579458963] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14054064 kB) free!
[1579459083] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36063032 kB) free.
[1579459202] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.7% (29927628 kB) free.
[1579459217] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189560 kB) free!
[1579459250] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185392 kB) free!
[1579459370] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179744 kB) free!
[1579459490] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (201076 kB) free!
[1579459683] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15660196 kB) free!
[1579459803] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12227012 kB) free!
[1579459923] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33937504 kB) free.
[1579460524] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13822920 kB) free!
[1579460601] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183568 kB) free!
[1579460644] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36159648 kB) free.
[1579460661] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 68.4% (16775672 kB) free.
[1579461244] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15094548 kB) free!
[1579461364] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36210592 kB) free.
[1579461680] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.1% (26551416 kB) free.
[1579461964] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16913860 kB) free!
[1579462084] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.2% (13465264 kB) free!
[1579462204] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34887120 kB) free.
[1579462804] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15071552 kB) free!
[1579462924] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1579462924] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36331824 kB) free.
[1579463480] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.0% (4692652 kB) free!
[1579463524] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16286216 kB) free!
[1579463644] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36405124 kB) free.
[1579463937] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.5% (8619516 kB) free!
[1579464057] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 14.8% (7271840 kB) free!
[1579464080] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167012 kB) free!
[1579464177] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 12.7% (6237096 kB) free!
[1579464244] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16930356 kB) free!
[1579464364] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (36457896 kB) free.
[1579464778] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (742152 kB) free!
[1579464964] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17552348 kB) free!
[1579465084] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.9% (20904796 kB) free!
[1579465114] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(161) < 2020/01/20 01:48:28 [error] 21672#0: *25519718 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:58/17, bytes from/to upstream:17/58
[1579465205] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33872592 kB) free.
[1579465235] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2) < 2020/01/20 01:49:20 [error] 21682#0: *25520018 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:54/17, bytes from/to upstream:17/54
[1579465277] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(20) < 2020/01/20 01:51:08 [error] 18357#0: *41567902 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:90/17, bytes from/to upstream:17/90
[1579465355] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579465397] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579465997] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(19) < 2020/01/20 01:58:06 [error] 18358#0: *41569190 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:90/49, bytes from/to upstream:49/90
[1579466117] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579466404] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.9% (17025244 kB) free!
[1579466524] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36643656 kB) free.
[1579466555] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1347) < 2020/01/20 02:12:17 [error] 21682#0: *25524937 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.27.86.30, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579466598] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6) < 203.212.70.170 [20/Jan/2020:02:11:02 +0530] TCP 502 0 0 67.019 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "7.018, 60.001, 0.000"
[1579466675] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(134) < 2020/01/20 02:14:19 [error] 21680#0: *25525528 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:5272/70, bytes from/to upstream:70/5272
[1579466717] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(121) < 2020/01/20 02:14:45 [error] 18370#0: *41572753 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1579466718] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579466795] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(611) < 2020/01/20 02:16:35 [error] 21682#0: *25526259 recv() failed (104: Connection reset by peer) while proxying connection, client: 172.105.59.173, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:41/0, bytes from/to upstream:0/41
[1579466837] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579467124] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17398356 kB) free!
[1579467244] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36716304 kB) free.
[1579467318] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 203.212.70.170 [20/Jan/2020:02:15:44 +0530] TCP 502 0 0 75.028 "10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "15.027, 60.001, 0.000"
[1579467438] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579467844] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18484972 kB) free!
[1579467964] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (15108136 kB) free!
[1579468039] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(286) < 203.212.70.170 [20/Jan/2020:02:37:17 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579468084] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33347432 kB) free.
[1579468159] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(29) < 49.50.86.152 [20/Jan/2020:02:38:26 +0530] TCP 502 0 0 60.000 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1579468279] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579468684] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9176912 kB) free!
[1579468804] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35387888 kB) free.
[1579468817] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.9% (10268460 kB) free.
[1579469195] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579469404] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10788916 kB) free!
[1579469479] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2249) < 49.50.86.152 [20/Jan/2020:03:00:32 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579469524] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35888492 kB) free.
[1579469599] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(217) < 172.105.59.173 [20/Jan/2020:03:02:22 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1579469720] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579469795] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2907) < 2020/01/20 03:02:47 [error] 21682#0: *25536103 recv() failed (104: Connection reset by peer) while proxying connection, client: 172.105.59.173, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:41/0, bytes from/to upstream:0/41
[1579469837] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(263) < 2020/01/20 03:01:32 [error] 18356#0: *41583445 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579469915] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579469957] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579469994] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(214) < 103.16.101.72 [20/Jan/2020:03:00:22 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1579470114] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579470724] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18825204 kB) free!
[1579470845] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39070732 kB) free.
[1579472044] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22408232 kB) free!
[1579472164] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38642472 kB) free.
[1579473364] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22983492 kB) free!
[1579473484] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38808184 kB) free.
[1579474684] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23323840 kB) free!
[1579474804] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38750788 kB) free.
[1579476004] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (22844888 kB) free!
[1579476124] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37681608 kB) free.
[1579476617] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4648660 kB) free!
[1579476738] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174136 kB) free!
[1579476858] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169216 kB) free!
[1579477490] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.9% (21550776 kB) free.
[1579477880] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 61.3% (15137948 kB) free.
[1579477924] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (22819656 kB) free!
[1579478045] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39163416 kB) free.
[1579479244] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24378384 kB) free!
[1579479364] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39164320 kB) free.
[1579481164] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (23530764 kB) free!
[1579481284] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (39381768 kB) free.
[1579482484] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22942264 kB) free!
[1579482604] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38484048 kB) free.
[1579483757] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(160) < 2020/01/20 06:51:18 [error] 18370#0: *41629178 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:90/0, bytes from/to upstream:0/90
[1579483804] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23950356 kB) free!
[1579483877] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579483924] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38812344 kB) free.
[1579485078] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(88) < 2020/01/20 07:20:29 [error] 18358#0: *41636933 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:58/17, bytes from/to upstream:17/58
[1579485198] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579485515] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(459) < 2020/01/20 07:28:33 [error] 21675#0: *25557878 recv() failed (104: Connection reset by peer) while proxying connection, client: 172.105.59.173, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:40/0, bytes from/to upstream:0/40
[1579485544] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1579485635] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(474) < 2020/01/20 07:30:34 [error] 21681#0: *25558499 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:101/0, bytes from/to upstream:0/101
[1579485664] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 398 processes
[1579485755] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(22) < 2020/01/20 07:31:32 [error] 21675#0: *25559748 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:85/0, bytes from/to upstream:0/85
[1579486355] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579487524] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (23806748 kB) free!
[1579487644] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38599368 kB) free.
[1579488064] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1579488155] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(341) < 2020/01/20 08:10:50 [error] 21681#0: *25567138 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:118/0, bytes from/to upstream:0/118
[1579488185] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1579488275] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(150) < 2020/01/20 08:14:04 [error] 21672#0: *25569140 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:39/0, bytes from/to upstream:0/39
[1579488395] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(25) < 2020/01/20 08:14:48 [error] 21681#0: *25569126 upstream timed out (110: Connection timed out) while connecting to upstream, client: 172.105.59.173, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579488798] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(18) < 2020/01/20 08:17:30 [error] 18357#0: *41637194 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:6473/7646, bytes from/to upstream:7646/6473
[1579488919] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579488995] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1579488995] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579489368] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.9% (40351724 kB) free.
[1579489400] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 76.8% (18821552 kB) free.
[1579489519] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(20) < 203.212.70.170 [20/Jan/2020:08:29:29 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579489595] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(277) < 2020/01/20 08:30:22 [error] 21679#0: *25572746 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579489639] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579489715] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(168) < 2020/01/20 08:38:35 [error] 21683#0: *25574702 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.144.246, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:410/0, bytes from/to upstream:0/410
[1579489836] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(412) < 2020/01/20 08:40:34 [error] 21675#0: *25574973 recv() failed (104: Connection reset by peer) while proxying connection, client: 172.105.59.173, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:41/0, bytes from/to upstream:0/41
[1579490044] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.4% (21619092 kB) free!
[1579490118] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(7) < 2020/01/20 08:45:13 [error] 18362#0: *41663780 connect() failed (111: Connection refused) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579490164] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37679428 kB) free.
[1579490238] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(171) < 2020/01/20 08:46:59 [error] 18366#0: *41664251 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579490239] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(31) < 203.212.70.170 [20/Jan/2020:08:41:42 +0530] TCP 502 0 0 120.001 "10.147.212.155:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.001, 0.000"
[1579490358] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579490359] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579491036] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579492758] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(52) < 2020/01/20 09:22:17 [error] 18362#0: *41676700 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1579492759] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(9) < 203.212.70.170 [20/Jan/2020:09:25:06 +0530] TCP 502 0 0 123.009 "10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "3.008, 60.001, 60.000, 0.000"
[1579492836] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1070) < 2020/01/20 09:30:08 [error] 21675#0: *25588430 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:3001, upstream: "10.147.212.154:3001", bytes from/to client:91/0, bytes from/to upstream:0/91
[1579492878] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579492879] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579492956] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579492984] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1579493104] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1579493478] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(174) < 2020/01/20 09:40:50 [error] 18356#0: *41681853 upstream timed out (110: Connection timed out) while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579493480] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.0% (4202988 kB) free!
[1579493598] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579493600] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.5% (4561512 kB) free!
[1579493720] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.0% (4443684 kB) free!
[1579494079] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(224) < 182.18.144.246 [20/Jan/2020:09:44:31 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579494114] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(10) < 49.50.86.152 [20/Jan/2020:09:46:07 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1579494156] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(611) < 2020/01/20 09:47:02 [error] 21683#0: *25591334 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:150/0, bytes from/to upstream:0/150
[1579494198] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(47) < 2020/01/20 09:46:07 [error] 18370#0: *41684799 no live upstreams while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579494200] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579494235] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579494276] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579494318] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/01/20 09:55:14 [error] 18365#0: *41688192 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:223/17, bytes from/to upstream:17/223
[1579494438] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579495039] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/20 09:58:59 [error] 18360#0: *41689584 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:218/17, bytes from/to upstream:17/218
[1579495159] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579495564] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (17505416 kB) free!
[1579495684] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.4% (34829344 kB) free.
[1579496720] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (1219888 kB) free!
[1579498084] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15151744 kB) free!
[1579498204] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30207240 kB) free.
[1579499135] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.4% (22057816 kB) free.
[1579499676] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/20 11:19:03 [error] 21682#0: *25602054 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:2112/17, bytes from/to upstream:17/2112
[1579499796] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579500604] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10796040 kB) free!
[1579500724] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31080256 kB) free.
[1579501596] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(865) < 2020/01/20 11:55:57 [error] 21674#0: *25605647 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579501716] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579501924] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (8341720 kB) free!
[1579501999] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(863) < 103.255.146.164 [20/Jan/2020:11:55:57 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1579502044] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28548420 kB) free.
[1579502119] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579502316] SERVICE FLAPPING ALERT: primary-server;Nginx Error Log Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1579502959] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/20 12:11:33 [error] 18370#0: *41738716 open() "/usr/share/nginx/html/js/tmpl/404.ejs" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /js/tmpl/404.ejs?0.02862410669631621 HTTP/1.1", host: "qservercluster1.com:8091", referrer: "http://qservercluster1.com:8091/"
[1579503080] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579503516] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(14) < 2020/01/20 12:23:24 [error] 21675#0: *25606976 connect() failed (111: Connection refused) while connecting to upstream, client: 185.255.8.58, server: 0.0.0.0:13001, upstream: "10.147.212.154:13001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579503637] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579503844] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (6153000 kB) free!
[1579503964] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27419856 kB) free.
[1579504203] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579504323] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579504443] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1579504628] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (165588 kB) free!
[1579504748] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (282044 kB) free!
[1579504868] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (285172 kB) free!
[1579504879] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/20 12:45:16 [error] 18368#0: *41755204 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:870/17, bytes from/to upstream:17/870
[1579504999] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579505043] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579505104] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1579505123] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.2% (14340604 kB) free.
[1579505225] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1579505739] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (2094460 kB) free!
[1579505764] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4451068 kB) free!
[1579505859] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (608608 kB) free!
[1579505884] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27311780 kB) free.
[1579505979] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169716 kB) free!
[1579506036] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/20 13:08:01 [error] 21679#0: *25609189 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:3066/123, bytes from/to upstream:123/3066
[1579506156] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579506209] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174748 kB) free!
[1579506329] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190488 kB) free!
[1579506388] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1579506388] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.6% (21859464 kB) free.
[1579506430] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.8% (20796048 kB) free.
[1579506484] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17381936 kB) free!
[1579506604] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.8% (20880832 kB) free!
[1579506641] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1579506724] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15290028 kB) free!
[1579506761] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1579507040] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181652 kB) free!
[1579507160] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198992 kB) free!
[1579507280] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190580 kB) free!
[1579507324] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34913080 kB) free.
[1579507357] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/20 13:32:15 [error] 21683#0: *25610313 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:6093/70, bytes from/to upstream:70/6093
[1579507443] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579507477] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579507490] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180952 kB) free!
[1579507563] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579507588] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169688 kB) free!
[1579507600] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.1% (3228164 kB) free!
[1579507610] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198060 kB) free!
[1579507683] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1579507708] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178980 kB) free!
[1579507720] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (1758088 kB) free!
[1579507730] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191880 kB) free!
[1579507828] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173632 kB) free!
[1579507840] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.2% (302324 kB) free!
[1579508077] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/20 13:35:04 [error] 21681#0: *25610429 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:4015/87, bytes from/to upstream:87/4015
[1579508155] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 339 processes
[1579508197] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579508275] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1579508303] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.3% (21892572 kB) free.
[1579508335] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.3% (22284496 kB) free.
[1579508360] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.8% (22505728 kB) free.
[1579508388] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.7% (22220268 kB) free.
[1579509537] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.6% (4595804 kB) free!
[1579509657] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (1675504 kB) free!
[1579509724] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (18446780 kB) free!
[1579509778] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188620 kB) free!
[1579509844] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40865236 kB) free.
[1579510075] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1579510163] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182356 kB) free!
[1579510195] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1579510283] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190536 kB) free!
[1579510315] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1579510403] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176980 kB) free!
[1579510519] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(596) < 103.255.146.164 [20/Jan/2020:14:22:06 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1579510597] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(607) < 2020/01/20 14:22:06 [error] 21680#0: *25613741 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579510639] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579510717] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579511044] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18073188 kB) free!
[1579511164] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40291780 kB) free.
[1579511239] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(540) < 103.255.146.164 [20/Jan/2020:14:30:45 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1579511317] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(542) < 2020/01/20 14:30:45 [error] 21677#0: *25615195 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579511359] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579511399] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164524 kB) free!
[1579511437] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579511519] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181656 kB) free!
[1579511640] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167392 kB) free!
[1579512364] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.5% (17732504 kB) free!
[1579512484] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41687380 kB) free.
[1579513084] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14695364 kB) free!
[1579513140] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.2% (22250696 kB) free.
[1579513175] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1579513175] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.7% (22486484 kB) free.
[1579513200] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1579513200] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.0% (22060396 kB) free.
[1579513204] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39032876 kB) free.
[1579514348] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.7% (3873924 kB) free!
[1579514400] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195780 kB) free!
[1579514404] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13711380 kB) free!
[1579514468] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196196 kB) free!
[1579514520] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193348 kB) free!
[1579514524] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36864216 kB) free.
[1579514588] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190076 kB) free!
[1579514640] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167452 kB) free!
[1579514977] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163836 kB) free!
[1579515097] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193864 kB) free!
[1579515217] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178092 kB) free!
[1579515461] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.9% (22185468 kB) free.
[1579515483] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579515493] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.7% (22477416 kB) free.
[1579515536] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.7% (21737108 kB) free.
[1579515724] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12902332 kB) free!
[1579515844] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (37498680 kB) free.
[1579516083] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579516145] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.2% (4216668 kB) free!
[1579516203] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579516265] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180176 kB) free!
[1579516323] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (24.4% change >= 20.0% threshold)
[1579516323] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1579516385] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168648 kB) free!
[1579516444] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (1328412 kB) free!
[1579516564] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.5% (33677636 kB) free.
[1579516668] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.6% (4587804 kB) free!
[1579516702] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173840 kB) free!
[1579516788] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.3% (3038740 kB) free!
[1579516822] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189696 kB) free!
[1579516837] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/20 16:06:04 [error] 21680#0: *25619708 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:7724/70, bytes from/to upstream:70/7724
[1579516909] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171164 kB) free!
[1579516923] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579516942] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187276 kB) free!
[1579516957] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579517164] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (452908 kB) free!
[1579517284] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35252704 kB) free.
[1579517884] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (562132 kB) free!
[1579518004] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34082840 kB) free.
[1579518078] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.9% (21949756 kB) free.
[1579518113] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.3% (21883972 kB) free.
[1579518123] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579518142] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.0% (22315584 kB) free.
[1579518243] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579518604] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (2559496 kB) free!
[1579518724] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33277896 kB) free.
[1579518749] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.8% (3370904 kB) free!
[1579518757] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/20 16:41:41 [error] 21676#0: *25621355 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:26613/70, bytes from/to upstream:70/26613
[1579518869] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197668 kB) free!
[1579518877] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579518989] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197704 kB) free!
[1579519317] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190084 kB) free!
[1579519324] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19056404 kB) free!
[1579519361] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1579519437] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176444 kB) free!
[1579519444] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30678480 kB) free.
[1579519481] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 329 processes
[1579519557] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174080 kB) free!
[1579519602] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1579519889] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180744 kB) free!
[1579519986] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191264 kB) free!
[1579520098] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.3% (22041712 kB) free.
[1579520184] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.9% (21795852 kB) free.
[1579520215] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.0% (21809192 kB) free.
[1579520224] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1579520344] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1579520576] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9348408 kB) free!
[1579520643] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579520644] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4473784 kB) free!
[1579520696] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.0% (9342692 kB) free!
[1579520763] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579520764] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.8% (43192016 kB) free.
[1579520788] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.5% (2812072 kB) free!
[1579520816] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (9667368 kB) free!
[1579520883] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579520908] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198780 kB) free!
[1579521028] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165840 kB) free!
[1579521290] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.6% (22210108 kB) free.
[1579521302] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.0% (3454576 kB) free!
[1579521365] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16272820 kB) free!
[1579521415] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187420 kB) free!
[1579521422] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.1% (525376 kB) free!
[1579521485] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40519968 kB) free.
[1579521535] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186992 kB) free!
[1579521543] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167596 kB) free!
[1579521655] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166020 kB) free!
[1579522497] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170064 kB) free!
[1579522617] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170428 kB) free!
[1579522634] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.1% (21977700 kB) free.
[1579522642] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177252 kB) free!
[1579522667] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.7% (22240860 kB) free.
[1579522683] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1579522685] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16298960 kB) free!
[1579522692] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.6% (21721100 kB) free.
[1579522803] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579522805] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.2% (42397128 kB) free.
[1579522818] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9660240 kB) free!
[1579522924] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1579522938] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9434936 kB) free!
[1579523046] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.5% (4143756 kB) free!
[1579523058] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.7% (9202704 kB) free!
[1579523166] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.5% (3937880 kB) free!
[1579523286] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.9% (3794240 kB) free!
[1579523405] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13697188 kB) free!
[1579523525] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39081972 kB) free.
[1579523839] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.3% (4034224 kB) free!
[1579523867] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195580 kB) free!
[1579523886] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 75.1% (16853024 kB) free.
[1579523898] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174104 kB) free!
[1579523959] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.2% (2758708 kB) free!
[1579523988] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183016 kB) free!
[1579524018] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184844 kB) free!
[1579524041] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.5% (22179440 kB) free.
[1579524080] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165864 kB) free!
[1579524101] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.9% (21780728 kB) free.
[1579524166] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.2% (21772504 kB) free.
[1579524724] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579524725] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13247664 kB) free!
[1579524845] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39833260 kB) free.
[1579525241] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192824 kB) free!
[1579525310] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180068 kB) free!
[1579525361] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193512 kB) free!
[1579525369] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.0% (3447996 kB) free!
[1579525430] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188476 kB) free!
[1579525445] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11615420 kB) free!
[1579525481] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188988 kB) free!
[1579525489] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (2093268 kB) free!
[1579525550] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173392 kB) free!
[1579525565] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37646064 kB) free.
[1579525609] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (855228 kB) free!
[1579525616] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4210272 kB) free!
[1579526765] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16790276 kB) free!
[1579526885] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38599136 kB) free.
[1579527259] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4880028 kB) free!
[1579527485] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12778604 kB) free!
[1579527606] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.2% (35817912 kB) free.
[1579528805] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10798196 kB) free!
[1579528925] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (36000816 kB) free.
[1579530125] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12873464 kB) free!
[1579530245] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36342652 kB) free.
[1579530845] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (9981740 kB) free!
[1579530877] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(698) < 2020/01/20 20:03:45 [error] 21682#0: *25632503 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579530965] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (34131960 kB) free.
[1579530998] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579531159] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(686) < 103.255.146.164 [20/Jan/2020:20:03:45 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1579531279] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579532165] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10623564 kB) free!
[1579532285] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.5% (34987416 kB) free.
[1579532885] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.6% (12605416 kB) free!
[1579533005] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32950428 kB) free.
[1579533350] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1579534805] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14168468 kB) free!
[1579534924] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1579534925] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33869204 kB) free.
[1579535081] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1579535810] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1579536125] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14045344 kB) free!
[1579536245] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33070684 kB) free.
[1579537445] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.8% (14276144 kB) free!
[1579537565] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33049472 kB) free.
[1579539365] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11206500 kB) free!
[1579539485] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31397908 kB) free.
[1579541285] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (12526484 kB) free!
[1579541406] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31888684 kB) free.