[1544848795] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544848915] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1544855907] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 328 processes
[1544856507] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 328 processes
[1544856715] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544856835] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544856955] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544857555] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 288 processes
[1544863555] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544863676] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1544869675] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1544869795] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1544869915] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1544870515] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1544871115] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1544871235] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1544874562] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (1604208 kB) free!
[1544874846] SERVICE ALERT: sig6;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1544874849] HOST ALERT: sig6;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.58)
[1544874966] SERVICE ALERT: sig6;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38040 MB (74% inode=98%):
[1544874970] HOST ALERT: sig6;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1544875660] SERVICE ALERT: sig1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 299 processes
[1544875780] SERVICE ALERT: sig1;CPU_Procs;OK;HARD;1;CPU OK: 297 processes
[1544876372] SERVICE ALERT: sig6;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1544876402] HOST ALERT: sig6;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1544876428] SERVICE ALERT: sig6;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig6' was out of bounds)
[1544876431] SERVICE ALERT: sig6;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig6' was out of bounds)
[1544876466] HOST ALERT: sig6;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1544876492] SERVICE ALERT: sig6;Total Processes;OK;HARD;1;PROCS OK: 296 processes
[1544876548] SERVICE ALERT: sig6;CPU_Procs;OK;HARD;1;CPU OK: 305 processes
[1544876551] SERVICE ALERT: sig6;Memory;OK;HARD;1;OK - 94.8% (15400476 kB) free.
[1544876862] SERVICE ALERT: sig6;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1544876865] HOST ALERT: sig6;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.58)
[1544876906] SERVICE ALERT: sig6;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig6' was out of bounds)
[1544876929] HOST FLAPPING ALERT: sig6;STARTED; Host appears to have started flapping (21.2% change > 20.0% threshold)
[1544876929] HOST ALERT: sig6;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1544876982] SERVICE ALERT: sig6;CPU Load;OK;HARD;1;OK - load average: 0.19, 0.10, 0.04
[1544877026] SERVICE ALERT: sig6;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 2 processes with command name 'ntpd'
[1544877366] SERVICE ALERT: sig6;HDD_Root status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Root status' on host 'sig6' was out of bounds)
[1544877486] SERVICE ALERT: sig6;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38040 MB (74% inode=98%):
[1544877749] HOST FLAPPING ALERT: sig6;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1544885035] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544885155] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1544885275] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1544892396] HOST ALERT: sig6;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1544892550] HOST ALERT: sig6;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.41 ms