[1570127883] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (910652 kB) free!
[1570127967] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190572 kB) free!
[1570127995] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186448 kB) free!
[1570128003] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (866128 kB) free!
[1570128087] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174480 kB) free!
[1570128115] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191932 kB) free!
[1570128123] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (910224 kB) free!
[1570128207] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195536 kB) free!
[1570128235] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189344 kB) free!
[1570129241] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 501 processes
[1570129301] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 501 processes
[1570129362] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 0 warn out of 503 processes
[1570129422] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 498 processes
[1570131765] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 42.7% (10526464 kB) free.
[1570132078] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 449 processes
[1570132365] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (1378816 kB) free!
[1570132485] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.8% (455372 kB) free!
[1570132605] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200232 kB) free!
[1570136176] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (240312 kB) free!
[1570136296] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (226272 kB) free!
[1570136416] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233904 kB) free!
[1570137016] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 85.3% (42050484 kB) free.
[1570138923] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.4% (21419940 kB) free.
[1570139607] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 84.9% (20810648 kB) free.
[1570139636] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 83.1% (20377248 kB) free.
[1570141014] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 18.71, 11.52, 4.97
[1570141134] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.23, 7.95, 4.46
[1570144005] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1570144678] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 593 processes
[1570144798] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 593 processes
[1570144918] SERVICE ALERT: smscapp2;Total Processes;WARNING;HARD;3;PROCS WARNING: 594 processes
[1570146378] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 32.81, 8.13, 3.02
[1570146498] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 87.16, 37.09, 14.26
[1570146618] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 94.71, 56.52, 24.14
[1570147818] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.38, 3.10, 13.22
[1570149018] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.35, 0.50, 3.99
[1570149119] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1570149719] SERVICE ALERT: smscapp2;Total Processes;WARNING;HARD;3;PROCS WARNING: 599 processes
[1570150606] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 34.8% (8580748 kB) free.
[1570151806] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (180412 kB) free!
[1570151922] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 506 processes
[1570151926] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179756 kB) free!
[1570151982] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 498 processes
[1570152046] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191284 kB) free!
[1570152282] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 510 processes
[1570152342] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 508 processes
[1570153242] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 504 processes
[1570153302] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 495 processes
[1570153452] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1570153573] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1570153815] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.22, 7.85, 5.85
[1570153935] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.22, 8.31, 6.26
[1570154055] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.07, 8.56, 6.60
[1570162776] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 72.3% (17850564 kB) free.
[1570162808] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 59.3% (29241932 kB) free.
[1570163384] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183492 kB) free!
[1570163505] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196200 kB) free!
[1570163625] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177980 kB) free!
[1570164476] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.03, 6.29, 5.37
[1570164596] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.23, 5.61, 5.24
[1570164716] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.59, 5.24, 5.14
[1570166357] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9702240 kB) free!
[1570166477] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.4% (9539468 kB) free!
[1570166597] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (9357788 kB) free!
[1570166607] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171392 kB) free!
[1570166727] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198728 kB) free!
[1570166847] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183164 kB) free!
[1570167119] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1570167123] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189660 kB) free!
[1570167236] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178616 kB) free!
[1570167243] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180124 kB) free!
[1570167293] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 81.1% (20003048 kB) free.
[1570167333] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.9% (22029348 kB) free.
[1570167357] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.3% (3015668 kB) free!
[1570167362] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.3% (22133864 kB) free.
[1570167363] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 79.3% (19428448 kB) free.
[1570167719] SERVICE ALERT: smscapp2;Total Processes;WARNING;HARD;3;PROCS WARNING: 596 processes
[1570167896] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178124 kB) free!
[1570168016] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186104 kB) free!
[1570168074] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 18 users currently logged in
[1570168137] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1570168137] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193932 kB) free!
[1570168194] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 18 users currently logged in
[1570168314] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 18 users currently logged in
[1570168319] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 650 processes
[1570169597] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4342420 kB) free!
[1570170114] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 22 users currently logged in
[1570171314] SERVICE ALERT: pservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 18 users currently logged in
[1570172597] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 53.6% (26405912 kB) free.
[1570173336] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.6% (4812660 kB) free!
[1570173457] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (4618484 kB) free!
[1570173466] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24125132 kB) free!
[1570173577] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.9% (4398740 kB) free!
[1570173586] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 37.2% (49065444 kB) free.
[1570174186] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23247040 kB) free!
[1570174306] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.4% (48035824 kB) free.
[1570174397] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236168 kB) free!
[1570174517] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (245980 kB) free!
[1570174637] SERVICE FLAPPING ALERT: qservercluster2;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1570174637] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (239640 kB) free!
[1570174777] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1570174777] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2090404 kB) free!
[1570174906] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (22795732 kB) free!
[1570174914] SERVICE FLAPPING ALERT: pservercluster1;Current Users;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1570174914] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 13 users currently logged in
[1570175027] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.7% (47041212 kB) free.
[1570175167] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.1% (4424736 kB) free!
[1570175237] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 59.9% (29501944 kB) free.
[1570175288] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.1% (4184428 kB) free!
[1570175408] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.2% (3964664 kB) free!
[1570175626] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22451916 kB) free!
[1570175746] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.9% (45954216 kB) free.
[1570176013] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9092124 kB) free!
[1570176133] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.0% (8355060 kB) free!
[1570176253] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.0% (8354360 kB) free!
[1570176346] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (21822368 kB) free!
[1570176466] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.0% (44859588 kB) free.
[1570176608] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1856060 kB) free!
[1570177066] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20043104 kB) free!
[1570177186] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 42.8% (56410716 kB) free.
[1570177216] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8785988 kB) free!
[1570177336] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.2% (8474176 kB) free!
[1570177455] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.23, 3.29, 4.47
[1570177456] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.6% (8158080 kB) free!
[1570178606] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.5% (4780900 kB) free!
[1570178726] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.4% (4753108 kB) free!
[1570178846] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.3% (4719108 kB) free!
[1570178937] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1570179256] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (3552152 kB) free!
[1570179457] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 13.7% (3218648 kB) free!
[1570179578] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 11.7% (2761216 kB) free!
[1570179698] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2295332 kB) free!
[1570179853] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (3562748 kB) free!
[1570180176] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.8% (2897348 kB) free!
[1570180297] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.8% (2896300 kB) free!
[1570180417] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 17.8% (2898388 kB) free!
[1570181035] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 16.4% (5713568 kB) free!
[1570181138] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;1;WARNING - 14.5% (5086660 kB) free!
[1570181155] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 15.4% (5364284 kB) free!
[1570181258] SERVICE ALERT: sig5-vm2;Memory;WARNING;SOFT;2;WARNING - 13.4% (4693476 kB) free!
[1570181275] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 14.3% (4968368 kB) free!
[1570181378] SERVICE ALERT: sig5-vm2;Memory;WARNING;HARD;3;WARNING - 12.2% (4262172 kB) free!
[1570181498] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 96.8% (22810924 kB) free.
[1570181522] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 128.06, 76.22, 33.15
[1570181642] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 129.79, 94.06, 44.91
[1570181762] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 129.76, 105.84, 55.18
[1570181867] SERVICE ALERT: sig3-vm2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 335 processes
[1570181875] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 93.5% (32568920 kB) free.
[1570181978] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 93.1% (32614452 kB) free.
[1570181987] SERVICE ALERT: sig3-vm2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 335 processes
[1570182107] SERVICE ALERT: sig3-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 335 processes
[1570183314] SERVICE FLAPPING ALERT: pservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1570183562] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.46, 2.55, 11.95
[1570184386] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23142056 kB) free!
[1570184506] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.7% (45802652 kB) free.
[1570185106] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22905268 kB) free!
[1570185226] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.7% (45696724 kB) free.
[1570185362] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.60, 1.83, 3.50
[1570185577] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1570185826] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22982500 kB) free!
[1570185946] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44552660 kB) free.
[1570186546] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.3% (22854104 kB) free!
[1570186637] SERVICE FLAPPING ALERT: qservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1570186666] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.4% (18970728 kB) free!
[1570186786] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41077416 kB) free.
[1570187386] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.5% (19122252 kB) free!
[1570187506] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41068548 kB) free.
[1570188106] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19478968 kB) free!
[1570188226] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40767520 kB) free.
[1570188437] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9516608 kB) free!
[1570188446] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2363016 kB) free!
[1570188557] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.0% (9387012 kB) free!
[1570188678] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (9211332 kB) free!
[1570188827] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (19041708 kB) free!
[1570188947] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40884428 kB) free.
[1570189418] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (159768 kB) free!
[1570189547] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (19382824 kB) free!
[1570189547] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3209900 kB) free!
[1570189667] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40263568 kB) free.
[1570189667] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 20.0% (3247464 kB) free!
[1570189787] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3579828 kB) free.
[1570190027] SERVICE ALERT: sig6-vm1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1570190029] SERVICE ALERT: sig6-vm1;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'sig6-vm1' was out of bounds)
[1570190029] HOST ALERT: sig6-vm1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.58)
[1570190037] SERVICE ALERT: sig6-vm1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig6-vm1' was out of bounds)
[1570190037] SERVICE ALERT: sig6-vm1;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig6-vm1' was out of bounds)
[1570190037] SERVICE ALERT: sig6-vm1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig6-vm1' was out of bounds)
[1570190037] SERVICE ALERT: sig6-vm1;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig6-vm1' was out of bounds)
[1570190037] SERVICE ALERT: sig6-vm1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig6-vm1' was out of bounds)
[1570190150] SERVICE ALERT: sig6-vm1;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig6-vm1' was out of bounds)
[1570190150] SERVICE ALERT: sig6-vm1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig6-vm1' was out of bounds)
[1570190158] SERVICE ALERT: sig6-vm2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1570190161] HOST ALERT: sig6-vm2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.193)
[1570190267] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19546088 kB) free!
[1570190281] SERVICE ALERT: sig6-vm2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig6-vm2' was out of bounds)
[1570190387] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39679048 kB) free.
[1570190615] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig6-vm2' was out of bounds)
[1570190618] SERVICE ALERT: sig6-vm2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig6-vm2' was out of bounds)
[1570190618] SERVICE ALERT: sig6-vm2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig6-vm2' was out of bounds)
[1570190637] SERVICE ALERT: sig6-vm2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig6-vm2' was out of bounds)
[1570190637] SERVICE ALERT: sig6-vm2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig6-vm2' was out of bounds)
[1570190637] SERVICE ALERT: sig6-vm2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig6-vm2' was out of bounds)
[1570190728] HOST ALERT: sig6-vm2;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.193)
[1570190836] HOST ALERT: sig6-vm1;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.58)
[1570190987] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18895844 kB) free!
[1570191107] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.8% (24828908 kB) free!
[1570191227] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.4% (33525596 kB) free.
[1570191595] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 3 warn out of 333 processes
[1570191698] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 19.1% (4504300 kB) free!
[1570191715] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1570191818] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4341496 kB) free!
[1570191938] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 17.6% (4137268 kB) free!
[1570192302] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 534 processes
[1570192362] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 535 processes
[1570192422] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;3;CPU CRITICAL: 1 crit, 0 warn out of 543 processes
[1570192482] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 538 processes
[1570192538] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 97.7% (23028084 kB) free.
[1570192608] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 93.5% (22911108 kB) free.
[1570192652] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.5% (23066036 kB) free.
[1570192691] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.6% (23192272 kB) free.
[1570192807] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.3% (39569900 kB) free.
[1570192856] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 83.3% (41066336 kB) free.
[1570192878] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (4715540 kB) free!
[1570192912] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.7% (22971300 kB) free.
[1570192963] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 92.8% (15086656 kB) free.
[1570193037] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 95.8% (47212764 kB) free.
[1570193052] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 330 processes
[1570193387] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (3119736 kB) free!
[1570193507] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3376724 kB) free.
[1570193655] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.00, 9.29, 5.75
[1570193775] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.54, 6.65, 5.21
[1570193855] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.1% (3237088 kB) free!
[1570193895] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.91, 4.80, 4.72
[1570193934] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.27, 15.31, 7.57
[1570193975] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.8% (3159448 kB) free!
[1570194054] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.06, 10.53, 6.75
[1570194095] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1828612 kB) free!
[1570194115] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 3 warn out of 340 processes
[1570194174] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.22, 7.38, 6.06
[1570194235] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 4 crit, 3 warn out of 340 processes
[1570194253] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 326 processes
[1570194355] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 1 warn out of 334 processes
[1570194363] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 64.46, 22.22, 8.76
[1570194483] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.72, 15.26, 7.84
[1570194603] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.42, 10.63, 7.04
[1570194774] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.19, 1.73, 3.55
[1570194955] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1570195075] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 773 processes
[1570195095] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.92, 5.55, 5.38
[1570195195] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 773 processes
[1570195215] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.89, 3.98, 4.82
[1570195315] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 340 processes
[1570195814] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (221220 kB) free!
[1570195815] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.25, 5.19, 6.20
[1570195907] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (3050588 kB) free!
[1570195934] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (230348 kB) free!
[1570195935] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.12, 4.18, 5.70
[1570196027] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (3162584 kB) free!
[1570196054] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (223732 kB) free!
[1570196055] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.13, 3.56, 5.30
[1570196147] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.6% (3346204 kB) free.
[1570196155] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 333 processes
[1570196163] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (4716056 kB) free!
[1570196275] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1570196284] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.1% (4686408 kB) free!
[1570196404] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (4615352 kB) free!
[1570196516] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.82, 3.66, 4.34
[1570196655] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 5.54, 4.81, 4.87
[1570196875] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1570196995] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1570197003] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.41, 1.73, 4.23
[1570197167] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.6% (2854276 kB) free!
[1570197255] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.81, 5.12, 5.12
[1570197287] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.6% (2858120 kB) free!
[1570197375] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.35, 4.15, 4.77
[1570197407] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 17.6% (2855144 kB) free!
[1570200932] HOST ALERT: sig6-vm2;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.79 ms
[1570200975] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.63, 26.54, 21.64
[1570201095] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.64, 17.97, 19.09
[1570201215] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.96, 12.34, 16.89
[1570201663] SERVICE ALERT: qservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1570201784] SERVICE ALERT: qservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1570201904] SERVICE ALERT: qservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1570203015] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1570203015] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.13, 0.96, 3.17
[1570203104] SERVICE ALERT: qservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1570203704] SERVICE ALERT: qservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1570203824] SERVICE ALERT: qservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1570203944] SERVICE ALERT: qservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1570205403] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.11, 6.63, 5.94
[1570205524] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 20.90, 19.04, 11.04
[1570205644] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 3.78, 13.10, 9.84
[1570206774] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.70, 7.20, 6.93
[1570206844] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 34.56, 17.17, 9.35
[1570206894] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.60, 5.01, 6.16
[1570207014] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.44, 3.49, 5.47
[1570207435] SERVICE ALERT: rq-2;CPU_Procs;WARNING;HARD;1;CPU WARNING: 1 warn out of 239 processes
[1570207444] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1570207444] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.37, 3.34, 5.44
[1570207546] SERVICE ALERT: rq-2;CPU Load;WARNING;HARD;1;WARNING - load average: 7.39, 7.28, 6.09
[1570207614] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.42, 0.91, 3.14
[1570208044] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.01, 1.32, 3.32
[1570208755] SERVICE ALERT: rq-2;CPU_Procs;CRITICAL;HARD;1;CPU CRITICAL: 1 crit, 0 warn out of 238 processes
[1570209244] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 61.31, 23.91, 12.53
[1570209364] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.92, 16.19, 11.08
[1570209484] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.06, 11.13, 9.84
[1570209618] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.47, 9.50, 6.64
[1570209738] SERVICE ALERT: sservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.37, 6.69, 5.95
[1570209858] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 29.72, 17.19, 9.89
[1570210084] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 80.82, 31.07, 18.50
[1570210108] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 544 processes
[1570210614] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.09, 11.17, 6.13
[1570210684] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.70, 4.97, 10.15
[1570210735] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.08, 7.65, 5.45
[1570210855] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.28, 5.45, 4.91
[1570211455] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.85, 4.78, 6.46
[1570211575] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.21, 3.57, 5.81
[1570211658] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.64, 1.40, 3.81
[1570211695] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 68.90, 17.68, 10.12
[1570212295] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1570212295] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.62, 2.94, 5.62
[1570213204] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2411180 kB) free!
[1570213215] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1570213458] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 142.84, 92.39, 41.41
[1570213578] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 144.03, 109.48, 53.85
[1570213698] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 23.58, 76.76, 48.77
[1570213708] SERVICE ALERT: smscapp1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 551 processes