[1568572716] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (852404 kB) free!
[1568572754] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185340 kB) free!
[1568572836] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (839252 kB) free!
[1568572874] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196364 kB) free!
[1568572956] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (825204 kB) free!
[1568572994] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184708 kB) free!
[1568573190] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1568573286] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.7% (3124604 kB) free!
[1568573407] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184300 kB) free!
[1568573527] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187036 kB) free!
[1568573978] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 297 processes
[1568574098] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1568577408] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (3553140 kB) free!
[1568577597] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.76, 7.11, 5.33
[1568577717] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.39, 7.35, 5.62
[1568577837] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.61, 7.30, 5.83
[1568579748] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 226 processes
[1568580127] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.0% (3467036 kB) free!
[1568580237] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.43, 3.25, 4.08
[1568580728] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191640 kB) free!
[1568584356] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 46.4% (11380216 kB) free.
[1568584394] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 61.6% (15088516 kB) free.
[1568596622] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 447 processes
[1568596682] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 442 processes
[1568597528] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 44.9% (11077660 kB) free.
[1568598355] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1568598475] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1568598728] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198000 kB) free!
[1568598757] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171940 kB) free!
[1568598795] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170484 kB) free!
[1568598848] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187060 kB) free!
[1568598877] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181236 kB) free!
[1568598915] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197696 kB) free!
[1568598968] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202144 kB) free!
[1568598997] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176244 kB) free!
[1568599035] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (163432 kB) free!
[1568608075] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1568608195] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1568610032] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.9% (22936872 kB) free.
[1568610129] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.3% (22629684 kB) free.
[1568610244] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.3% (40576392 kB) free.
[1568610344] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.6% (23181428 kB) free.
[1568610592] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1568610712] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1568611423] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 26644 MB (20% inode=99%):
[1568611543] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 26644 MB (20% inode=99%):
[1568611663] SERVICE ALERT: sig4-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 26580 MB (20% inode=99%):
[1568613046] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181796 kB) free!
[1568613166] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173092 kB) free!
[1568613286] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178756 kB) free!
[1568613723] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.8% (4850220 kB) free!
[1568613843] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.7% (4832496 kB) free!
[1568613963] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.7% (4819532 kB) free!
[1568614687] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.0% (22941268 kB) free.
[1568615386] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.36, 6.80, 5.96
[1568615506] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.80, 7.02, 6.15
[1568615626] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 9.34, 7.83, 6.55
[1568620265] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.1% (8907184 kB) free!
[1568620386] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.2% (8498712 kB) free!
[1568620506] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.4% (8102364 kB) free!
[1568620696] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.0% (4453184 kB) free!
[1568620816] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.1% (4222112 kB) free!
[1568620936] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.6% (3594548 kB) free!
[1568621706] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4830244 kB) free!
[1568623336] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (855556 kB) free!
[1568624533] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4718772 kB) free!
[1568624653] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (4547808 kB) free!
[1568624746] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 14.7% (3475744 kB) free!
[1568624773] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.8% (4371804 kB) free!
[1568624866] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 12.5% (2940684 kB) free!
[1568624987] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 10.5% (2472476 kB) free!
[1568625493] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 96.2% (47399428 kB) free.
[1568625587] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (385556 kB) free!
[1568625703] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 333 processes
[1568625823] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1568625895] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.4% (22552948 kB) free.
[1568626009] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 97.8% (23054180 kB) free.
[1568626573] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1722732 kB) free!
[1568627105] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178340 kB) free!
[1568627153] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (4896832 kB) free!
[1568627225] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179564 kB) free!
[1568627268] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.7% (22237884 kB) free.
[1568627273] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 21.6% (5291740 kB) free.
[1568627335] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.2% (21023732 kB) free.
[1568627626] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.59, 4.30, 4.98
[1568627873] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4821004 kB) free!
[1568627947] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171920 kB) free!
[1568627993] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (4736576 kB) free!
[1568628067] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166820 kB) free!
[1568628113] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.9% (4629920 kB) free!
[1568628124] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.1% (22715788 kB) free.
[1568628262] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.9% (22763308 kB) free.
[1568629678] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (4524852 kB) free!
[1568629798] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.9% (4387064 kB) free!
[1568629918] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1568629918] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.8% (4112112 kB) free!
[1568630016] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (413784 kB) free!
[1568630136] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (3225796 kB) free!
[1568630147] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (223388 kB) free!
[1568630255] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 85.4% (112573080 kB) free.
[1568630267] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 66.6% (32824256 kB) free.
[1568630535] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.5% (3083144 kB) free!
[1568630655] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (1824048 kB) free!
[1568630666] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (4774944 kB) free!
[1568630775] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (984692 kB) free!
[1568630787] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.0% (4668228 kB) free!
[1568630867] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (316292 kB) free!
[1568630907] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.7% (4592716 kB) free!
[1568630987] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (219568 kB) free!
[1568631107] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 91.1% (44914732 kB) free.
[1568631253] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (216416 kB) free!
[1568631373] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (259452 kB) free!
[1568631493] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (239068 kB) free!
[1568631719] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2195452 kB) free!
[1568631833] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 74.4% (18364440 kB) free.
[1568631984] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.6% (21965212 kB) free.
[1568632107] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1568632107] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1982612 kB) free!
[1568632693] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 34.8% (17158612 kB) free.
[1568632707] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.1% (3457700 kB) free!
[1568633026] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.16, 5.76, 5.07
[1568633146] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.73, 5.10, 4.91
[1568633293] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (4080876 kB) free!
[1568633413] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (1351228 kB) free!
[1568633534] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (214392 kB) free!
[1568633633] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176044 kB) free!
[1568633753] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193708 kB) free!
[1568633873] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1568633873] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195076 kB) free!
[1568634082] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.6% (22356272 kB) free.
[1568634134] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1568634134] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 31.8% (15656936 kB) free.
[1568634507] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2351640 kB) free!
[1568634595] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1568634686] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173740 kB) free!
[1568634715] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1568634807] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 73.1% (18046364 kB) free.
[1568635412] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194196 kB) free!
[1568635532] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184420 kB) free!
[1568635652] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189648 kB) free!
[1568636314] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.0% (23184100 kB) free.
[1568637223] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1568637343] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1568637398] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170260 kB) free!
[1568637518] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197404 kB) free!
[1568637614] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.2% (21863788 kB) free.
[1568637734] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9813400 kB) free!
[1568637854] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9213148 kB) free!
[1568637974] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.0% (9386972 kB) free!
[1568638093] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (433816 kB) free!
[1568638124] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.9% (2677780 kB) free!
[1568638213] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (360912 kB) free!
[1568638244] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.8% (2665956 kB) free!
[1568638333] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (250800 kB) free!
[1568638364] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.6% (2618640 kB) free!
[1568638964] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (170476 kB) free!
[1568640374] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (247684 kB) free!
[1568640625] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1568642256] SERVICE ALERT: smscapp1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 552 processes
[1568642377] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 550 processes
[1568643390] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 62.22, 28.62, 11.18
[1568643510] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.85, 19.36, 9.90
[1568643630] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.90, 13.23, 8.80
[1568643959] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1568644080] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1568644707] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1568644830] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.95, 1.13, 3.14
[1568647176] SERVICE ALERT: smscapp1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 558 processes
[1568647296] SERVICE ALERT: smscapp1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 559 processes
[1568647416] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 560 processes
[1568648882] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 490 processes
[1568648942] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 489 processes
[1568649974] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1568650364] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)