[1573065272] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573065439] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179024 kB) free!
[1573065829] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183240 kB) free!
[1573065859] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174004 kB) free!
[1573065950] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177392 kB) free!
[1573066070] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182356 kB) free!
[1573067110] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184880 kB) free!
[1573067170] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 57.3% (14040524 kB) free.
[1573067178] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24839924 kB) free!
[1573067298] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.6% (37640104 kB) free.
[1573068848] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 21.6% (5329164 kB) free.
[1573069448] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167100 kB) free!
[1573069568] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171620 kB) free!
[1573069689] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198552 kB) free!
[1573069698] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (23011428 kB) free!
[1573069818] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.7% (41725400 kB) free.
[1573070418] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24781920 kB) free!
[1573070538] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42804796 kB) free.
[1573071138] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26358540 kB) free!
[1573071258] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.3% (22863092 kB) free!
[1573071378] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (41139632 kB) free.
[1573071978] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24229992 kB) free!
[1573072098] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42591652 kB) free.
[1573072699] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24566944 kB) free!
[1573072819] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42584980 kB) free.
[1573073419] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (25148928 kB) free!
[1573073539] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42886784 kB) free.
[1573073736] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (241144 kB) free!
[1573073856] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (229856 kB) free!
[1573073891] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (226040 kB) free!
[1573073976] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240272 kB) free!
[1573074011] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (228476 kB) free!
[1573074131] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 85.8% (42276976 kB) free.
[1573074139] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25888336 kB) free!
[1573074259] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42907648 kB) free.
[1573074859] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26225488 kB) free!
[1573074979] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (42935096 kB) free.
[1573075176] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.9% (45296196 kB) free.
[1573075579] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26068632 kB) free!
[1573075689] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 55.0% (13568644 kB) free.
[1573075699] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (42953868 kB) free.
[1573076299] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26035600 kB) free!
[1573076419] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (42965508 kB) free.
[1573076532] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1764) < 2019/11/07 03:11:22 [error] 24125#0: *3147617 no live upstreams while connecting to upstream, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "backend_3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573076652] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573076702] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1484) < 203.212.70.170 [07/Nov/2019:03:11:22 +0530] TCP 502 0 0 60.000 "10.147.212.155:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1573076822] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573076851] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(971) < 2019/11/07 03:11:21 [error] 15232#0: *1576109 upstream timed out (110: Connection timed out) while connecting to upstream, client: 121.241.242.109, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573076902] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(905) < 103.255.146.164 [07/Nov/2019:03:10:13 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1573076971] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573077019] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26094724 kB) free!
[1573077022] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573077139] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (42987720 kB) free.
[1573077739] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25884248 kB) free!
[1573077859] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (43009764 kB) free.
[1573078460] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (25641728 kB) free!
[1573078580] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (43030812 kB) free.
[1573079180] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25453952 kB) free!
[1573079300] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.7% (43048868 kB) free.
[1573079900] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (25172412 kB) free!
[1573080020] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.4% (42737592 kB) free.
[1573080620] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24968412 kB) free!
[1573080740] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42188436 kB) free.
[1573081340] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24746916 kB) free!
[1573081460] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41977496 kB) free.
[1573082060] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24432132 kB) free!
[1573082147] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 302 processes
[1573082180] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.9% (41989028 kB) free.
[1573082267] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1573082780] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24814344 kB) free!
[1573082900] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42161040 kB) free.
[1573083500] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25528160 kB) free!
[1573083621] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.7% (43132356 kB) free.
[1573084220] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26388416 kB) free!
[1573084340] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.7% (43140916 kB) free.
[1573085433] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1573085553] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 399 processes
[1573087089] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177820 kB) free!
[1573087209] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186792 kB) free!
[1573087329] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185128 kB) free!
[1573087340] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26036580 kB) free!
[1573087460] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (39038508 kB) free.
[1573087670] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.9% (21304508 kB) free.
[1573088059] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 39.1% (9594052 kB) free.
[1573088239] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 79.7% (19537612 kB) free.
[1573089467] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 295 processes
[1573089587] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1573091907] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.8% (22409176 kB) free.
[1573091947] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.9% (40381564 kB) free.
[1573091981] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 47.4% (23354988 kB) free.
[1573092260] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23421648 kB) free!
[1573092380] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.9% (19672140 kB) free!
[1573092500] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40502184 kB) free.
[1573093100] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20557152 kB) free!
[1573093220] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40907652 kB) free.
[1573093820] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22102896 kB) free!
[1573093940] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.8% (18169280 kB) free!
[1573094060] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39629736 kB) free.
[1573094660] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19889064 kB) free!
[1573094780] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39902612 kB) free.
[1573095380] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.1% (21260708 kB) free!
[1573095500] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.1% (17324960 kB) free!
[1573095517] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (4741880 kB) free!
[1573095620] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39290504 kB) free.
[1573095638] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.3% (4030236 kB) free!
[1573095758] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 13.3% (3289708 kB) free!
[1573096220] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19182852 kB) free!
[1573096340] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.6% (15347588 kB) free!
[1573096461] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37084508 kB) free.
[1573097060] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16882220 kB) free!
[1573097180] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (38093660 kB) free.
[1573097558] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1987568 kB) free!
[1573097780] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17641056 kB) free!
[1573097900] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.4% (13758124 kB) free!
[1573098020] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36774576 kB) free.
[1573098259] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (4780184 kB) free!
[1573098379] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4697456 kB) free!
[1573098499] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (4596432 kB) free!
[1573098620] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15615388 kB) free!
[1573098740] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29165772 kB) free.
[1573099039] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.1% (4187200 kB) free!
[1573099159] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.5% (4034304 kB) free!
[1573099279] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.7% (3848648 kB) free!
[1573099340] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16849664 kB) free!
[1573099460] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (12671340 kB) free!
[1573099580] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35657396 kB) free.
[1573100180] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (13179756 kB) free!
[1573100300] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (8864660 kB) free!
[1573100421] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34056288 kB) free.
[1573100479] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (1765804 kB) free!
[1573100899] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2299232 kB) free!
[1573101020] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10127248 kB) free!
[1573101140] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (5907816 kB) free!
[1573101252] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(345) < 2019/11/07 10:02:02 [error] 24130#0: *3354427 upstream timed out (110: Connection timed out) while connecting to upstream, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.155:3001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573101260] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31535916 kB) free.
[1573101372] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573101423] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(197) < 182.18.182.41 [07/Nov/2019:10:00:53 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1573101543] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573101860] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (7018256 kB) free!
[1573101980] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.0% (2676472 kB) free!
[1573102100] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29250952 kB) free.
[1573102392] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 387 processes
[1573102512] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 459 processes
[1573102572] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(38) < 2019/11/07 10:24:25 [error] 24132#0: *3364973 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573102692] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573102700] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5197340 kB) free!
[1573102820] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (1047728 kB) free!
[1573102940] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26785156 kB) free.
[1573102951] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 361.27, 296.94, 138.43
[1573103071] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 71.50, 207.71, 124.97
[1573103191] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 32.35, 148.02, 113.16
[1573103540] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (1390860 kB) free!
[1573103660] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (441396 kB) free!
[1573103780] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.5% (24391016 kB) free!
[1573103971] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/11/07 10:48:12 [error] 15230#0: *1577917 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: qserverc-2-1gbps.server, request: "GET /favicon.ico HTTP/1.1", host: "qserverc-2-1gbps.server:8091", referrer: "http://qserverc-2-1gbps.server:8091/"
[1573104091] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573104380] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (389300 kB) free!
[1573104680] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.8% (2636652 kB) free!
[1573104818] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.5% (4771540 kB) free!
[1573104938] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 19.4% (4748692 kB) free!
[1573104961] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 20.0% (4500272 kB) free!
[1573105058] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 19.2% (4697056 kB) free!
[1573105081] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.3% (4105352 kB) free!
[1573105199] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 2002 processes
[1573105201] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 16.5% (3698888 kB) free!
[1573105280] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1573105280] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.3% (1301576 kB) free!
[1573105319] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 2003 processes
[1573105440] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 2003 processes
[1573105580] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16179100 kB) free!
[1573105801] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1714568 kB) free!
[1573106180] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1573106180] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (425872 kB) free!
[1573106492] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(23) < 2019/11/07 11:27:32 [error] 15227#0: *1587924 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573106612] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573106919] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 17.8% (6201948 kB) free!
[1573107039] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 16.4% (5697224 kB) free!
[1573107159] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 14.9% (5193232 kB) free!
[1573107380] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.5% (24442508 kB) free!
[1573107391] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 27.27, 26.24, 35.37
[1573107576] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.0% (8368912 kB) free!
[1573107696] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.3% (8048720 kB) free!
[1573107759] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (2603400 kB) free!
[1573107816] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.7% (7736144 kB) free!
[1573107980] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (394888 kB) free!
[1573108143] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.1% (22492172 kB) free.
[1573108174] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 40.9% (10034184 kB) free.
[1573108218] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 50.8% (12462956 kB) free.
[1573108261] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 42.5% (55982108 kB) free.
[1573108297] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 92.7% (32285580 kB) free.
[1573108332] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 83.2% (18689024 kB) free.
[1573109016] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4529956 kB) free!
[1573109469] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22489540 kB) free!
[1573109589] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.3% (17591028 kB) free!
[1573109710] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (13071992 kB) free!
[1573109791] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 22.61, 25.09, 27.95
[1573110310] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20875232 kB) free!
[1573110631] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (4513592 kB) free!
[1573110751] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.7% (4343404 kB) free!
[1573110871] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.0% (4159212 kB) free!
[1573110910] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (34012832 kB) free.
[1573111471] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (183900 kB) free!
[1573111510] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7509024 kB) free!
[1573111630] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.3% (39939736 kB) free.
[1573111670] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.6% (3339604 kB) free!
[1573111751] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197932 kB) free!
[1573111785] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196400 kB) free!
[1573111790] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 20.6% (5051908 kB) free.
[1573111871] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184516 kB) free!
[1573111905] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166108 kB) free!
[1573111991] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164040 kB) free!
[1573112025] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194016 kB) free!
[1573112230] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.0% (14547668 kB) free!
[1573112350] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.6% (10007368 kB) free!
[1573112390] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.7% (4580696 kB) free!
[1573112470] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (5720320 kB) free!
[1573112511] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.4% (4514208 kB) free!
[1573112531] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9499260 kB) free!
[1573112631] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (4454576 kB) free!
[1573112651] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (9330324 kB) free!
[1573112771] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.6% (9143064 kB) free!
[1573112956] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.2% (8494632 kB) free!
[1573112981] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.8% (22161524 kB) free.
[1573113063] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.9% (47272136 kB) free.
[1573113070] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15140912 kB) free!
[1573113076] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.2% (8494008 kB) free!
[1573113145] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.2% (22096656 kB) free.
[1573113196] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.2% (8494632 kB) free!
[1573113207] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1573113207] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 50.1% (12278376 kB) free.
[1573113283] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 92.8% (22737404 kB) free.
[1573113326] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.7% (40263444 kB) free.
[1573113360] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 92.6% (22696868 kB) free.
[1573113449] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 95.0% (46807100 kB) free.
[1573113517] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 47.6% (62718212 kB) free.
[1573114190] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175476 kB) free!
[1573114310] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195740 kB) free!
[1573114419] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178924 kB) free!
[1573114430] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1573114430] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177420 kB) free!
[1573114512] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1573114539] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167944 kB) free!
[1573114632] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 361 processes
[1573114659] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183676 kB) free!
[1573114802] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 50.3% (12332648 kB) free.
[1573114847] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.9% (20939136 kB) free.
[1573115013] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(180) < 2019/11/07 13:47:29 [error] 23647#0: *1645488 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573115133] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573115458] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176476 kB) free!
[1573115549] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.2% (3976248 kB) free!
[1573115578] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177752 kB) free!
[1573115669] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.9% (3416724 kB) free!
[1573115698] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (178020 kB) free!
[1573115790] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 13.3% (3255440 kB) free!
[1573115928] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (13686992 kB) free!
[1573116048] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.1% (9369996 kB) free!
[1573116139] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.1% (4068756 kB) free!
[1573116168] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (5303888 kB) free!
[1573116259] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.5% (3703144 kB) free!
[1573116379] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.0% (3360320 kB) free!
[1573116390] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2402268 kB) free!
[1573116606] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4646152 kB) free!
[1573116663] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8802064 kB) free!
[1573116726] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.2% (4465040 kB) free!
[1573116768] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14605244 kB) free!
[1573116783] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.3% (8533388 kB) free!
[1573116846] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.3% (4252004 kB) free!
[1573116903] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1573116903] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.7% (8251884 kB) free!
[1573116979] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (1852680 kB) free!
[1573116993] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.4% (22298384 kB) free.
[1573117109] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.6% (21959056 kB) free.
[1573117368] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (11609360 kB) free!
[1573117968] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (26509592 kB) free.
[1573118046] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (180088 kB) free!
[1573118086] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.7% (2871176 kB) free!
[1573118196] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183592 kB) free!
[1573118206] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 11.4% (2791292 kB) free!
[1573118292] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/07 14:47:08 [error] 24130#0: *3508222 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:3996/121, bytes from/to upstream:121/3996
[1573118316] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181480 kB) free!
[1573118319] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163728 kB) free!
[1573118326] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.0% (2697040 kB) free!
[1573118412] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573118436] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190076 kB) free!
[1573118439] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (178420 kB) free!
[1573118506] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 17.9% (6241368 kB) free!
[1573118559] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191036 kB) free!
[1573118569] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (1331792 kB) free!
[1573118626] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 17.2% (5980432 kB) free!
[1573118689] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (390044 kB) free!
[1573118703] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4949372 kB) free!
[1573118746] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 16.5% (5738508 kB) free!
[1573118809] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (383492 kB) free!
[1573118926] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1573118926] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2267100 kB) free!
[1573119346] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 92.6% (32275384 kB) free.
[1573119379] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 83.2% (18674720 kB) free.
[1573119487] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1573119607] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 366 processes
[1573119612] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(45) < 2019/11/07 15:10:11 [error] 24132#0: *3518784 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1573119727] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 388 processes
[1573119733] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(17) < 2019/11/07 15:12:12 [error] 24134#0: *3519688 recv() failed (104: Connection reset by peer) while proxying connection, client: 121.241.242.109, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1573119850] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1573119853] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(19) < 2019/11/07 15:14:02 [error] 24131#0: *3519860 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.209.99.7, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1573119926] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 15.8% (7776132 kB) free!
[1573119970] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1573120009] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.5% (23100292 kB) free!
[1573120046] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.8% (7776180 kB) free!
[1573120166] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.5% (8615956 kB) free!
[1573120327] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1573120609] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (12017476 kB) free!
[1573121048] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.6% (22843896 kB) free.
[1573121053] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573121137] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.6% (22452296 kB) free.
[1573121265] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.7% (22242452 kB) free.
[1573121304] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 53.9% (13214284 kB) free.
[1573121352] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.8% (47199420 kB) free.
[1573121424] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.1% (47643116 kB) free.
[1573121525] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1573121525] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.9% (40344960 kB) free.
[1573122370] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1573122490] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1573122587] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.3% (8522776 kB) free!
[1573122707] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 22.8% (11224556 kB) free.
[1573123228] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11481448 kB) free!
[1573123307] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9686708 kB) free!
[1573123348] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (43020724 kB) free.
[1573123428] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.0% (9367572 kB) free!
[1573123457] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193296 kB) free!
[1573123548] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 18.3% (9031916 kB) free!
[1573123577] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179580 kB) free!
[1573123697] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177788 kB) free!
[1573123853] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.0% (22709940 kB) free.
[1573123948] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16516664 kB) free!
[1573124068] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (11945528 kB) free!
[1573124188] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (7491288 kB) free!
[1573124592] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.8% (2731080 kB) free!
[1573124713] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (2955552 kB) free!
[1573124788] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.6% (17918936 kB) free!
[1573124833] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.5% (3164752 kB) free!
[1573125348] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (3308972 kB) free!
[1573125388] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (31004972 kB) free.
[1573125988] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6515024 kB) free!
[1573126033] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3432040 kB) free.
[1573126078] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 12.7% (6236804 kB) free!
[1573126107] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4883328 kB) free!
[1573126108] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38951292 kB) free.
[1573126199] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (224468 kB) free!
[1573126227] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (4732356 kB) free!
[1573126256] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181588 kB) free!
[1573126319] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (226660 kB) free!
[1573126347] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.4% (4273780 kB) free!
[1573126376] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184108 kB) free!
[1573126453] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2019/11/07 17:00:38 [error] 24130#0: *3559808 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573126497] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169332 kB) free!
[1573126573] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573126709] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12723996 kB) free!
[1573126829] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8249464 kB) free!
[1573126928] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.4% (9550756 kB) free!
[1573126949] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.9% (3823868 kB) free!
[1573127048] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (8916540 kB) free!
[1573127168] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (8904200 kB) free!
[1573127233] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2993520 kB) free!
[1573127353] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.0% (3414764 kB) free.
[1573127549] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.5% (13893056 kB) free!
[1573127779] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.3% (3659552 kB) free!
[1573127849] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (3218496 kB) free!
[1573127899] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.8% (3326828 kB) free!
[1573127969] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.8% (3216452 kB) free!
[1573128019] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.3% (2980020 kB) free!
[1573128089] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 19.8% (3210260 kB) free!
[1573128147] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2042980 kB) free!
[1573128559] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.5% (9124548 kB) free!
[1573128619] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1573128619] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (1122840 kB) free!
[1573128679] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.8% (8766416 kB) free!
[1573128749] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (1330140 kB) free!
[1573128799] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.1% (8422672 kB) free!
[1573128946] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1573129546] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 18.6% (6471352 kB) free!
[1573129666] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 17.5% (6106704 kB) free!
[1573129786] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 16.3% (5680564 kB) free!
[1573129819] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 83.0% (18646904 kB) free.
[1573129949] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26037388 kB) free!
[1573130149] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 76.5% (37712176 kB) free.
[1573130269] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1573130387] SERVICE FLAPPING ALERT: sig2-vm2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1573130387] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 92.5% (32212624 kB) free.
[1573130519] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 41.8% (20600636 kB) free.
[1573130549] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (1313472 kB) free!
[1573130599] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 49.5% (24414572 kB) free.
[1573130746] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.8% (4119616 kB) free!
[1573130769] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4486916 kB) free!
[1573130866] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.1% (3955264 kB) free!
[1573130953] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.4% (2671892 kB) free!
[1573130986] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.4% (3782600 kB) free!
[1573131004] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1573131073] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.4% (3312092 kB) free.
[1573131124] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1573131149] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (13993160 kB) free!
[1573131244] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1573131673] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (1445216 kB) free!
[1573131749] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27637936 kB) free.
[1573131793] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 12.8% (2073608 kB) free!
[1573131906] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.23044.415> (10.147.212.158:33956 -> 10.147.212.204:5672, vhost: '/', user: 'guest'), channel 1:
[1573131913] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 17.6% (2858812 kB) free!
[1573132026] SERVICE ALERT: rq-2;rq2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573132187] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2012412 kB) free!
[1573132349] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (3272016 kB) free!
[1573132469] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (400336 kB) free!
[1573132513] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1573132513] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (1489300 kB) free!
[1573132533] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2019/11/07 18:44:48 [error] 23644#0: *1688691 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:26648936/7666952, bytes from/to upstream:7666952/26648936
[1573132589] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31349212 kB) free.
[1573132653] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573132889] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (1450488 kB) free!
[1573133044] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1573133113] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 15.4% (2504736 kB) free!
[1573133190] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (6591272 kB) free!
[1573133253] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/11/07 18:54:08 [error] 23644#0: *1688645 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:29524302/8499956, bytes from/to upstream:8499956/29524302
[1573133310] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2199824 kB) free!
[1573133373] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573133430] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.0% (23702928 kB) free!
[1573133714] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 31.9% (5182440 kB) free.
[1573134030] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (7317472 kB) free!
[1573134630] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.6% (13940628 kB) free!
[1573134914] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.8% (3048944 kB) free!
[1573135034] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.2% (3941536 kB) free.
[1573135497] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1573135634] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.8% (2560132 kB) free!
[1573135754] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.1% (3431788 kB) free.
[1573135830] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1152636 kB) free!
[1573136354] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.0% (1786672 kB) free!
[1573136430] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.6% (21848156 kB) free!
[1573136474] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.4% (2498344 kB) free!
[1573136594] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.9% (3077624 kB) free!
[1573136987] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 303 processes
[1573137030] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (1012580 kB) free!
[1573137107] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1573137194] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (907860 kB) free!
[1573137573] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(56) < 2019/11/07 20:00:57 [error] 23646#0: *1735755 connect() failed (111: Connection refused) while connecting to upstream, client: 49.50.67.175, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573137693] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573137794] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 34.9% (5670856 kB) free.
[1573138127] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.9% (4878664 kB) free!
[1573138230] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.5% (24441648 kB) free!
[1573138247] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.8% (4859152 kB) free!
[1573138348] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1573138368] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.8% (4845456 kB) free!
[1573138399] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9704376 kB) free!
[1573138519] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.5% (9627884 kB) free!
[1573138639] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.4% (9549252 kB) free!
[1573138830] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.7% (4814000 kB) free!
[1573138994] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.1% (2618488 kB) free!
[1573139114] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (3126616 kB) free!
[1573139170] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1573139234] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.9% (3716260 kB) free.
[1573139385] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.8% (22418196 kB) free.
[1573139388] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 89.2% (21856688 kB) free.
[1573139404] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 40.8% (53813532 kB) free.
[1573139416] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 92.4% (22645712 kB) free.
[1573139433] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.8% (47193520 kB) free.
[1573139529] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.8% (40325208 kB) free.
[1573139547] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 52.3% (12819392 kB) free.
[1573139571] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.1% (14637500 kB) free.
[1573140619] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1573140693] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(161) < 2019/11/07 20:56:40 [error] 23638#0: *1742097 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1573140813] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573141787] SERVICE FLAPPING ALERT: sig2-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573141789] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183972 kB) free!
[1573141909] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (190936 kB) free!
[1573142029] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164920 kB) free!
[1573142409] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.4% (22948664 kB) free!
[1573142519] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (9791164 kB) free!
[1573142529] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.7% (19374376 kB) free!
[1573142639] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (9767324 kB) free!
[1573142649] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.5% (44151036 kB) free.
[1573142759] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.8% (9748112 kB) free!
[1573143249] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25593844 kB) free!
[1573143369] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.7% (22016160 kB) free!
[1573143434] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.4% (2828596 kB) free!
[1573143489] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44527304 kB) free.
[1573143554] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.3% (2979064 kB) free!
[1573143674] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (3062484 kB) free!
[1573144274] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.9% (3399024 kB) free.
[1573144353] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1573144473] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1573144573] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(224) < 2019/11/07 22:06:13 [error] 24130#0: *3617667 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.64.122, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:39/0, bytes from/to upstream:0/39
[1573144694] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(39) < 2019/11/07 22:07:54 [error] 24130#0: *3618087 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:58/0, bytes from/to upstream:0/58
[1573144814] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(18) < 2019/11/07 22:09:35 [error] 24132#0: *3618274 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.182.41, server: 0.0.0.0:3001, upstream: "10.147.212.142:3001", bytes from/to client:37/0, bytes from/to upstream:0/37
[1573145289] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19452724 kB) free!
[1573145409] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43311816 kB) free.
[1573145414] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573146009] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25391528 kB) free!
[1573146074] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.3% (2813944 kB) free!
[1573146129] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.7% (21968260 kB) free!
[1573146194] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 35.6% (5794436 kB) free.
[1573146249] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44313860 kB) free.
[1573146614] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(10) < 2019/11/07 22:32:08 [error] 24130#0: *3630023 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.65.17, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1573146734] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573147934] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2019/11/07 23:00:41 [error] 24130#0: *3645331 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.155:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573148049] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19740632 kB) free!
[1573148055] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573148150] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (4688924 kB) free!
[1573148169] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42559188 kB) free.
[1573148186] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.7% (25990948 kB) free!
[1573148270] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 24.5% (12072356 kB) free.
[1573148306] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (12826196 kB) free!
[1573148426] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (7618672 kB) free!
[1573148870] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (214884 kB) free!
[1573148990] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236736 kB) free!
[1573149110] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 24.9% (12273592 kB) free.
[1573149216] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 783 processes
[1573149250] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 193750 MB (18% inode=99%):
[1573149254] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(9) < 2019/11/07 23:16:14 [error] 24131#0: *3653917 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573149370] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 193750 MB (18% inode=99%):
[1573149374] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573149490] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 193750 MB (18% inode=99%):
[1573149626] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.3% (14901740 kB) free!
[1573149636] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1573149969] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (22064112 kB) free!
[1573150089] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.7% (43082504 kB) free.
[1573150347] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4727952 kB) free!
[1573150467] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (4698380 kB) free!
[1573150587] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.0% (4664168 kB) free!
[1573150826] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (11899176 kB) free!
[1573151289] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)