[1579804213] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579804741] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187140 kB) free!
[1579804847] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187584 kB) free!
[1579804861] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188320 kB) free!
[1579804881] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197888 kB) free!
[1579804968] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194668 kB) free!
[1579804982] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192112 kB) free!
[1579805001] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184940 kB) free!
[1579805088] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183316 kB) free!
[1579805121] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184004 kB) free!
[1579808088] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 22.0% (5439108 kB) free.
[1579808688] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197248 kB) free!
[1579808808] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182072 kB) free!
[1579808928] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181964 kB) free!
[1579809198] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (215180 kB) free!
[1579812921] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 43.4% (10632524 kB) free.
[1579812977] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (219804 kB) free!
[1579813097] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (241748 kB) free!
[1579813217] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (223588 kB) free!
[1579813813] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1579813817] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.1% (42923620 kB) free.
[1579815205] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/24 03:00:25 [error] 18365#0: *43132073 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1579815324] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1432) < 203.212.70.170 [24/Jan/2020:03:01:32 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579815325] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579815444] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579822522] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189772 kB) free!
[1579822644] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163036 kB) free!
[1579822764] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 29.3% (7173132 kB) free.
[1579822982] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.5% (21933180 kB) free.
[1579823328] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 49.6% (12240896 kB) free.
[1579823365] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195552 kB) free!
[1579823485] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (201892 kB) free!
[1579823605] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197744 kB) free!
[1579827015] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.8% (40325164 kB) free.
[1579827049] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.7% (22719060 kB) free.
[1579828128] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (4479104 kB) free!
[1579828248] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (4476348 kB) free!
[1579828368] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 18.1% (4472884 kB) free!
[1579830168] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2455816 kB) free!
[1579830444] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6838) < 103.255.146.164 [24/Jan/2020:07:09:47 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1579830564] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579831164] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(823) < 103.255.146.164 [24/Jan/2020:07:20:57 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1579831285] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(526) < 103.255.146.164 [24/Jan/2020:07:31:24 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1579831405] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(2738) < 103.255.146.164 [24/Jan/2020:07:32:09 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1579832005] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579833325] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2020/01/24 07:59:20 [error] 18370#0: *43133196 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:99257/53642, bytes from/to upstream:53642/99257
[1579833445] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579833489] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1579833489] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.7% (22634532 kB) free.
[1579836806] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(4743) < 103.255.146.164 [24/Jan/2020:08:58:56 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1579836926] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579837213] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1579837333] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1579837453] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1579839361] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 506.15, 279.64, 115.78
[1579839481] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 82.81, 192.78, 103.84
[1579839602] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 19.68, 130.60, 92.14
[1579839680] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.7% (4838844 kB) free!
[1579839800] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.6% (4812108 kB) free!
[1579839920] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.5% (4784492 kB) free!
[1579840696] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1579841402] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 8.31, 10.05, 27.45
[1579841726] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6101) < 103.255.146.164 [24/Jan/2020:10:23:15 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1579841847] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579842446] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(7549) < 103.255.146.164 [24/Jan/2020:10:32:05 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1579842566] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579843376] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.1% (9399836 kB) free!
[1579843496] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 19.1% (9400292 kB) free!
[1579843616] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.1% (9399492 kB) free!
[1579843802] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 90.88, 43.81, 25.22
[1579843922] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 25.00, 34.22, 23.96
[1579844053] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579844522] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 50.38, 86.11, 51.76
[1579844642] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 17.58, 61.95, 47.08
[1579844762] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 11.29, 45.35, 42.76
[1579844945] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1579845065] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1579845362] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1579845362] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 10.18, 14.60, 27.20
[1579845445] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(65) < 2020/01/24 11:24:19 [error] 18359#0: *43227143 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579845565] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579845566] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(40) < 203.212.70.170 [24/Jan/2020:11:26:42 +0530] TCP 502 0 0 180.000 "10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 60.000, 60.000, 0.000"
[1579845597] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 103.16.101.72 [24/Jan/2020:11:24:19 +0530] TCP 502 0 0 60.000 "10.147.212.49:4000, backend3" "0, 0" "0, 0" "60.000, 0.000"
[1579845687] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579845718] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579847296] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4696788 kB) free!
[1579847416] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.7% (4605452 kB) free!
[1579847418] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 55.1% (27170488 kB) free.
[1579847537] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.3% (4281280 kB) free!
[1579847609] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26216968 kB) free!
[1579847729] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.4% (21564008 kB) free!
[1579847849] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17134316 kB) free!
[1579848019] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (240116 kB) free!
[1579848137] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193096 kB) free!
[1579848139] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (215308 kB) free!
[1579848236] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.5% (39663492 kB) free.
[1579848285] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.7% (22133640 kB) free.
[1579848350] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 51.7% (68203524 kB) free.
[1579848840] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236592 kB) free!
[1579848916] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (184428 kB) free!
[1579848960] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (214040 kB) free!
[1579848962] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 14.60, 43.65, 38.80
[1579849036] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185848 kB) free!
[1579849080] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (217132 kB) free!
[1579849083] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 14.28, 33.98, 35.87
[1579849156] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1579849156] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190460 kB) free!
[1579849209] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 13.12, 26.97, 33.06
[1579849320] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.9% (4626180 kB) free!
[1579849440] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.3% (4485560 kB) free!
[1579849560] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.4% (21807548 kB) free.
[1579849560] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.6% (4303464 kB) free!
[1579849626] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 76.5% (37679044 kB) free.
[1579849661] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.9% (21796284 kB) free.
[1579849809] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 49.84, 107.68, 71.10
[1579849857] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187480 kB) free!
[1579849977] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (184864 kB) free!
[1579850097] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190936 kB) free!
[1579850158] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1579850172] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194932 kB) free!
[1579850229] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (232192 kB) free!
[1579850278] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 334 processes
[1579850292] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165508 kB) free!
[1579850349] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (248168 kB) free!
[1579850398] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1579850412] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (240956 kB) free!
[1579850469] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (247856 kB) free!
[1579850965] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(40) < 2020/01/24 12:49:49 [error] 18358#0: *43264871 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:2444/87, bytes from/to upstream:87/2444
[1579850998] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1579851042] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 331 processes
[1579851085] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579851162] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1579851609] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 10.79, 13.21, 27.78
[1579851808] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.0% (22704836 kB) free.
[1579851843] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1579851843] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.0% (40404832 kB) free.
[1579851879] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 79.3% (19440556 kB) free.
[1579852418] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191608 kB) free!
[1579852448] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (229600 kB) free!
[1579852487] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181900 kB) free!
[1579852538] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (196940 kB) free!
[1579852568] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (237280 kB) free!
[1579852607] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177036 kB) free!
[1579852658] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (193540 kB) free!
[1579852688] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (234648 kB) free!
[1579852727] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188540 kB) free!
[1579852781] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.0% (21969296 kB) free.
[1579852824] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.1% (39952160 kB) free.
[1579852860] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1579852860] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 74.7% (18309276 kB) free.
[1579853388] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188828 kB) free!
[1579853409] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 12.78, 35.22, 39.51
[1579853428] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236248 kB) free!
[1579853469] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172480 kB) free!
[1579853508] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (185000 kB) free!
[1579853529] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 375.84, 136.46, 73.86
[1579853548] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (225800 kB) free!
[1579853590] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182336 kB) free!
[1579853628] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197392 kB) free!
[1579853638] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.9% (22759140 kB) free.
[1579853649] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 99.03, 125.29, 78.61
[1579853668] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (237380 kB) free!
[1579853677] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.3% (22277576 kB) free.
[1579853710] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.7% (39774720 kB) free.
[1579853998] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1579854085] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(34) < 2020/01/24 13:49:58 [error] 18368#0: *43265468 recv() failed (104: Connection reset by peer) while proxying connection, client: 124.153.112.226, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:1306/1265, bytes from/to upstream:1265/1306
[1579854087] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 182.18.144.246 [24/Jan/2020:13:51:09 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1579854118] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 362 processes
[1579854205] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579854207] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1) < 203.212.70.170 [24/Jan/2020:13:52:08 +0530] TCP 502 0 0 120.002 "10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.001, 0.000"
[1579854278] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188976 kB) free!
[1579854318] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (226360 kB) free!
[1579854327] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579854398] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185828 kB) free!
[1579854438] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (215020 kB) free!
[1579854518] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164612 kB) free!
[1579854558] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (249748 kB) free!
[1579854765] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.6% (22594932 kB) free.
[1579854865] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.0% (40425724 kB) free.
[1579855449] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.74, 16.75, 35.67
[1579856049] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 8.65, 10.13, 23.07
[1579856569] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (1212976 kB) free!
[1579856689] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (367048 kB) free!
[1579856809] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (238328 kB) free!
[1579857443] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.11882.5699> (10.147.212.191:41196 -> 10.147.212.203:5672, vhost: '/', user: 'guest'), channel 1:
[1579857564] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1579858737] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.5% (22340700 kB) free.
[1579859160] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11565420 kB) free!
[1579859280] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39532820 kB) free.
[1579859337] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.3% (3292388 kB) free!
[1579859457] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.0% (3200560 kB) free!
[1579859577] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.6% (3098332 kB) free!
[1579859880] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9423104 kB) free!
[1579860000] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36335044 kB) free.
[1579860177] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188092 kB) free!
[1579860850] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1579861156] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.1% (22238200 kB) free.
[1579861200] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10184204 kB) free!
[1579861281] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 56.3% (74277948 kB) free.
[1579861382] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (4752984 kB) free!
[1579861502] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.3% (4725656 kB) free!
[1579861622] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.2% (4696528 kB) free!
[1579862051] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166952 kB) free!
[1579862120] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;HARD;3;CRITICAL - 10.0% (2444100 kB) free!
[1579862171] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173248 kB) free!
[1579862194] SERVICE ALERT: db2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 339 processes
[1579862291] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188900 kB) free!
[1579862299] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.7% (21991624 kB) free.
[1579862314] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1579862434] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1579862446] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.8% (22248696 kB) free.
[1579862487] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 90.9% (22268348 kB) free.
[1579862908] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193104 kB) free!
[1579863028] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195960 kB) free!
[1579863148] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164860 kB) free!
[1579863162] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 333 processes
[1579863282] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1579863282] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.8% (22739228 kB) free.
[1579863282] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 331 processes
[1579863634] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1579865668] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1579866205] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/24 17:08:00 [error] 18366#0: *43379065 open() "/usr/share/nginx/html/js/tmpl/login.ejs" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /js/tmpl/login.ejs?0.3143749237211735 HTTP/1.1", host: "qservercluster1.com:8091", referrer: "http://qservercluster1.com:8091/"
[1579866325] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579867449] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1579869328] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(678) < 103.255.146.164 [24/Jan/2020:17:56:21 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1579869448] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579870217] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9842312 kB) free!
[1579870337] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.9% (9814208 kB) free!
[1579870457] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.9% (9783152 kB) free!
[1579871437] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1579872292] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173760 kB) free!
[1579872412] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180676 kB) free!
[1579872532] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194556 kB) free!
[1579873237] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.7% (4115240 kB) free!
[1579873357] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.1% (4213180 kB) free!
[1579873477] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.0% (4193780 kB) free!
[1579873482] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1579873602] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1579874077] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (1508900 kB) free!
[1579874125] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(171) < 2020/01/24 19:19:25 [error] 18368#0: *43424269 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579874245] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579874845] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(323) < 2020/01/24 19:33:25 [error] 18367#0: *43429175 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1579874849] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(702) < 49.50.64.122 [24/Jan/2020:19:35:19 +0530] TCP 502 0 0 60.000 "10.147.212.143:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1579874965] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579874969] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579875118] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(231) < 103.16.101.72 [24/Jan/2020:19:32:23 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1579875238] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579875565] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/24 19:47:36 [error] 18367#0: *43431808 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET / HTTP/1.0", upstream: "http://10.147.212.40:80/", host: "altsmsc.tesync.net"
[1579875685] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2) < 2020/01/24 19:51:19 [error] 18370#0: *43432340 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET / HTTP/1.0", upstream: "http://10.147.212.40:80/", host: "altsmsc.tesync.net"
[1579875805] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579876406] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/24 20:00:33 [error] 18365#0: *43433825 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /main.php?page=cligrptosenderid HTTP/1.0", upstream: "http://10.147.212.40:80/main.php?page=cligrptosenderid", host: "altsmsc.tesync.net"
[1579876526] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1) < 2020/01/24 20:03:45 [error] 18369#0: *43434334 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET / HTTP/1.0", upstream: "http://10.147.212.40:80/", host: "altsmsc.tesync.net"
[1579876646] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579877846] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/24 20:24:42 [error] 18370#0: *43437373 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "GET /main.php?page=cligrptosenderid&action=add HTTP/1.0", upstream: "http://10.147.212.40:80/main.php?page=cligrptosenderid&action=add", host: "altsmsc.tesync.net"
[1579877966] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579878532] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1579879265] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1579879385] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1579883657] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4888488 kB) free!
[1579889966] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/01/24 23:40:30 [error] 18361#0: *43463609 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1579890086] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found