[1562351432] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (13795768 kB) free.
[1562351744] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1562351930] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8298432 kB) free!
[1562351996] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562352032] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (9402592 kB) free!
[1562352051] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11926136 kB) free.
[1562352116] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562352152] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (5678308 kB) free!
[1562352272] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (13797184 kB) free.
[1562352409] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.20, 0.70, 3.49
[1562352650] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8255596 kB) free!
[1562352716] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562352770] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11932856 kB) free.
[1562352837] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562352872] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (8125256 kB) free!
[1562352992] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.9% (4519060 kB) free!
[1562353112] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (12838780 kB) free!
[1562353370] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8263636 kB) free!
[1562353436] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562353490] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11933632 kB) free.
[1562353556] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562353609] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 6.00, 28.90, 17.09
[1562353729] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.08, 19.44, 15.06
[1562353849] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.42, 13.16, 13.29
[1562354090] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8243548 kB) free!
[1562354156] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562354210] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11928324 kB) free.
[1562354276] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562354744] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 40.86, 43.33, 20.27
[1562354810] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8233196 kB) free!
[1562354864] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 7.76, 29.83, 18.12
[1562354876] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562354912] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (6540160 kB) free!
[1562354930] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11933088 kB) free.
[1562354985] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.08, 20.98, 16.29
[1562354996] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562355050] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.44, 0.62, 3.88
[1562355060] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 437 processes
[1562355120] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 432 processes
[1562355225] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (2199000 kB) free!
[1562355285] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (2196908 kB) free!
[1562355345] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 9.0% (2200024 kB) free!
[1562355405] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 9.0% (2203044 kB) free!
[1562355512] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (13029888 kB) free!
[1562355530] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8273052 kB) free!
[1562355585] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 82.82, 43.98, 25.04
[1562355596] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562355650] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11922856 kB) free.
[1562355717] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562356250] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8249236 kB) free!
[1562356316] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562356353] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562356371] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11928040 kB) free.
[1562356436] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562356970] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8241400 kB) free!
[1562357036] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562357090] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11931536 kB) free.
[1562357156] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562357312] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (6352108 kB) free!
[1562357385] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1562357385] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.78, 3.53, 12.39
[1562357690] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8188020 kB) free!
[1562357756] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562357810] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11929376 kB) free.
[1562357876] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562357912] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.9% (11117620 kB) free!
[1562358410] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.7% (8230268 kB) free!
[1562358476] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562358513] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (5115492 kB) free!
[1562358530] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11932956 kB) free.
[1562358597] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562359113] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.5% (10873196 kB) free!
[1562359130] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.8% (8276436 kB) free!
[1562359185] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.91, 2.56, 3.97
[1562359196] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562359250] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11924964 kB) free.
[1562359316] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562359713] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (4713224 kB) free!
[1562359850] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8333496 kB) free!
[1562359916] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562359970] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11931964 kB) free.
[1562360036] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562360287] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.88, 5.08, 5.02
[1562360313] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (9833720 kB) free!
[1562360407] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.00, 5.42, 5.15
[1562360527] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 5.37, 5.43, 5.18
[1562360570] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.9% (8335008 kB) free!
[1562360636] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562360691] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11933352 kB) free.
[1562360756] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562360913] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (13847016 kB) free.
[1562361127] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.46, 4.76, 4.90
[1562361153] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 75.22, 68.87, 34.51
[1562361273] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 10.35, 46.18, 30.36
[1562361290] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8198740 kB) free!
[1562361356] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562361393] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.87, 31.06, 26.74
[1562361410] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11929580 kB) free.
[1562361476] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562361513] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (9222680 kB) free!
[1562361566] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.8% (2656360 kB) free!
[1562361633] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.5% (5579308 kB) free!
[1562361650] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 21.91, 5.48, 2.01
[1562361753] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (13347996 kB) free.
[1562361770] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 51.11, 22.06, 8.50
[1562361890] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 54.81, 33.09, 14.18
[1562361993] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.26, 4.42, 14.16
[1562362010] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.3% (8039996 kB) free!
[1562362130] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11929960 kB) free.
[1562362327] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.09, 5.23, 5.05
[1562362353] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.4% (6856916 kB) free!
[1562362447] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.54, 5.01, 4.99
[1562362473] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (13848756 kB) free.
[1562362490] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.60, 13.43, 15.76
[1562362588] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.9% (2672764 kB) free!
[1562362676] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562362730] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.9% (7856240 kB) free!
[1562362796] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562362850] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11919540 kB) free.
[1562363047] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.32, 5.65, 5.35
[1562363073] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (9016312 kB) free!
[1562363167] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.38, 5.47, 5.32
[1562363193] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (5441976 kB) free!
[1562363193] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.45, 0.57, 4.24
[1562363287] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.63, 5.05, 5.17
[1562363313] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (13269476 kB) free.
[1562363366] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.5% (2334264 kB) free!
[1562363396] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562363450] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.7% (7745444 kB) free!
[1562363516] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562363570] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11925908 kB) free.
[1562363690] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.29, 0.56, 4.57
[1562363913] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (7579392 kB) free!
[1562364033] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (13841948 kB) free.
[1562364116] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562364170] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7594404 kB) free!
[1562364236] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562364290] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11924292 kB) free.
[1562364388] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2459160 kB) free!
[1562364487] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1562364487] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.93, 4.33, 4.88
[1562364633] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.0% (9225004 kB) free!
[1562364754] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (5506064 kB) free!
[1562364836] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562364874] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.1% (13241072 kB) free.
[1562364890] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7431244 kB) free!
[1562364957] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562365011] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11925592 kB) free.
[1562365474] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.1% (6674672 kB) free!
[1562365490] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.15, 13.96, 15.75
[1562365556] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562365588] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.5% (2573408 kB) free!
[1562365594] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (13848084 kB) free.
[1562365610] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.1% (7442184 kB) free!
[1562365610] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.38, 9.49, 13.89
[1562365676] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562365730] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11927080 kB) free.
[1562365730] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.19, 6.40, 12.23
[1562366194] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (8276068 kB) free!
[1562366260] SERVICE ALERT: pservercluster2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 577 processes
[1562366276] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562366314] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.3% (4818936 kB) free!
[1562366330] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7392136 kB) free!
[1562366380] SERVICE ALERT: pservercluster2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 576 processes
[1562366385] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 10.83, 53.08, 34.77
[1562366396] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562366434] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.3% (12669008 kB) free!
[1562366450] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11918788 kB) free.
[1562366500] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 358 processes
[1562366505] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 4.19, 36.62, 30.97
[1562366625] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.76, 25.36, 27.53
[1562366930] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.19, 0.38, 3.58
[1562366996] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562367034] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (6389144 kB) free!
[1562367050] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.9% (7340984 kB) free!
[1562367116] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562367170] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11925508 kB) free.
[1562367225] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.72, 5.81, 15.78
[1562367388] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1562367388] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (2405172 kB) free!
[1562367393] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.10, 43.67, 36.87
[1562367487] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.99, 5.97, 5.19
[1562367513] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.22, 29.26, 32.42
[1562367566] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.9% (3931684 kB) free!
[1562367607] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.35, 6.83, 5.60
[1562367633] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1562367633] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.47, 19.74, 28.54
[1562367634] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.8% (11020084 kB) free!
[1562367717] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562367727] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 9.77, 7.87, 6.13
[1562367770] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7249876 kB) free!
[1562367794] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 3 warn out of 361 processes
[1562367837] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562367890] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11917980 kB) free.
[1562367914] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 3 warn out of 363 processes
[1562368035] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 360 processes
[1562368233] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.53, 3.00, 15.14
[1562368234] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (13875928 kB) free.
[1562368437] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562368490] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.2% (7469648 kB) free!
[1562368557] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562368610] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11915612 kB) free.
[1562368635] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1562368834] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (9400084 kB) free!
[1562368954] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (5845944 kB) free!
[1562369025] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.70, 2.89, 4.72
[1562369074] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.4% (13408644 kB) free.
[1562369157] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562369210] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7532400 kB) free!
[1562369277] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562369331] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11915852 kB) free.
[1562369433] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.20, 0.45, 4.48
[1562369674] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (7278580 kB) free!
[1562369794] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (13875752 kB) free.
[1562369877] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562369930] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7611060 kB) free!
[1562369997] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562370050] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11910712 kB) free.
[1562370394] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (8421672 kB) free!
[1562370515] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (4747976 kB) free!
[1562370598] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562370635] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.8% (12339904 kB) free!
[1562370650] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.4% (7567920 kB) free!
[1562370718] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562370770] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.2% (11918528 kB) free.
[1562370988] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.2% (2493316 kB) free!
[1562371235] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (6146880 kB) free!
[1562371318] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562371370] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7548784 kB) free!
[1562371438] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562371490] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.4% (11537168 kB) free.
[1562371835] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 20.0% (13152628 kB) free!
[1562371927] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.65, 4.74, 4.99
[1562372038] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562372158] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562372278] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562372690] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 17.4% (8559688 kB) free!
[1562372810] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11896824 kB) free.
[1562373635] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (13614996 kB) free.
[1562374010] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.5% (9601764 kB) free!
[1562374131] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8291096 kB) free!
[1562374235] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (9891224 kB) free!
[1562374251] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.1% (11864308 kB) free.
[1562374355] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.4% (6810124 kB) free!
[1562374359] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.2% (3607368 kB) free.
[1562374475] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (3585388 kB) free!
[1562375075] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.6% (11585852 kB) free!
[1562375330] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.30, 18.31, 8.60
[1562375450] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.65, 12.30, 7.57
[1562375451] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9242740 kB) free!
[1562375570] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.43, 8.31, 6.68
[1562375571] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 15.9% (7856964 kB) free!
[1562375625] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.25, 13.46, 12.25
[1562375691] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.9% (11793976 kB) free.
[1562375745] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.57, 9.80, 11.05
[1562375865] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.73, 7.47, 10.05
[1562376170] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.63, 1.46, 3.67
[1562376275] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (14838820 kB) free.
[1562376388] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2455040 kB) free!
[1562376875] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (11850620 kB) free!
[1562376995] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.1% (8594408 kB) free!
[1562377065] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.47, 2.61, 4.60
[1562377078] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562377116] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (14804752 kB) free.
[1562377234] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 26.17, 34.42, 15.52
[1562377354] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.95, 23.16, 13.69
[1562377474] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.72, 15.59, 12.06
[1562377491] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.6% (7203752 kB) free!
[1562377611] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11843604 kB) free.
[1562377715] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (11861980 kB) free!
[1562377835] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.8% (8427780 kB) free!
[1562377955] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (14786804 kB) free.
[1562378211] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.3% (7042628 kB) free!
[1562378331] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11845372 kB) free.
[1562378555] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (10324480 kB) free!
[1562378674] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.05, 0.46, 3.50
[1562378675] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.9% (7177696 kB) free!
[1562378795] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (14763968 kB) free.
[1562378879] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562378999] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562379395] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (9260124 kB) free!
[1562379515] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.0% (5895524 kB) free!
[1562379531] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.6% (8159952 kB) free!
[1562379599] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562379635] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (13681660 kB) free.
[1562379651] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11837960 kB) free.
[1562379719] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562379727] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1562380065] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 162.08, 42.64, 16.20
[1562380185] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 62.91, 59.41, 26.73
[1562380235] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (7478700 kB) free!
[1562380251] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.4% (8096856 kB) free!
[1562380305] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 10.73, 40.58, 23.79
[1562380319] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562380355] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (14716764 kB) free.
[1562380371] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11840564 kB) free.
[1562380439] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562380559] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562380905] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.90, 8.10, 13.93
[1562380955] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (10680056 kB) free!
[1562381075] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.2% (7379764 kB) free!
[1562381195] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (14660828 kB) free.
[1562381505] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 23.8% (5834576 kB) free.
[1562381752] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 370 processes
[1562381759] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1562381759] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1562381796] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (12516532 kB) free!
[1562381872] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 374 processes
[1562381916] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.8% (9759716 kB) free!
[1562381992] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 375 processes
[1562382036] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.8% (7084816 kB) free!
[1562382171] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.8% (7287036 kB) free!
[1562382291] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11839464 kB) free.
[1562382636] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (14516636 kB) free.
[1562382891] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.0% (7405592 kB) free!
[1562382989] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1562383011] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.0% (11835856 kB) free.
[1562383166] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183572 kB) free!
[1562383236] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (11517300 kB) free!
[1562383255] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1562383345] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 323 processes
[1562383356] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.1% (8614840 kB) free!
[1562383375] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1562383465] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 321 processes
[1562383474] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.40, 9.42, 6.45
[1562383476] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (5483600 kB) free!
[1562383559] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1562383594] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.37, 6.39, 5.70
[1562383611] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.0% (5911512 kB) free!
[1562383620] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 435 processes
[1562383679] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562383680] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 432 processes
[1562383714] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.21, 4.36, 5.04
[1562383732] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (4640468 kB) free!
[1562383799] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562383852] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.5% (3181744 kB) free!
[1562384076] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.6% (12205524 kB) free!
[1562384280] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 434 processes
[1562384314] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.38, 1.47, 3.04
[1562384341] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 433 processes
[1562384505] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 3.06, 3.14, 4.34
[1562384575] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1562384676] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (5933108 kB) free!
[1562384695] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1562385052] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.4% (7589608 kB) free!
[1562385276] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.1% (12576508 kB) free!
[1562385476] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.58, 6.39, 5.18
[1562385596] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.08, 7.15, 5.61
[1562385705] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.46, 5.77, 6.78
[1562385717] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 9.18, 7.71, 5.99
[1562385825] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.56, 4.70, 6.27
[1562385876] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (3910240 kB) free!
[1562385945] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.69, 4.04, 5.84
[1562386188] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 19.8% (1455904 kB) free!
[1562386309] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 16.6% (1219172 kB) free!
[1562386429] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 12.5% (922836 kB) free!
[1562386476] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.9% (18980748 kB) free.
[1562386545] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.89, 3.08, 4.43
[1562386852] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3746664 kB) free!
[1562387029] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (115108 kB) free!
[1562387076] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (12331000 kB) free!
[1562387196] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (21611580 kB) free.
[1562387399] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562387452] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1562387452] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.4% (7570964 kB) free!
[1562387570] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562387629] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 91.9% (6761212 kB) free.
[1562387999] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562388052] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (2986520 kB) free!
[1562388119] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562388159] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (167976 kB) free!
[1562388170] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 22.50, 5.73, 2.32
[1562388279] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (178324 kB) free!
[1562388290] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 69.91, 29.48, 11.44
[1562388396] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (9737104 kB) free!
[1562388399] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (166716 kB) free!
[1562388410] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 76.34, 45.31, 19.43
[1562388514] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 42.26, 50.16, 26.50
[1562388516] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (20025076 kB) free.
[1562388566] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.8% (2675780 kB) free!
[1562388634] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 6.58, 33.91, 23.41
[1562388652] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.4% (7568672 kB) free!
[1562388754] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 1.57, 22.94, 20.67
[1562388988] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.0% (2688668 kB) free!
[1562389116] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (13118516 kB) free!
[1562389166] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2253616 kB) free!
[1562389236] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.0% (9204092 kB) free!
[1562389354] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.04, 3.91, 11.28
[1562389356] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.9% (19685932 kB) free.
[1562389588] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2388264 kB) free!
[1562389610] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.39, 3.05, 12.28
[1562389852] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4507364 kB) free!
[1562389919] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562390040] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1562390453] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.8% (5804624 kB) free!
[1562390556] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.6% (7610004 kB) free!
[1562390639] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562390676] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (18021260 kB) free.
[1562390759] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562390810] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.45, 0.55, 3.80
[1562391154] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.85, 0.90, 3.55
[1562391276] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (11973312 kB) free!
[1562391295] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1562391359] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562391396] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (7814996 kB) free!
[1562391415] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1562391479] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562391516] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (18057512 kB) free.
[1562391599] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562391653] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3950856 kB) free!
[1562392116] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (12248980 kB) free!
[1562392199] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562392236] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.4% (8186160 kB) free!
[1562392253] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.1% (7430788 kB) free!
[1562392356] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (17761492 kB) free.
[1562392956] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (12062636 kB) free!
[1562393068] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 20.7% (3368600 kB) free.
[1562393076] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.7% (7720892 kB) free!
[1562393197] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.9% (17039964 kB) free.
[1562393400] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562393453] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.1% (2995072 kB) free!
[1562393520] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562393668] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 11.9% (1929304 kB) free!
[1562393788] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 10.3% (1679640 kB) free!
[1562393796] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (12699144 kB) free!
[1562393908] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1419164 kB) free!
[1562393916] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.1% (8590376 kB) free!
[1562394036] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (4590940 kB) free!
[1562394053] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.0% (5406028 kB) free!
[1562394229] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (383420 kB) free!
[1562394349] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (167932 kB) free!
[1562394470] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1562394470] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (132196 kB) free!
[1562394636] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 20.0% (13129304 kB) free!
[1562394720] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562394840] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1562394945] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 73.70, 60.52, 32.43
[1562394960] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562395065] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 58.29, 61.49, 36.33
[1562395070] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.4% (6798792 kB) free.
[1562395185] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 10.36, 42.13, 32.29
[1562395354] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 30.16, 10.54, 4.45
[1562395474] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 5.66, 7.62, 4.11
[1562395560] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562395786] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.32, 7.99, 18.29
[1562395836] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (14976496 kB) free.
[1562396161] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1562396281] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562396436] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (10660328 kB) free!
[1562396556] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (6542264 kB) free!
[1562396676] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (2472336 kB) free!
[1562396881] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562397001] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562397276] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.4% (11430636 kB) free!
[1562397586] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.59, 2.68, 4.91
[1562397876] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (5582384 kB) free!
[1562398201] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562398253] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (3316004 kB) free!
[1562398321] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562398476] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (19533696 kB) free.
[1562398921] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562399042] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562399641] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562399676] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (8266260 kB) free!
[1562399761] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562399797] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (19134952 kB) free.
[1562399881] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562400053] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.5% (7144964 kB) free!
[1562400117] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.34, 3.89, 4.81
[1562400727] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.70, 6.60, 5.32
[1562400847] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.86, 6.62, 5.49
[1562400967] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.99, 6.81, 5.69
[1562400996] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (8293288 kB) free!
[1562401116] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (19400488 kB) free.
[1562401474] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.33, 9.41, 5.55
[1562401567] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.00, 0.91, 2.98
[1562401594] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562401594] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.97, 6.51, 4.95
[1562402270] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 11.0% (807772 kB) free!
[1562402316] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (8268572 kB) free!
[1562402390] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.1% (670008 kB) free!
[1562402436] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (19576672 kB) free.
[1562402510] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (541304 kB) free!
[1562402517] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.04, 6.18, 5.78
[1562402637] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.77, 5.57, 5.60
[1562402665] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 321 processes
[1562402757] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.74, 5.74, 5.66
[1562402785] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1562403410] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 278.46, 167.70, 71.50
[1562403530] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 281.30, 205.43, 96.99
[1562403636] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (8684404 kB) free!
[1562403650] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 112.91, 183.39, 103.22
[1562403653] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4852512 kB) free!
[1562403756] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (18675168 kB) free.
[1562403988] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.1% (2464788 kB) free!
[1562404253] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.2% (7473044 kB) free!
[1562404567] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.90, 6.55, 5.55
[1562404588] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (1867224 kB) free!
[1562404594] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.37, 5.19, 5.52
[1562404687] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.87, 6.34, 5.62
[1562404714] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.89, 4.17, 5.11
[1562404786] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.95, 5.52, 5.95
[1562404807] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.77, 4.27, 4.95
[1562404834] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.53, 3.24, 4.65
[1562404906] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.92, 5.00, 5.71
[1562404956] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (8676308 kB) free!
[1562405026] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.63, 4.18, 5.32
[1562405076] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.0% (4591436 kB) free!
[1562405196] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (16195944 kB) free.
[1562405434] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.40, 7.52, 7.07
[1562405510] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 91.7% (6743008 kB) free.
[1562405554] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.71, 5.28, 6.30
[1562405626] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.91, 2.98, 4.08
[1562405674] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.68, 3.77, 5.62
[1562405796] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (10086244 kB) free!
[1562405917] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.8% (5768744 kB) free!
[1562406037] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (16850840 kB) free.
[1562406050] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.74, 1.08, 11.09
[1562406275] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.90, 1.16, 3.31
[1562406637] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (11649600 kB) free!
[1562406653] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4109736 kB) free!
[1562406757] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.0% (7236680 kB) free!
[1562406877] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.9% (3243400 kB) free!
[1562407250] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.40, 0.47, 3.39
[1562407253] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.6% (7685720 kB) free!
[1562407475] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 120.44, 43.73, 17.51
[1562407477] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (14098828 kB) free.
[1562407595] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 173.16, 89.28, 37.38
[1562407681] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562407715] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 180.32, 119.78, 54.85
[1562408077] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.0% (8548060 kB) free!
[1562408197] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (4245980 kB) free!
[1562408282] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562408317] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (15072824 kB) free.
[1562408366] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.7% (2652444 kB) free!
[1562408402] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562408407] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.55, 6.35, 5.34
[1562408522] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562408527] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.54, 6.75, 5.61
[1562408647] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 7.71, 7.09, 5.87
[1562408917] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (10823308 kB) free!
[1562408966] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2299216 kB) free!
[1562409037] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.9% (6483280 kB) free!
[1562409053] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (3278208 kB) free!
[1562409122] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562409157] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (2277164 kB) free!
[1562409247] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1562409247] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.00, 0.95, 3.08
[1562409399] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.1% (1636608 kB) free!
[1562409515] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1562409515] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.65, 1.37, 13.29
[1562409654] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 15.3% (7543316 kB) free!
[1562409722] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562409757] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.3% (12716200 kB) free!
[1562409842] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562409962] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562409999] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (634924 kB) free!
[1562410562] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562410854] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (4791132 kB) free!
[1562410957] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (2661360 kB) free!
[1562411162] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562411283] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562411315] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.88, 0.92, 3.42
[1562411403] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562411558] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (14286972 kB) free.
[1562411915] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.79, 6.07, 7.77
[1562412035] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.15, 4.46, 6.98
[1562412050] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.30, 10.31, 5.33
[1562412054] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.9% (6856940 kB) free!
[1562412110] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 16.4% (1209136 kB) free!
[1562412155] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.59, 3.55, 6.34
[1562412158] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.3% (8770544 kB) free!
[1562412170] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.34, 7.11, 4.76
[1562412226] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 5.13, 9.74, 8.96
[1562412230] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 14.1% (1035320 kB) free!
[1562412248] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.44, 7.05, 5.98
[1562412278] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.7% (4412764 kB) free!
[1562412346] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 83.51, 25.19, 14.02
[1562412351] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 11.3% (832812 kB) free!
[1562412368] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.10, 4.81, 5.29
[1562412398] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (274672 kB) free!
[1562412466] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 13.56, 17.75, 12.65
[1562412488] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.15, 3.22, 4.65
[1562412951] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 1.7% (124484 kB) free!
[1562412998] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (9908136 kB) free!
[1562413254] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3910336 kB) free!
[1562413356] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.67, 1.61, 4.16
[1562413370] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 113.23, 69.41, 31.02
[1562413490] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 114.40, 84.37, 41.16
[1562413557] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.41, 3.52, 4.73
[1562413588] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 10.8% (2644304 kB) free!
[1562413610] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1562413610] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 25.84, 64.71, 39.53
[1562413854] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 10.4% (5131992 kB) free!
[1562413956] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.05, 12.43, 8.85
[1562414076] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.62, 8.41, 7.81
[1562414188] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1562414188] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2428136 kB) free!
[1562414196] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.47, 5.77, 6.91
[1562414198] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (20343112 kB) free.
[1562414266] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.90, 3.01, 4.35
[1562414796] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.29, 1.10, 3.83
[1562414810] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.46, 9.89, 17.92
[1562415398] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (10884908 kB) free!
[1562415518] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.2% (20530840 kB) free.
[1562415604] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562416088] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.76, 6.51, 5.51
[1562416208] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.96, 6.97, 5.80
[1562416328] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.54, 5.68, 5.50
[1562416551] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 92.8% (6826816 kB) free.
[1562416610] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.46, 0.67, 3.13
[1562416718] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (12140532 kB) free!
[1562416804] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562416838] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.6% (8291812 kB) free!
[1562416924] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562416928] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.32, 2.35, 3.45
[1562416958] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (19183284 kB) free.
[1562417455] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3712816 kB) free!
[1562417796] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.54, 10.79, 9.56
[1562417916] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.54, 7.39, 8.46
[1562418036] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.65, 5.18, 7.52
[1562418124] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562418128] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.25, 5.90, 5.16
[1562418158] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (9299696 kB) free!
[1562418244] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562418248] SERVICE ALERT: webservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.05, 6.34, 5.41
[1562418279] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (19742792 kB) free.
[1562418368] SERVICE ALERT: webservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.29, 6.35, 5.53
[1562418636] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.48, 1.05, 4.12
[1562418655] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.0% (5927700 kB) free!
[1562418844] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562418965] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562419478] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.6% (10894920 kB) free!
[1562419564] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562419598] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.0% (7204452 kB) free!
[1562419684] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1562419718] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (17924504 kB) free.
[1562420169] SERVICE ALERT: webservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.03, 2.03, 4.15
[1562420284] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562420404] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562420799] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 10.6% (1727232 kB) free!
[1562420810] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.12, 8.48, 9.36
[1562420815] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1562420918] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.7% (9679596 kB) free!
[1562420930] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.49, 5.84, 8.29
[1562420935] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 347 processes
[1562421004] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562421038] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (19263136 kB) free.
[1562421050] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.19, 3.94, 7.30
[1562421055] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1562421124] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562421399] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (815792 kB) free!
[1562421650] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.23, 0.85, 4.02
[1562421655] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3870552 kB) free!
[1562421724] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562421845] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562422066] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 13.48, 9.80, 5.74
[1562422186] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 13.57, 11.05, 6.69
[1562422236] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 67.31, 38.85, 15.35
[1562422238] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.5% (10879708 kB) free!
[1562422306] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 14.02, 12.00, 7.56
[1562422356] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 9.25, 26.07, 13.52
[1562422358] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.0% (7264956 kB) free!
[1562422477] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.41, 17.53, 11.91
[1562422478] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (17582928 kB) free.
[1562422855] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6053260 kB) free!
[1562423044] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562423164] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1562423506] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.37, 2.70, 4.03
[1562423678] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (9895248 kB) free!
[1562423764] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562423798] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (18898852 kB) free.
[1562423884] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562424484] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562424604] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1562424604] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562424877] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.39, 0.72, 4.02
[1562424951] SERVICE FLAPPING ALERT: sig3-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1562424998] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (12105896 kB) free!
[1562425118] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.2% (8689920 kB) free!
[1562425205] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1562425238] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (18333844 kB) free.
[1562425255] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 350 processes
[1562425325] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562425375] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1562425477] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.36, 3.53, 5.16
[1562425588] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562425597] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.51, 2.50, 4.58
[1562425855] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4238672 kB) free!
[1562425925] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562426045] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562426438] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (11851832 kB) free!
[1562426455] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.2% (5498512 kB) free!
[1562426558] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.6% (8947468 kB) free!
[1562426645] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562426678] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562426678] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.4% (18659620 kB) free.
[1562426765] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562427050] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 10.23, 85.28, 52.21
[1562427170] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.40, 57.08, 45.88
[1562427290] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.34, 38.26, 40.34
[1562427365] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 341 processes
[1562427485] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1562427878] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (11822620 kB) free!
[1562427998] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.5% (8849056 kB) free!
[1562428086] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562428118] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (18473732 kB) free.
[1562428206] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1562428306] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 3.37, 86.97, 93.29
[1562428375] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1562428426] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 3.05, 59.17, 82.33
[1562428490] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.23, 0.94, 11.29
[1562428495] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1562428546] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.66, 40.40, 72.64
[1562428806] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1562428926] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1562429318] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (12548412 kB) free!
[1562429438] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.1% (9948076 kB) free!
[1562429455] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (3693500 kB) free!
[1562429558] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (18532760 kB) free.
[1562429690] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.22, 0.39, 3.35
[1562429797] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1323.43, 589.33, 219.47
[1562429917] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 178.79, 394.45, 192.90
[1562430037] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 24.50, 264.09, 169.57
[1562430126] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562430246] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562430346] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.59, 2.70, 12.67
[1562430369] SERVICE FLAPPING ALERT: webservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562430655] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.1% (5961100 kB) free!
[1562430758] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (9583232 kB) free!
[1562430846] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1562430878] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.7% (18215272 kB) free.
[1562430966] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1562431567] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 340 processes
[1562431687] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562432146] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.62, 2.89, 4.39
[1562432287] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562432407] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562432437] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.41, 0.62, 13.23
[1562432499] SERVICE ALERT: sigqscluster1;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1562432799] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 11.7% (1894160 kB) free!
[1562433007] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1562433127] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1562433278] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.1% (9241980 kB) free!
[1562433398] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (17571056 kB) free.
[1562433637] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.42, 53.58, 66.67
[1562433655] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (3836524 kB) free!
[1562433998] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (11552040 kB) free!
[1562434118] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (18423840 kB) free.
[1562434328] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1562434448] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1562434718] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (12826348 kB) free!
[1562434837] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.41, 1.44, 18.62
[1562434838] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (8996412 kB) free!
[1562434855] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.1% (5981076 kB) free!
[1562434958] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (17150400 kB) free.
[1562435048] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1562435168] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1562435558] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.9% (10478256 kB) free!
[1562435678] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (18312168 kB) free.
[1562435768] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1562435888] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1562436278] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (12462268 kB) free!
[1562436399] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.7% (9029332 kB) free!
[1562436488] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1562436519] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.8% (16978544 kB) free.
[1562436608] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1562436637] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.51, 0.49, 3.07
[1562436895] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1562437015] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1562437119] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (9709756 kB) free!
[1562437209] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 342 processes
[1562437239] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (17514160 kB) free.
[1562437255] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (2957968 kB) free!
[1562437329] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes