[1555354030] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53827 MB (23% inode=99%):
[1555354630] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46659 MB (20% inode=99%):
[1555354750] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42553 MB (18% inode=99%):
[1555354870] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42553 MB (18% inode=99%):
[1555356071] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54848 MB (24% inode=99%):
[1555356671] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46656 MB (20% inode=99%):
[1555356791] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42552 MB (18% inode=99%):
[1555356911] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42552 MB (18% inode=99%):
[1555357511] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47319 MB (21% inode=99%):
[1555358711] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46655 MB (20% inode=99%):
[1555358831] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42557 MB (18% inode=99%):
[1555358951] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42557 MB (18% inode=99%):
[1555360151] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54782 MB (24% inode=99%):
[1555360751] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46494 MB (20% inode=99%):
[1555360871] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 38284 MB (17% inode=99%):
[1555360991] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 38284 MB (17% inode=99%):
[1555362225] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 315 processes
[1555362826] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1555362946] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 322 processes
[1555363066] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555363391] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52733 MB (23% inode=99%):
[1555363992] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45053 MB (20% inode=99%):
[1555364112] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43004 MB (19% inode=99%):
[1555364232] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54781 MB (24% inode=99%):
[1555364266] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1555364832] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44637 MB (19% inode=99%):
[1555364866] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 323 processes
[1555364952] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44637 MB (19% inode=99%):
[1555364987] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 325 processes
[1555365072] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 36427 MB (16% inode=99%):
[1555365107] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1555365107] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1555366272] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54525 MB (24% inode=99%):
[1555366872] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42486 MB (18% inode=99%):
[1555366992] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42486 MB (18% inode=99%):
[1555367112] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42493 MB (18% inode=99%):
[1555368312] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53756 MB (23% inode=99%):
[1555368912] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42415 MB (18% inode=99%):
[1555369032] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42415 MB (18% inode=99%):
[1555369152] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42422 MB (18% inode=99%):
[1555370352] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52662 MB (23% inode=99%):
[1555370952] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42415 MB (18% inode=99%):
[1555371072] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42422 MB (18% inode=99%):
[1555371192] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54710 MB (24% inode=99%):
[1555371792] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46518 MB (20% inode=99%):
[1555371912] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44982 MB (20% inode=99%):
[1555372032] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42966 MB (19% inode=99%):
[1555372372] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 136.07, 319.76, 196.60
[1555372492] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 18.54, 214.05, 172.79
[1555372514] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 504.62, 344.28, 155.57
[1555372612] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 3.05, 143.41, 151.91
[1555372634] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 394.18, 373.13, 189.99
[1555372754] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 53.56, 249.80, 167.00
[1555372764] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.2% (2758064 kB) free!
[1555372895] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.2% (2754224 kB) free!
[1555373232] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53685 MB (23% inode=99%):
[1555373832] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45845 MB (20% inode=99%):
[1555373952] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42772 MB (19% inode=99%):
[1555374073] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42773 MB (19% inode=99%):
[1555375012] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.66, 0.53, 11.88
[1555375154] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.55, 0.53, 12.98
[1555375273] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54709 MB (24% inode=99%):
[1555375295] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2445580 kB) free!
[1555375307] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 316 processes
[1555375764] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (2255296 kB) free!
[1555375873] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47124 MB (20% inode=99%):
[1555375907] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1555375993] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47124 MB (20% inode=99%):
[1555376027] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 312 processes
[1555376113] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45589 MB (20% inode=99%):
[1555376148] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555376212] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 512.36, 402.65, 203.62
[1555376354] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.33, 0.46, 3.99
[1555376364] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.9% (2932496 kB) free!
[1555376713] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54707 MB (24% inode=99%):
[1555376785] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.8% (3210916 kB) free!
[1555376905] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 19.5% (3175840 kB) free!
[1555377025] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 19.3% (3136204 kB) free!
[1555377313] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46515 MB (20% inode=99%):
[1555377433] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46515 MB (20% inode=99%):
[1555377553] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42412 MB (18% inode=99%):
[1555378153] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52654 MB (23% inode=99%):
[1555378612] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1555378612] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.47, 0.64, 16.40
[1555378753] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46510 MB (20% inode=99%):
[1555378873] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44974 MB (20% inode=99%):
[1555378993] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44902 MB (19% inode=99%):
[1555379593] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54110 MB (24% inode=99%):
[1555379812] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.36, 0.47, 4.86
[1555380194] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46430 MB (20% inode=99%):
[1555380314] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42327 MB (18% inode=99%):
[1555380434] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42327 MB (18% inode=99%):
[1555381034] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 51668 MB (22% inode=99%):
[1555381164] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1555381164] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (1556928 kB) free!
[1555381634] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44500 MB (19% inode=99%):
[1555381754] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42452 MB (18% inode=99%):
[1555381874] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54603 MB (24% inode=99%):
[1555382474] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46412 MB (20% inode=99%):
[1555382594] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42309 MB (18% inode=99%):
[1555382714] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42309 MB (18% inode=99%):
[1555383625] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (1571244 kB) free!
[1555383949] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 324 processes
[1555384549] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 323 processes
[1555384669] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 319 processes
[1555384789] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555385114] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53385 MB (23% inode=99%):
[1555385714] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45705 MB (20% inode=99%):
[1555385834] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44681 MB (19% inode=99%):
[1555385954] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44681 MB (19% inode=99%):
[1555386554] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50310 MB (22% inode=99%):
[1555387155] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42112 MB (18% inode=99%):
[1555387275] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42112 MB (18% inode=99%):
[1555387395] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 39903 MB (17% inode=99%):
[1555388364] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1555388595] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54399 MB (24% inode=99%):
[1555388991] SERVICE ALERT: secondary-server;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 12 users currently logged in
[1555389195] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46207 MB (20% inode=99%):
[1555389315] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42104 MB (18% inode=99%):
[1555389435] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42104 MB (18% inode=99%):
[1555390259] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555390320] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555390379] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555390440] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555390499] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1555390561] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555390612] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1555391235] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50299 MB (22% inode=99%):
[1555391515] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.35, 8.82, 9.52
[1555391635] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.05, 6.62, 8.63
[1555391755] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.68, 5.01, 7.80
[1555391835] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42100 MB (18% inode=99%):
[1555391955] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 42107 MB (18% inode=99%):
[1555392075] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54393 MB (24% inode=99%):
[1555392675] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46401 MB (20% inode=99%):
[1555392795] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46401 MB (20% inode=99%):
[1555392916] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45633 MB (20% inode=99%):
[1555392955] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.53, 1.91, 3.56
[1555393516] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54527 MB (24% inode=99%):
[1555394116] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46335 MB (20% inode=99%):
[1555394236] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46335 MB (20% inode=99%):
[1555394356] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43740 MB (19% inode=99%):
[1555394425] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 83.0% (13489784 kB) free.
[1555394956] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54536 MB (24% inode=99%):
[1555394989] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1555395299] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555395336] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.0% (12306928 kB) free.
[1555395387] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.6% (20619260 kB) free.
[1555395430] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.9% (22761636 kB) free.
[1555395474] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.7% (15064484 kB) free.
[1555395544] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.9% (22518348 kB) free.
[1555395556] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47137 MB (20% inode=99%):
[1555395586] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.8% (22989808 kB) free.
[1555395590] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 318 processes
[1555395676] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43036 MB (19% inode=99%):
[1555395710] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555395796] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43036 MB (19% inode=99%):
[1555395830] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 317 processes
[1555395943] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 10.2% (5010684 kB) free!
[1555396063] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (3005748 kB) free!
[1555396184] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (2386864 kB) free!
[1555396396] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49366 MB (21% inode=99%):
[1555396499] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555396619] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555396683] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (157724 kB) free!
[1555396739] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555396753] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.1% (3697408 kB) free!
[1555396803] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (154028 kB) free!
[1555396850] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 340 processes
[1555396869] SERVICE ALERT: sservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50646 MB (22% inode=99%):
[1555396873] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.0% (3677508 kB) free!
[1555396923] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (172888 kB) free!
[1555396993] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.9% (3654064 kB) free!
[1555396996] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43214 MB (19% inode=99%):
[1555397116] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43221 MB (19% inode=99%):
[1555397237] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43205 MB (19% inode=99%):
[1555397339] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1555397339] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1555397340] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47619 MB (21% inode=99%):
[1555397524] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.1% (14313148 kB) free.
[1555397940] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47139 MB (20% inode=99%):
[1555398060] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47073 MB (20% inode=99%):
[1555398164] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.3% (11996084 kB) free.
[1555398181] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 47040 MB (20% inode=99%):
[1555398188] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.0% (22308056 kB) free.
[1555398361] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555398771] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (3322624 kB) free!
[1555398891] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.9% (1407160 kB) free!
[1555399012] SERVICE FLAPPING ALERT: db2;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1555399012] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (543820 kB) free!
[1555399037] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 54495 MB (24% inode=99%):
[1555399139] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555399259] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555399379] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555399555] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.86, 11.90, 6.31
[1555399637] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46303 MB (20% inode=99%):
[1555399676] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.06, 8.61, 5.78
[1555399718] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 25.1% (12367128 kB) free.
[1555399757] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46303 MB (20% inode=99%):
[1555399768] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.52, 6.20, 5.03
[1555399796] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.43, 6.50, 5.35
[1555399799] SERVICE ALERT: sservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 324 processes
[1555399877] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46303 MB (20% inode=99%):
[1555399888] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.19, 4.20, 4.44
[1555400212] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.31, 17.32, 8.60
[1555400327] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 12.9% (6373180 kB) free!
[1555400332] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.20, 11.93, 7.68
[1555400396] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1555400396] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 2.27, 2.68, 3.78
[1555400447] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 12.0% (5931564 kB) free!
[1555400452] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 1.18, 8.34, 6.88
[1555400477] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53422 MB (23% inode=99%):
[1555400567] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (4072264 kB) free!
[1555400792] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.0% (2725504 kB) free!
[1555400913] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.0% (2709224 kB) free!
[1555400999] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 322 processes
[1555401033] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.9% (2690760 kB) free!
[1555401052] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.58, 1.76, 3.97
[1555401077] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46252 MB (20% inode=99%):
[1555401129] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 584 processes
[1555401189] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 590 processes
[1555401197] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44718 MB (19% inode=99%):
[1555401317] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44718 MB (19% inode=99%):
[1555401374] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 23.8% (11731876 kB) free.
[1555401429] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.4% (20817624 kB) free.
[1555401984] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (3006848 kB) free!
[1555402104] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.4% (1204520 kB) free!
[1555402224] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (219608 kB) free!
[1555402850] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 338 processes
[1555405517] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 50247 MB (22% inode=99%):
[1555406117] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46151 MB (20% inode=99%):
[1555406237] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46151 MB (20% inode=99%):
[1555406357] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 42048 MB (18% inode=99%):
[1555407230] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1555407247] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 24.8% (12242724 kB) free.
[1555407557] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48895 MB (21% inode=99%):
[1555407830] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1555407853] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.3% (3124188 kB) free!
[1555407950] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 322 processes
[1555407973] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.8% (1863916 kB) free!
[1555408070] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 316 processes
[1555408094] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (221752 kB) free!
[1555408157] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41727 MB (18% inode=99%):
[1555408226] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.7% (2870528 kB) free!
[1555408277] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39678 MB (17% inode=99%):
[1555408347] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (161072 kB) free!
[1555408392] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.3% (3505812 kB) free!
[1555408397] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53779 MB (23% inode=99%):
[1555408467] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (149848 kB) free!
[1555408512] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180228 kB) free!
[1555408632] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182616 kB) free!
[1555408633] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.9% (4903172 kB) free!
[1555408753] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (4853472 kB) free!
[1555408873] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 19.4% (4798632 kB) free!
[1555408979] SERVICE FLAPPING ALERT: sservercluster1;Current Users;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1555408997] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45843 MB (20% inode=99%):
[1555409117] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45843 MB (20% inode=99%):
[1555409237] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 37626 MB (16% inode=99%):
[1555409579] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555409667] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 71.0% (11546228 kB) free.
[1555409699] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555409819] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555410267] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (175096 kB) free!
[1555410387] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (174104 kB) free!
[1555410507] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (158984 kB) free!
[1555410673] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1555410673] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173636 kB) free!
[1555410844] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 22.6% (11125204 kB) free.
[1555411269] SERVICE ALERT: sservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46309 MB (20% inode=99%):
[1555411389] SERVICE ALERT: sservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46308 MB (20% inode=99%):
[1555411456] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.4% (6611472 kB) free!
[1555411509] SERVICE ALERT: sservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46228 MB (20% inode=99%):
[1555411576] SERVICE ALERT: db2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (4583336 kB) free!
[1555411619] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555411642] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1555411642] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 35.9% (5827536 kB) free.
[1555411696] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (2477780 kB) free!
[1555411796] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1555412201] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1555412242] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (169148 kB) free!
[1555412362] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (161548 kB) free!
[1555412482] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (156496 kB) free!
[1555414096] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 19.5% (9617060 kB) free!
[1555414282] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 75.4% (12249268 kB) free.
[1555414324] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.1% (2938548 kB) free!
[1555414444] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.6% (2858048 kB) free!
[1555414564] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.1% (2776256 kB) free!
[1555414637] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49850 MB (22% inode=99%):
[1555414696] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 4.3% (2109856 kB) free!
[1555415237] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41651 MB (18% inode=99%):
[1555415357] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41651 MB (18% inode=99%):
[1555415477] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41651 MB (18% inode=99%):
[1555416077] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52914 MB (23% inode=99%):
[1555416677] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 43860 MB (19% inode=99%):
[1555416797] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43860 MB (19% inode=99%):
[1555416917] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43866 MB (19% inode=99%):
[1555416964] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (1336716 kB) free!
[1555418117] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49876 MB (22% inode=99%):
[1555418296] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 13.5% (6649084 kB) free!
[1555418717] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 42355 MB (18% inode=99%):
[1555418838] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53962 MB (24% inode=99%):
[1555418844] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1555418870] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1555418896] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 1.4% (697236 kB) free!
[1555418965] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1555419085] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1555419361] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1555419437] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45866 MB (20% inode=99%):
[1555419470] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 320 processes
[1555419481] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1555419557] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45866 MB (20% inode=99%):
[1555419590] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 321 processes
[1555419595] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177404 kB) free!
[1555419601] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1555419677] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41761 MB (18% inode=99%):
[1555419710] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 327 processes
[1555419715] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189432 kB) free!
[1555419835] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192240 kB) free!
[1555420089] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 605 processes
[1555420149] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 603 processes
[1555420509] SERVICE ALERT: sservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 97345 MB (44% inode=99%):
[1555420754] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 112.97, 306.26, 193.64
[1555420801] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1555420874] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15.68, 205.06, 170.21
[1555420877] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49859 MB (22% inode=99%):
[1555420994] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.52, 137.37, 149.63
[1555421035] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 64.4% (15774228 kB) free.
[1555421477] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41916 MB (18% inode=99%):
[1555421597] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 41923 MB (18% inode=99%):
[1555421717] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 39873 MB (17% inode=99%):
[1555423394] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.18, 0.48, 11.73
[1555423791] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 8 users currently logged in
[1555424482] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1555424594] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.71, 0.61, 3.61
[1555425085] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 10 users currently logged in
[1555425496] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 12.3% (6050656 kB) free!
[1555425688] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.09, 9.57, 6.24
[1555425808] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.48, 6.54, 5.54
[1555425928] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 10.20, 8.55, 6.38
[1555426097] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (236248 kB) free!
[1555426222] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8777804 kB) free!
[1555426342] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 32.0% (15788812 kB) free.
[1555426528] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.32, 1.39, 3.48
[1555426942] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9673144 kB) free!
[1555427062] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 34.3% (16925204 kB) free.
[1555427717] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 48672 MB (21% inode=99%):
[1555428317] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44576 MB (19% inode=99%):
[1555428437] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 44576 MB (19% inode=99%):
[1555428557] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 44576 MB (19% inode=99%):
[1555428862] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9791308 kB) free!
[1555428982] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 14.2% (6982524 kB) free!
[1555429103] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 29.7% (14649324 kB) free.
[1555429310] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 321 processes
[1555429702] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9174792 kB) free!
[1555429822] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 31.9% (15719708 kB) free.
[1555429910] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 320 processes
[1555430030] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 319 processes
[1555430150] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 324 processes
[1555430957] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 53642 MB (23% inode=99%):
[1555430996] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 43.66, 16.13, 6.86
[1555431116] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.28, 11.35, 6.24
[1555431236] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.36, 8.15, 5.69
[1555431557] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45450 MB (20% inode=99%):
[1555431622] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9765044 kB) free!
[1555431677] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45450 MB (20% inode=99%):
[1555431742] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.5% (6646896 kB) free!
[1555431797] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41347 MB (18% inode=99%):
[1555431836] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.24, 2.27, 3.66
[1555431862] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 28.0% (13803396 kB) free.
[1555432462] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 16.4% (8083408 kB) free!
[1555432582] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 30.6% (15078396 kB) free.
[1555433182] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9793068 kB) free!
[1555433302] SERVICE ALERT: db1;Memory;WARNING;SOFT;2;WARNING - 13.7% (6744888 kB) free!
[1555433423] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.9% (13748312 kB) free.
[1555434022] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9044392 kB) free!
[1555434142] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 30.2% (14874384 kB) free.
[1555434797] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52414 MB (23% inode=99%):
[1555435397] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45246 MB (20% inode=99%):
[1555435517] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43710 MB (19% inode=99%):
[1555435637] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 43710 MB (19% inode=99%):
[1555436237] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49340 MB (21% inode=99%):
[1555436542] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 14.9% (7318756 kB) free!
[1555436662] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 27.2% (13424920 kB) free.
[1555436835] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4867616 kB) free!
[1555436837] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 41141 MB (18% inode=99%):
[1555436955] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (4848040 kB) free!
[1555436958] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 39099 MB (17% inode=99%):
[1555437075] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (4837884 kB) free!
[1555437078] SERVICE ALERT: db2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 52923 MB (23% inode=99%):
[1555437262] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 17.3% (8508816 kB) free!
[1555437382] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 29.7% (14651284 kB) free.
[1555437497] SERVICE FLAPPING ALERT: db2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1555437678] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 44731 MB (19% inode=99%):
[1555437798] SERVICE ALERT: db2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 43707 MB (19% inode=99%):
[1555437918] SERVICE ALERT: db2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 41657 MB (18% inode=99%):
[1555437982] SERVICE ALERT: db1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9389768 kB) free!
[1555438102] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 29.7% (14628556 kB) free.
[1555439150] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 324 processes