[1546021829] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1546021887] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9970996 kB) free.
[1546021949] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 293 processes
[1546022192] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1546022312] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 332 processes
[1546022432] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1546023087] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 14.7% (7253824 kB) free!
[1546023207] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9969380 kB) free.
[1546024374] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 456 processes
[1546024434] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 456 processes
[1546024476] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1546025754] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 11 users currently logged in
[1546028432] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1546028552] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 333 processes
[1546028672] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 333 processes
[1546038554] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1.59, 72.44, 60.55
[1546038674] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.41, 48.56, 53.24
[1546038794] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.18, 32.55, 46.80
[1546039994] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 7.49, 2.93, 13.98
[1546041194] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.20, 0.66, 4.48
[1546055594] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 611.04, 248.04, 90.54
[1546055714] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 83.05, 166.16, 79.63
[1546055719] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 606 processes
[1546055834] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 11.27, 111.22, 69.99
[1546055839] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 606 processes
[1546055959] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 349 processes
[1546057034] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.35, 2.29, 19.42
[1546057949] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 295 processes
[1546058008] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 18.8% (9287616 kB) free!
[1546058070] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1546058128] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.3% (8050972 kB) free!
[1546058190] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546058248] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9968608 kB) free.
[1546058790] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1546058834] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.98, 0.85, 3.42
[1546058848] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 19.6% (9648736 kB) free!
[1546058968] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8222952 kB) free!
[1546059088] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9965012 kB) free.
[1546059234] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 429 processes
[1546059273] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 335 processes
[1546059294] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 433 processes
[1546059354] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;3;CPU WARNING: 2 warn out of 435 processes
[1546059390] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546059414] SERVICE ALERT: webserver2;CPU_Procs;WARNING;HARD;4;CPU WARNING: 2 warn out of 431 processes
[1546059510] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546059630] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1546059688] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 13.9% (6860832 kB) free!
[1546059808] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9957292 kB) free.
[1546060614] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 430 processes
[1546060830] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1546060830] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 288 processes
[1546061073] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 332 processes
[1546064430] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546064550] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 292 processes
[1546065954] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1546066074] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1546066194] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1546067073] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 336 processes
[1546067550] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546067608] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 16.0% (7863956 kB) free!
[1546067670] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1546067728] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9948524 kB) free.
[1546067790] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546070073] SERVICE FLAPPING ALERT: smscapp2;CPU_Procs;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1546070073] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 335 processes
[1546071990] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1546074328] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9858268 kB) free!
[1546074390] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546074448] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 17.4% (8579116 kB) free!
[1546074510] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546074569] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 14.2% (7009752 kB) free!
[1546074631] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546074926] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1546075046] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1546075169] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9946416 kB) free.
[1546075510] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.2% (3119960 kB) free!
[1546075630] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 20.5% (3325720 kB) free.
[1546075769] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 15.3% (7519976 kB) free!
[1546075889] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 20.2% (9944364 kB) free.
[1546076830] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.9% (3069692 kB) free!
[1546076950] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.9% (3069776 kB) free!
[1546077070] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.9% (3071108 kB) free!
[1546077670] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 20.3% (3304184 kB) free.
[1546079073] SERVICE FLAPPING ALERT: smscapp2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1546080070] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.6% (3028456 kB) free!
[1546080190] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.6% (3029648 kB) free!
[1546080311] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.6% (3028804 kB) free!
[1546080434] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 5.83, 80.79, 53.92
[1546080554] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1.13, 54.18, 47.43
[1546080674] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.42, 36.34, 41.71
[1546080911] SERVICE FLAPPING ALERT: smscapp1;Memory;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1546080911] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 20.4% (3310148 kB) free.
[1546081107] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 360 processes
[1546081874] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.26, 0.94, 11.66
[1546082111] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.3% (3139380 kB) free!
[1546082231] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.3% (3137956 kB) free!
[1546082351] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.3% (3138292 kB) free!
[1546082394] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1546082907] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 359 processes
[1546083074] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.39, 0.47, 3.51
[1546084794] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 17 users currently logged in
[1546084914] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 17 users currently logged in
[1546085034] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 17 users currently logged in
[1546086030] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546086089] SERVICE ALERT: db2;Memory;WARNING;SOFT;1;WARNING - 20.0% (9837120 kB) free!
[1546086150] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546086209] SERVICE ALERT: db2;Memory;WARNING;SOFT;2;WARNING - 16.8% (8261520 kB) free!
[1546086239] SERVICE ALERT: sig1;Memory;WARNING;HARD;3;WARNING - 18.2% (2951512 kB) free!
[1546086269] SERVICE ALERT: sig4;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 300 processes
[1546086270] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 290 processes
[1546086329] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 21.6% (10660708 kB) free.
[1546086389] SERVICE ALERT: sig4;CPU_Procs;OK;HARD;1;CPU OK: 298 processes
[1546093151] SERVICE FLAPPING ALERT: smscapp1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1546094474] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 0.75, 41.87, 47.65
[1546094594] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 0.45, 28.18, 41.94
[1546094670] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546094714] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 0.30, 18.94, 36.88
[1546094790] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546094907] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 368 processes
[1546094910] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1546095314] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.60, 3.36, 19.76
[1546095510] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1546095835] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 14 users currently logged in
[1546096107] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 357 processes
[1546097114] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.77, 1.93, 4.45
[1546097307] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 357 processes
[1546097907] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1546097907] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 355 processes
[1546103521] SERVICE ALERT: smscapp1;CPU_Procs;OK;HARD;1;CPU OK: 294 processes
[1546104273] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 328 processes
[1546104551] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 50.1% (8140916 kB) free.
[1546104702] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 5.50, 6.87, 8.65
[1546105710] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1546105830] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1546105921] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 1 warn out of 296 processes
[1546106035] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1546106041] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 1 warn out of 295 processes
[1546106155] SERVICE ALERT: pservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 16 users currently logged in
[1546106161] SERVICE ALERT: smscapp1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 1 warn out of 296 processes
[1546106275] SERVICE ALERT: pservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 16 users currently logged in
[1546106502] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 13.51, 13.05, 10.89
[1546106622] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 13.06, 13.10, 11.18
[1546106742] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 13.57, 13.35, 11.51
[1546107507] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 357 processes