[1582050759] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23482192 kB) free!
[1582050879] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.1% (19839852 kB) free!
[1582050999] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.4% (16399792 kB) free!
[1582051055] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171760 kB) free!
[1582051078] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 58.5% (28829412 kB) free.
[1582051116] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198944 kB) free!
[1582051175] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179520 kB) free!
[1582051237] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191540 kB) free!
[1582051296] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183348 kB) free!
[1582051357] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186488 kB) free!
[1582052199] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35553408 kB) free.
[1582052556] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (2382452 kB) free!
[1582052620] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167672 kB) free!
[1582052681] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (175088 kB) free!
[1582052741] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.7% (176276 kB) free!
[1582053041] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 69.3% (16982284 kB) free.
[1582053267] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 26.8% (6615200 kB) free.
[1582053399] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24471728 kB) free!
[1582053519] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.8% (20888820 kB) free!
[1582053639] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.4% (17611040 kB) free!
[1582053867] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187120 kB) free!
[1582053987] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176880 kB) free!
[1582054107] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188132 kB) free!
[1582054214] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(126) < 92.42.104.166 [19/Feb/2020:00:55:17 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1582054334] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582054839] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32430272 kB) free.
[1582054856] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9389608 kB) free!
[1582054976] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9048532 kB) free!
[1582055096] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.7% (8705552 kB) free!
[1582055307] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 10.3% (2532228 kB) free!
[1582055696] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4627260 kB) free!
[1582055907] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182832 kB) free!
[1582056040] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.3% (21469388 kB) free!
[1582056160] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.1% (46246852 kB) free.
[1582056760] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (25019072 kB) free!
[1582056880] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.2% (21346076 kB) free!
[1582057000] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.0% (46148820 kB) free.
[1582057600] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24545740 kB) free!
[1582057720] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.0% (21091232 kB) free!
[1582057840] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.0% (46099888 kB) free.
[1582058440] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24132360 kB) free!
[1582058560] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.8% (20788432 kB) free!
[1582058680] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.7% (45739708 kB) free.
[1582059280] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23515444 kB) free!
[1582059400] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.0% (22421124 kB) free!
[1582059520] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41499324 kB) free.
[1582059629] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (239104 kB) free!
[1582059750] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (231560 kB) free!
[1582059870] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (236648 kB) free!
[1582060086] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (231004 kB) free!
[1582060120] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15545332 kB) free!
[1582060206] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (238208 kB) free!
[1582060240] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.6% (45571504 kB) free.
[1582060301] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.4% (10135440 kB) free.
[1582060326] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (213356 kB) free!
[1582060470] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 77.7% (38280056 kB) free.
[1582060841] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (19009076 kB) free!
[1582060926] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 83.2% (41023236 kB) free.
[1582060961] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.7% (15423016 kB) free!
[1582061081] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.6% (45572108 kB) free.
[1582061535] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2842) < 203.212.70.170 [19/Feb/2020:03:01:41 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1582061655] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(16) < 49.50.64.122 [19/Feb/2020:03:02:31 +0530] TCP 502 0 0 120.000 "10.147.212.154:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1582061681] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20292784 kB) free!
[1582061684] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(225) < 103.16.101.72 [19/Feb/2020:03:00:21 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1582061775] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582061801] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (15668384 kB) free!
[1582061806] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582061921] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.6% (44291784 kB) free.
[1582061995] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(280) < 2020/02/19 03:01:32 [error] 12255#0: *56241101 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582062115] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582062521] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20185428 kB) free!
[1582062641] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.6% (45582804 kB) free.
[1582063241] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24066848 kB) free!
[1582063361] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.7% (20658568 kB) free!
[1582063481] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.9% (46060884 kB) free.
[1582064081] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25893668 kB) free!
[1582064201] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.9% (47278212 kB) free.
[1582067307] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582067801] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1582068702] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171948 kB) free!
[1582068822] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193140 kB) free!
[1582068942] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (148464 kB) free!
[1582069958] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 83.7% (20505424 kB) free.
[1582070308] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 51.1% (12620248 kB) free.
[1582071929] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1582072050] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1582073577] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 27.1% (35778184 kB) free.
[1582075222] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.3% (39596152 kB) free.
[1582075249] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.5% (21927756 kB) free.
[1582076908] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199828 kB) free!
[1582077029] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198856 kB) free!
[1582077149] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194112 kB) free!
[1582079523] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.0% (21967924 kB) free.
[1582080115] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(41) < 2020/02/19 08:11:45 [error] 12257#0: *56325560 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582080235] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(2) < 2020/02/19 08:12:06 [error] 12256#0: *56325604 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1582080355] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582088801] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.2% (22736292 kB) free!
[1582088921] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.5% (45434524 kB) free.
[1582089521] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15189936 kB) free!
[1582089627] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9681868 kB) free!
[1582089641] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.4% (11068916 kB) free!
[1582089747] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.6% (9681828 kB) free!
[1582089761] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.8% (43202356 kB) free.
[1582089867] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (9682556 kB) free!
[1582090328] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.2% (3500404 kB) free!
[1582090361] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.9% (18341356 kB) free!
[1582090448] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.7% (3368764 kB) free!
[1582090481] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.7% (14136752 kB) free!
[1582090568] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 12.7% (3133448 kB) free!
[1582090601] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40275572 kB) free.
[1582090869] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4334916 kB) free!
[1582090989] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 17.6% (3941540 kB) free!
[1582091067] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4347820 kB) free!
[1582091109] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.7% (3514556 kB) free!
[1582091168] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (990232 kB) free!
[1582091201] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14758664 kB) free!
[1582091322] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (10634300 kB) free!
[1582091401] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1582091442] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37571380 kB) free.
[1582091462] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1582091521] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1582091571] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1582091582] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1582091642] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1582091667] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 10.4% (5142816 kB) free!
[1582091683] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582091691] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1582091702] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1582091709] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1566620 kB) free!
[1582091803] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582091923] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582091997] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170992 kB) free!
[1582092042] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12319572 kB) free!
[1582092117] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195928 kB) free!
[1582092160] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 10.6% (2601908 kB) free!
[1582092162] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (8158708 kB) free!
[1582092237] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176740 kB) free!
[1582092267] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1582092267] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (226564 kB) free!
[1582092280] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (207516 kB) free!
[1582092282] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36180120 kB) free.
[1582092400] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (183944 kB) free!
[1582092445] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 15.7% (3840436 kB) free!
[1582092565] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 14.1% (3465896 kB) free!
[1582092685] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 14.0% (3425764 kB) free!
[1582092842] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1582092882] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.0% (10502324 kB) free!
[1582093002] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6276308 kB) free!
[1582093078] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 15.8% (7776636 kB) free!
[1582093122] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35004612 kB) free.
[1582093198] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 14.8% (7280736 kB) free!
[1582093257] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.4% (4263192 kB) free!
[1582093285] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 69.7% (17086516 kB) free.
[1582093318] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 13.7% (6766776 kB) free!
[1582093378] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.5% (4032696 kB) free!
[1582093498] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.6% (3823340 kB) free!
[1582093629] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 14.5% (3582528 kB) free!
[1582093722] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10170688 kB) free!
[1582093749] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 12.5% (3073552 kB) free!
[1582093842] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6174472 kB) free!
[1582093869] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 10.7% (2639852 kB) free!
[1582093901] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1582093919] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 61.3% (30231800 kB) free.
[1582093962] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.9% (32877700 kB) free.
[1582094021] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 697 processes
[1582094098] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (2090240 kB) free!
[1582094110] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.5% (16059772 kB) free.
[1582094155] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(39) < 2020/02/19 12:00:58 [error] 12249#0: *56423804 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582094175] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 172.105.59.173 [19/Feb/2020:12:02:18 +0530] TCP 502 0 0 60.001 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.001, 0.000"
[1582094275] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582094295] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582094469] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (811060 kB) free!
[1582094557] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.7% (21495760 kB) free.
[1582094562] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9697488 kB) free!
[1582094591] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.0% (21231892 kB) free.
[1582094638] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 51.1% (67313448 kB) free.
[1582094692] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.3% (39070948 kB) free.
[1582094753] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 88.2% (21608204 kB) free.
[1582094783] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.0% (21321060 kB) free.
[1582094901] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.6% (22849652 kB) free.
[1582095197] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190532 kB) free!
[1582095303] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1582095317] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182864 kB) free!
[1582095438] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1582095438] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195276 kB) free!
[1582095523] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1582095898] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (822244 kB) free!
[1582095903] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1582096018] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (527832 kB) free!
[1582096023] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1582096138] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (229936 kB) free!
[1582096143] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 343 processes
[1582096326] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.8% (8785824 kB) free!
[1582096368] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.6% (4551588 kB) free!
[1582096446] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.1% (8410260 kB) free!
[1582096488] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 71.7% (17578968 kB) free.
[1582096566] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.2% (7990944 kB) free!
[1582096586] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (1443820 kB) free!
[1582096706] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (1344260 kB) free!
[1582096826] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.1% (1250892 kB) free!
[1582097766] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4537836 kB) free!
[1582097895] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(443) < 103.255.146.164 [19/Feb/2020:13:01:53 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1582097923] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1582098015] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582098043] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1582098163] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1582098245] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11377948 kB) free!
[1582098365] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7501952 kB) free!
[1582098485] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37517796 kB) free.
[1582098568] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.6% (2853532 kB) free!
[1582098688] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 34.4% (5597136 kB) free.
[1582099085] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12264164 kB) free!
[1582099205] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (8025516 kB) free!
[1582099325] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.1% (4081440 kB) free!
[1582099925] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 10.1% (13321232 kB) free!
[1582100254] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.1% (21491344 kB) free.
[1582100283] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.5% (39673736 kB) free.
[1582100304] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 30.6% (15092580 kB) free.
[1582100339] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1582100339] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.0% (21326296 kB) free.
[1582100390] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 59.3% (78158776 kB) free.
[1582100946] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187492 kB) free!
[1582101066] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179060 kB) free!
[1582101186] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192884 kB) free!
[1582101910] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 14.9% (3336536 kB) free!
[1582102031] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 13.7% (3072980 kB) free!
[1582102056] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192524 kB) free!
[1582102151] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.5% (2813292 kB) free!
[1582102176] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176584 kB) free!
[1582102296] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202076 kB) free!
[1582102470] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9307288 kB) free!
[1582102557] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166396 kB) free!
[1582102590] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.6% (9155644 kB) free!
[1582102677] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187452 kB) free!
[1582102710] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (9015824 kB) free!
[1582102751] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (24.4% change >= 20.0% threshold)
[1582102751] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1563352 kB) free!
[1582102797] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (182220 kB) free!
[1582102896] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 20.8% (5127936 kB) free.
[1582103022] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1582103142] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1582103262] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1582103275] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/19 14:37:01 [error] 12257#0: *56512895 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1582103396] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582103497] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (2012696 kB) free!
[1582103542] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1582103617] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (1523840 kB) free!
[1582103662] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1582103737] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.5% (1355292 kB) free!
[1582103782] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1582103893] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.7% (21396448 kB) free.
[1582103939] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 77.7% (38282360 kB) free.
[1582103997] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13042620 kB) free!
[1582103998] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1582103998] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 87.7% (21494472 kB) free.
[1582104056] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 78.2% (19166172 kB) free.
[1582104117] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8552060 kB) free!
[1582104119] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9295432 kB) free!
[1582104133] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 57.5% (75830892 kB) free.
[1582104163] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 361 processes
[1582104221] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.6% (16068392 kB) free.
[1582104239] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.2% (8976228 kB) free!
[1582104359] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.4% (8584064 kB) free!
[1582104615] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(189) < 92.42.104.166 [19/Feb/2020:15:00:13 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1582104673] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179748 kB) free!
[1582104735] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1) < 92.42.104.166 [19/Feb/2020:15:00:15 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1582104763] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1582104763] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 367 processes
[1582104794] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188656 kB) free!
[1582104855] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582104914] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177496 kB) free!
[1582104960] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1582104960] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 60.1% (29613740 kB) free.
[1582105107] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 18.0% (4453188 kB) free!
[1582105227] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 16.6% (4093632 kB) free!
[1582105347] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 15.2% (3751368 kB) free!
[1582105488] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.4% (4758364 kB) free!
[1582105608] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.9% (4620368 kB) free!
[1582105728] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 18.2% (4461128 kB) free!
[1582105934] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.5% (15157648 kB) free!
[1582105947] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.3% (22793256 kB) free.
[1582106054] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11714056 kB) free!
[1582106174] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20360984 kB) free!
[1582106285] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (9638436 kB) free!
[1582106405] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (9608444 kB) free!
[1582106525] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.4% (8103244 kB) free!
[1582106899] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (4866376 kB) free!
[1582107019] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.1% (4723268 kB) free!
[1582107139] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.4% (4296972 kB) free!
[1582107222] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.7% (20275404 kB) free.
[1582107374] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1582107374] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40508252 kB) free.
[1582107449] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1582107569] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1582107595] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(37) < 2020/02/19 15:42:50 [error] 12246#0: *56546285 open() "/usr/share/nginx/html/js/tmpl/404.ejs" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /js/tmpl/404.ejs?0.051820887981984365 HTTP/1.1", host: "qservercluster1.com:8091", referrer: "http://qservercluster1.com:8091/"
[1582107716] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582107828] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184776 kB) free!
[1582107948] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (165504 kB) free!
[1582107974] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.2% (13465388 kB) free!
[1582108068] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194320 kB) free!
[1582108094] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.4% (9816832 kB) free!
[1582108207] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (1805560 kB) free!
[1582108214] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.0% (50077784 kB) free.
[1582108315] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(65) < 2020/02/19 16:01:34 [error] 12251#0: *56557820 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1582108327] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186996 kB) free!
[1582108340] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184260 kB) free!
[1582108435] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582108447] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (202668 kB) free!
[1582108455] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(84) < 49.50.64.122 [19/Feb/2020:16:03:11 +0530] TCP 502 0 0 120.000 "10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1582108576] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582108728] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1721892 kB) free!
[1582108814] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23681308 kB) free!
[1582108888] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.7% (2877576 kB) free!
[1582108934] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.6% (19222664 kB) free!
[1582109008] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.6% (3186556 kB) free!
[1582109054] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15120856 kB) free!
[1582109129] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.4% (3478988 kB) free.
[1582109654] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28797116 kB) free.
[1582110125] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.5% (4180340 kB) free!
[1582110328] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.1% (2460740 kB) free!
[1582110448] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.3% (2811084 kB) free!
[1582110568] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (3137800 kB) free!
[1582110854] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.5% (16475064 kB) free!
[1582110974] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.3% (12289872 kB) free!
[1582111094] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (8351016 kB) free!
[1582111168] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.0% (3738388 kB) free.
[1582111901] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.6% (21726576 kB) free.
[1582111935] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.4% (20816116 kB) free.
[1582111978] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 45.9% (60572036 kB) free.
[1582112018] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.5% (39651712 kB) free.
[1582112054] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.2% (22109244 kB) free.
[1582112087] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 60.7% (14867808 kB) free.
[1582112635] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/02/19 17:13:41 [error] 12256#0: *56601926 open() "/usr/share/nginx/html/js/tmpl/404.ejs" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /js/tmpl/404.ejs?0.10465307730099305 HTTP/1.1", host: "qservercluster1.com:8091", referrer: "http://qservercluster1.com:8091/"
[1582112696] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194964 kB) free!
[1582112755] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582112816] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177780 kB) free!
[1582112936] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186524 kB) free!
[1582113135] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170036 kB) free!
[1582113255] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195484 kB) free!
[1582113375] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171836 kB) free!
[1582113785] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25478480 kB) free!
[1582113905] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.0% (21037156 kB) free!
[1582114025] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.9% (16983720 kB) free!
[1582114422] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1582114625] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.2% (30517528 kB) free.
[1582115021] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.5% (3710080 kB) free!
[1582115141] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 15.6% (3508332 kB) free!
[1582115261] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.9% (3340156 kB) free!
[1582115825] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.5% (23043412 kB) free!
[1582115861] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 71.6% (16070980 kB) free.
[1582115945] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.2% (18676296 kB) free!
[1582115968] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.4% (2998884 kB) free!
[1582116066] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.1% (14677668 kB) free!
[1582116088] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.2% (3277932 kB) free.
[1582116163] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1582116283] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1582116361] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582116403] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 356 processes
[1582116666] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.1% (29090900 kB) free.
[1582117004] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1582117428] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.2% (8489108 kB) free!
[1582117549] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.2% (8486792 kB) free!
[1582117604] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1582117669] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.2% (8487276 kB) free!
[1582117724] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 359 processes
[1582117844] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 365 processes
[1582117866] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.0% (19758832 kB) free!
[1582117888] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.6% (2531564 kB) free!
[1582117986] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.8% (15613736 kB) free!
[1582118008] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.5% (2848376 kB) free!
[1582118059] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190280 kB) free!
[1582118106] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (11806272 kB) free!
[1582118128] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (3008876 kB) free!
[1582118179] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176508 kB) free!
[1582118299] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (201392 kB) free!
[1582118444] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 359 processes
[1582118706] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27934112 kB) free.
[1582118728] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1582118728] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.2% (3773972 kB) free.
[1582119355] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/02/19 19:02:34 [error] 12248#0: *56666258 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/uuid HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/uuid", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/uuid"
[1582119476] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582119535] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (182972 kB) free!
[1582119655] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (200052 kB) free!
[1582119775] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (212824 kB) free!
[1582119906] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17396048 kB) free!
[1582120026] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (12840540 kB) free!
[1582120146] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (8786484 kB) free!
[1582120746] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27937252 kB) free.
[1582121263] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 705 processes
[1582121383] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1582121503] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1582121946] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17879552 kB) free!
[1582122066] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.2% (13453048 kB) free!
[1582122187] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.2% (9473384 kB) free!
[1582122787] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.4% (24226712 kB) free!
[1582123075] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(98) < 2020/02/19 20:01:54 [error] 12247#0: *56695957 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1582123195] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1582123303] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1582123387] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.9% (43319472 kB) free.
[1582123582] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1582123669] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (4382432 kB) free!
[1582123736] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1582123987] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16840916 kB) free!
[1582124107] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (12562404 kB) free!
[1582124176] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1582124227] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.6% (8671356 kB) free!
[1582124299] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1582124827] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.9% (23625344 kB) free!
[1582125427] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42214056 kB) free.
[1582125703] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1582125823] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1582125982] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1582126027] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.6% (19283384 kB) free!
[1582126102] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1582126147] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.5% (15114736 kB) free!
[1582126222] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1582126244] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1582126267] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (11095844 kB) free!
[1582126364] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1582126867] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.3% (25439704 kB) free!
[1582127467] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.6% (41595648 kB) free.
[1582128067] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18846064 kB) free!
[1582128187] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.1% (14695848 kB) free!
[1582128304] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.8% (21913216 kB) free.
[1582128307] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (10778368 kB) free!
[1582128363] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.1% (39457508 kB) free.
[1582128413] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.1% (25191720 kB) free!
[1582128415] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.3% (33396844 kB) free.
[1582128904] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (260752 kB) free!
[1582129024] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198972 kB) free!
[1582129144] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197372 kB) free!
[1582129377] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.5% (21929308 kB) free.
[1582129451] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.7% (21897908 kB) free.
[1582129676] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 47.0% (11520648 kB) free.
[1582129707] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9865496 kB) free!
[1582129776] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.0% (14145044 kB) free.
[1582129827] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 20.0% (9849344 kB) free!
[1582129947] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 19.9% (9821708 kB) free!
[1582129964] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582130023] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1582130128] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582130144] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1582130215] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 20.0% (26386876 kB) free!
[1582130288] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186312 kB) free!
[1582130335] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.1% (22542576 kB) free!
[1582130408] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (225564 kB) free!
[1582130455] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40300472 kB) free.
[1582130529] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (204236 kB) free!
[1582131164] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1582131284] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1582131404] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 353 processes
[1582132255] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23351920 kB) free!
[1582132375] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.8% (19521648 kB) free!
[1582132496] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 39.4% (51893760 kB) free.
[1582132604] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1582133143] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 701 processes
[1582133204] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1582133263] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1582133384] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 700 processes
[1582133422] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1582134376] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1582135560] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 12.2% (5989452 kB) free!
[1582135680] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (212504 kB) free!
[1582135800] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (238164 kB) free!
[1582136095] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24145764 kB) free!
[1582136215] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 40.7% (53660440 kB) free.
[1582136577] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (1521260 kB) free!
[1582136697] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.9% (1228624 kB) free!
[1582136817] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (804580 kB) free!
[1582136969] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes