[1583433106] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1583433200] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 6.3% (8349568 kB) free!
[1583433268] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 58.4% (28796572 kB) free.
[1583433414] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 203.212.70.170 [06/Mar/2020:00:01:43 +0530] TCP 502 0 0 60.001 "10.147.212.155:3001, 10.147.212.142:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 0.000, 60.000, 0.000"
[1583433534] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583434400] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21540272 kB) free!
[1583435000] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32468180 kB) free.
[1583435003] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.0% (1215792 kB) free!
[1583435063] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (201580 kB) free!
[1583435123] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;3;CRITICAL - 0.7% (164564 kB) free!
[1583435183] SERVICE ALERT: webserver2;Memory;CRITICAL;HARD;4;CRITICAL - 0.8% (187480 kB) free!
[1583435334] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 14.0% (3450660 kB) free!
[1583435600] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.4% (11099388 kB) free!
[1583435720] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7555940 kB) free!
[1583435740] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1583435783] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 70.6% (17292604 kB) free.
[1583435840] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32431152 kB) free.
[1583435860] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1583435934] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (227664 kB) free!
[1583436440] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.9% (11797224 kB) free!
[1583436534] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 39.9% (9845104 kB) free.
[1583436560] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (8293612 kB) free!
[1583436680] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (22778164 kB) free!
[1583437134] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (1163284 kB) free!
[1583437254] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164932 kB) free!
[1583437280] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (11958792 kB) free!
[1583437374] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1583437374] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188140 kB) free!
[1583437880] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19222248 kB) free!
[1583438480] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28173748 kB) free.
[1583438813] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.2% (13504324 kB) free!
[1583439080] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.9% (5173948 kB) free!
[1583439201] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32000356 kB) free.
[1583439413] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (13013548 kB) free!
[1583439800] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10857548 kB) free!
[1583439920] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.6% (7396384 kB) free!
[1583440013] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1583440013] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.7% (14111452 kB) free!
[1583440040] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.2% (31919992 kB) free.
[1583440640] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9787704 kB) free!
[1583440760] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6295592 kB) free!
[1583440880] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.2% (29250592 kB) free.
[1583441164] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 41.7% (10213996 kB) free.
[1583441480] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7670572 kB) free!
[1583441600] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.5% (31043208 kB) free.
[1583441741] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (235956 kB) free!
[1583441862] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (212784 kB) free!
[1583441982] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240496 kB) free!
[1583442200] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.5% (9837384 kB) free!
[1583442320] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31824312 kB) free.
[1583442582] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 79.9% (39366216 kB) free.
[1583442920] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.1% (12049404 kB) free!
[1583442924] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.1% (42908084 kB) free.
[1583443040] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.5% (8614204 kB) free!
[1583443160] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31089076 kB) free.
[1583443760] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10195372 kB) free!
[1583443881] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31811912 kB) free.
[1583444116] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(214) < 103.16.101.72 [06/Mar/2020:03:00:23 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1583444173] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(272) < 2020/03/06 03:00:23 [error] 12242#0: *65723024 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583444238] SERVICE ALERT: secondary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583444294] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583444334] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3483) < 203.212.70.170 [06/Mar/2020:03:01:42 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583444455] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583444480] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.3% (10947272 kB) free!
[1583444600] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.1% (31756680 kB) free.
[1583444860] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1583444980] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1583445200] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11661728 kB) free!
[1583445320] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.1% (27859236 kB) free.
[1583445920] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12424020 kB) free!
[1583446040] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (9021136 kB) free!
[1583446160] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28862612 kB) free.
[1583446760] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9455140 kB) free!
[1583446880] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29042932 kB) free.
[1583447480] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9599596 kB) free!
[1583447600] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28775592 kB) free.
[1583447893] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(21) < 2020/03/06 04:00:38 [error] 12257#0: *65742132 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583448013] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583448774] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583448800] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12156800 kB) free!
[1583448920] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31568212 kB) free.
[1583449521] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.4% (12358716 kB) free!
[1583449641] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31504208 kB) free.
[1583450241] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6931876 kB) free!
[1583450361] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31380908 kB) free.
[1583450961] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7924116 kB) free!
[1583451081] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30072776 kB) free.
[1583451365] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170352 kB) free!
[1583451413] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583451485] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181788 kB) free!
[1583451606] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (148872 kB) free!
[1583452281] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.2% (10796220 kB) free!
[1583452375] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 50.8% (12528420 kB) free.
[1583452401] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583452401] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31112296 kB) free.
[1583452780] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1583452883] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 80.8% (19800844 kB) free.
[1583452900] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1583453601] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.4% (15042928 kB) free!
[1583453721] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.9% (31464624 kB) free.
[1583455521] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14907944 kB) free!
[1583455613] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 38.0% (50128688 kB) free.
[1583455641] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31335512 kB) free.
[1583457441] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15633516 kB) free!
[1583457561] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31301828 kB) free.
[1583459961] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17427272 kB) free!
[1583460081] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.7% (32535444 kB) free.
[1583460316] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 89.0% (21816228 kB) free.
[1583460346] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.7% (38782444 kB) free.
[1583460378] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 87.1% (21344536 kB) free.
[1583460400] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.0% (14133276 kB) free.
[1583460775] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190692 kB) free!
[1583460895] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (202120 kB) free!
[1583461015] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (203104 kB) free!
[1583461281] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.6% (16608748 kB) free!
[1583461300] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1583461401] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31966948 kB) free.
[1583461420] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1583462414] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(26) < 2020/03/06 08:01:07 [error] 12254#0: *65821928 connect() failed (111: Connection refused) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583462454] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(36) < 203.212.70.170 [06/Mar/2020:08:02:58 +0530] TCP 502 0 0 127.009 "10.147.212.154:3001, 10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "7.008, 60.000, 60.001, 0.000"
[1583462534] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583462574] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1) < 172.105.59.173 [06/Mar/2020:08:12:30 +0530] TCP 502 0 0 60.000 "10.147.212.142:3001, 10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.000, 0.000, 60.000, 0.000"
[1583462694] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583463295] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 182.18.144.246 [06/Mar/2020:08:19:20 +0530] TCP 502 0 0 60.000 "10.147.212.154:3001, backend_3001" "0, 0" "0, 0" "60.000, 0.000"
[1583463415] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583463802] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14701476 kB) free!
[1583463922] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30297420 kB) free.
[1583466322] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (23912472 kB) free!
[1583466442] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28087892 kB) free.
[1583468842] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (8383352 kB) free!
[1583468962] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.2% (26668444 kB) free.
[1583469562] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25487068 kB) free!
[1583469682] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.5% (20460008 kB) free!
[1583469802] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15715000 kB) free!
[1583469853] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 703 processes
[1583469973] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 701 processes
[1583470093] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1583470290] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.9% (21701436 kB) free.
[1583470403] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37593720 kB) free.
[1583472203] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20044380 kB) free!
[1583472323] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.1% (15936640 kB) free!
[1583472443] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40707820 kB) free.
[1583472734] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(32) < 2020/03/06 11:02:12 [error] 12256#0: *65891785 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1583472854] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(29) < 2020/03/06 11:02:30 [error] 12257#0: *65891864 recv() failed (104: Connection reset by peer) while proxying connection, client: 182.18.144.246, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1583472974] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583473015] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(2) < 203.212.70.170 [06/Mar/2020:11:04:27 +0530] TCP 502 0 0 120.000 "10.147.212.155:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1583473043] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.1% (14648652 kB) free!
[1583473135] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583473163] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.6% (38966172 kB) free.
[1583473306] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 12.3% (2768288 kB) free!
[1583473426] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 10.9% (2437372 kB) free!
[1583473546] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (2091240 kB) free!
[1583473763] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16338456 kB) free!
[1583473883] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.4% (12416276 kB) free!
[1583474003] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.0% (36902996 kB) free.
[1583474492] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.4% (3307884 kB) free!
[1583474603] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12802044 kB) free!
[1583474612] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.1% (2499300 kB) free!
[1583474668] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 19.8% (9776688 kB) free!
[1583474723] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (9024320 kB) free!
[1583474732] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2193332 kB) free!
[1583474751] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.7% (7223324 kB) free!
[1583474783] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 16.6% (4075060 kB) free!
[1583474788] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 18.6% (9146244 kB) free!
[1583474843] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (31992880 kB) free.
[1583474871] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 14.6% (7208568 kB) free!
[1583474903] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 15.5% (3787868 kB) free!
[1583474908] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 17.2% (8463684 kB) free!
[1583474991] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 14.6% (7202584 kB) free!
[1583475023] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 14.2% (3490592 kB) free!
[1583475364] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 16.1% (3968744 kB) free!
[1583475443] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9235928 kB) free!
[1583475484] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 64.2% (14409196 kB) free.
[1583475484] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 14.2% (3513864 kB) free!
[1583475509] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.5% (21583128 kB) free.
[1583475541] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.5% (20430584 kB) free!
[1583475543] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30088728 kB) free.
[1583475581] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 25.5% (12570016 kB) free.
[1583475602] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 55.1% (27154984 kB) free.
[1583475604] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 12.2% (3007388 kB) free!
[1583475623] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (1936456 kB) free!
[1583476093] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 707 processes
[1583476204] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 3.5% (861844 kB) free!
[1583476213] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 710 processes
[1583476333] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 707 processes
[1583476574] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(61) < 2020/03/06 12:02:01 [error] 12248#0: *65918383 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.38:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583476694] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583476712] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (4491264 kB) free!
[1583476735] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1) < 203.212.70.170 [06/Mar/2020:12:01:42 +0530] TCP 502 0 0 60.001 "10.147.212.154:3001, 10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "0.001, 0.000, 60.000, 0.000"
[1583476832] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 13.0% (3204872 kB) free!
[1583476855] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583476952] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1583476952] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (230612 kB) free!
[1583477404] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.3% (22794444 kB) free.
[1583477412] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8809568 kB) free!
[1583477483] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4662196 kB) free!
[1583477532] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.0% (8357580 kB) free!
[1583477533] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1583477603] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.1% (4429044 kB) free!
[1583477652] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.0% (8356716 kB) free!
[1583477723] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (4196952 kB) free!
[1583477813] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.7% (24624588 kB) free!
[1583477933] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.1% (31743948 kB) free.
[1583477952] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (25141200 kB) free!
[1583478055] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1031) < 103.255.146.164 [06/Mar/2020:12:25:54 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583478072] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.1% (21220100 kB) free!
[1583478175] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583478192] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35178144 kB) free.
[1583478252] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1583478252] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.1% (3508436 kB) free!
[1583478324] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9785096 kB) free!
[1583478444] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.1% (9396004 kB) free!
[1583478533] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 14.1% (18620600 kB) free!
[1583478564] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.2% (8992376 kB) free!
[1583478653] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 23.4% (30792424 kB) free.
[1583478792] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.2% (12097988 kB) free!
[1583478912] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8437364 kB) free!
[1583478923] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1849852 kB) free!
[1583479032] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35091556 kB) free.
[1583479253] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.1% (23830968 kB) free!
[1583479373] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.0% (33002860 kB) free.
[1583479523] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.5% (4044760 kB) free!
[1583479632] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11254548 kB) free!
[1583479644] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 15.7% (3860440 kB) free!
[1583479752] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7488392 kB) free!
[1583479764] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.0% (3675276 kB) free!
[1583479872] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34358956 kB) free.
[1583480304] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.2% (2629732 kB) free!
[1583480364] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (2352320 kB) free!
[1583480364] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.4% (3138640 kB) free!
[1583480424] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.2% (2309336 kB) free!
[1583480472] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.5% (11221840 kB) free!
[1583480544] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (3188484 kB) free!
[1583480573] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.4% (24222440 kB) free!
[1583480592] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.5% (7216792 kB) free!
[1583480693] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 24.6% (32383584 kB) free.
[1583480712] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.6% (3457716 kB) free!
[1583480982] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (9508664 kB) free!
[1583481102] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (9258660 kB) free!
[1583481144] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 35.0% (5691720 kB) free.
[1583481222] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.2% (8988820 kB) free!
[1583481893] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.0% (23777792 kB) free!
[1583481912] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.2% (23938816 kB) free!
[1583482013] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 13.2% (17429428 kB) free!
[1583482133] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.2% (26562848 kB) free.
[1583482512] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2433532 kB) free!
[1583482617] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.8% (21676064 kB) free.
[1583482640] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 55.2% (13535532 kB) free.
[1583482780] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 59.4% (14560492 kB) free.
[1583482816] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1583482816] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 53.8% (70974556 kB) free.
[1583482844] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.8% (39330080 kB) free.
[1583482874] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 72.2% (35573624 kB) free.
[1583482903] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 88.4% (21669328 kB) free.
[1583482962] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 89.3% (44002120 kB) free.
[1583483292] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 18.2% (4083264 kB) free!
[1583483333] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 17.7% (23347860 kB) free!
[1583483412] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.7% (3752236 kB) free!
[1583483453] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.1% (33047804 kB) free.
[1583483532] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.1% (3383988 kB) free!
[1583484094] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1583484114] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190168 kB) free!
[1583484132] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (1947776 kB) free!
[1583484144] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (3195368 kB) free!
[1583484199] SERVICE ALERT: primary-server;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 349 processes
[1583484214] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1583484234] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197904 kB) free!
[1583484264] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (3040928 kB) free!
[1583484319] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1583484335] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1583484354] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186376 kB) free!
[1583484385] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.5% (3011808 kB) free!
[1583484422] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179604 kB) free!
[1583484440] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1583484442] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198948 kB) free!
[1583484542] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (765248 kB) free!
[1583484562] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (181284 kB) free!
[1583484653] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.4% (25569820 kB) free!
[1583484662] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (598300 kB) free!
[1583484682] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1583484682] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179248 kB) free!
[1583484773] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.8% (19489460 kB) free!
[1583484775] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1216) < 103.255.146.164 [06/Mar/2020:14:19:05 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1583484893] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.4% (28221892 kB) free.
[1583484895] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583484935] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1583485040] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1583485221] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.2% (14822348 kB) free!
[1583485341] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.3% (10976912 kB) free!
[1583485461] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 5.8% (7655724 kB) free!
[1583485554] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 77.5% (19000396 kB) free.
[1583485585] SERVICE FLAPPING ALERT: webservercluster2;Memory;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1583485585] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.5% (3980292 kB) free.
[1583485852] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9856140 kB) free!
[1583485973] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.8% (9779064 kB) free!
[1583486061] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.8% (19568816 kB) free!
[1583486093] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.9% (9297032 kB) free!
[1583486093] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.2% (25259124 kB) free!
[1583486213] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.8% (34039948 kB) free.
[1583486383] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9054584 kB) free!
[1583486405] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 19.3% (4752684 kB) free!
[1583486503] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.4% (8578076 kB) free!
[1583486526] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 17.8% (4396472 kB) free!
[1583486623] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.5% (8133528 kB) free!
[1583486646] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 16.2% (3992620 kB) free!
[1583486661] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.7% (44415588 kB) free.
[1583486735] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1583486818] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1583486855] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1583486938] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1583486975] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1583487132] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 69.8% (15664564 kB) free.
[1583487230] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 59.7% (29401020 kB) free.
[1583487246] SERVICE FLAPPING ALERT: db1-vm2;Memory;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1583487246] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.7% (1907224 kB) free!
[1583487261] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.7% (11486892 kB) free!
[1583487295] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(186) < 92.42.104.166 [06/Mar/2020:15:00:23 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1583487381] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.8% (7676008 kB) free!
[1583487415] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583487501] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.4% (4441208 kB) free!
[1583487874] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1583487985] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.3% (1838196 kB) free!
[1583487994] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1583488013] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.1% (25173008 kB) free!
[1583488101] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.5% (20466876 kB) free!
[1583488105] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.2% (2145808 kB) free!
[1583488114] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1583488133] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 14.1% (18523256 kB) free!
[1583488225] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 14.7% (2389392 kB) free!
[1583488247] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.2% (22006692 kB) free.
[1583488253] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.8% (30044472 kB) free.
[1583488283] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 37.3% (9151900 kB) free.
[1583488321] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.6% (22852156 kB) free.
[1583488825] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 29.7% (4831588 kB) free.
[1583488853] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 10.1% (13373304 kB) free!
[1583488973] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 15.2% (20062680 kB) free!
[1583489093] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4335392 kB) free!
[1583489093] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.6% (17965312 kB) free!
[1583489240] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1583489301] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (840068 kB) free!
[1583489360] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 343 processes
[1583489425] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (2839940 kB) free!
[1583489480] SERVICE ALERT: primary-server;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1583489493] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167640 kB) free!
[1583489545] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (3204532 kB) free!
[1583489613] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177056 kB) free!
[1583489665] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.2% (3602432 kB) free.
[1583489693] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.1% (5363672 kB) free!
[1583489733] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181708 kB) free!
[1583489754] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (1517620 kB) free!
[1583489874] SERVICE ALERT: primary-server;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (1218520 kB) free!
[1583489901] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.0% (17155968 kB) free!
[1583489983] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (4389932 kB) free!
[1583489994] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1583489994] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 3.8% (929884 kB) free!
[1583490103] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.1% (4186592 kB) free!
[1583490223] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.0% (3932252 kB) free!
[1583490265] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (1322784 kB) free!
[1583490385] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.6% (1730888 kB) free!
[1583490502] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (378936 kB) free!
[1583490505] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 13.4% (2172664 kB) free!
[1583490576] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1583490576] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 353 processes
[1583490893] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.6% (17880852 kB) free!
[1583491015] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1608) < 203.212.70.170 [06/Mar/2020:16:03:22 +0530] TCP 502 0 0 120.001 "10.147.212.143:3001, 10.147.212.154:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.001, 60.000, 0.000"
[1583491094] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(176) < 2020/03/06 16:02:10 [error] 12257#0: *66037597 upstream timed out (110: Connection timed out) while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.37:4000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583491102] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.8% (16858144 kB) free!
[1583491105] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.0% (819560 kB) free!
[1583491135] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583491176] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1583491214] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583491423] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.0% (989036 kB) free!
[1583491702] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (416960 kB) free!
[1583491705] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 12.3% (2004140 kB) free!
[1583491852] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166048 kB) free!
[1583491871] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 78.9% (19338352 kB) free.
[1583491894] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 74.6% (18412984 kB) free.
[1583491925] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.7% (43165652 kB) free.
[1583491963] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 62.3% (30720100 kB) free.
[1583491996] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1583491996] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 76.4% (18714728 kB) free.
[1583492024] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.8% (22001656 kB) free.
[1583492093] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1583492093] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 4.5% (5925444 kB) free!
[1583492305] SERVICE ALERT: webservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 4.8% (784348 kB) free!
[1583492501] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187688 kB) free!
[1583492622] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177432 kB) free!
[1583492742] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187972 kB) free!
[1583492906] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 18.8% (3059700 kB) free!
[1583492938] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 334 processes
[1583492976] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1583493058] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1583493096] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 357 processes
[1583493171] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (244000 kB) free!
[1583493216] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1583493291] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (244896 kB) free!
[1583493293] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.1% (23834644 kB) free!
[1583493411] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (216292 kB) free!
[1583493731] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.3% (16187180 kB) free!
[1583493816] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1583493851] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.5% (12558124 kB) free!
[1583493893] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.5% (29643868 kB) free.
[1583493936] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1583493971] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.8% (45810716 kB) free.
[1583494056] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1583494332] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 16.1% (3606420 kB) free!
[1583494333] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1583494453] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.4% (3231452 kB) free!
[1583494453] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 699 processes
[1583494493] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 11.4% (14997508 kB) free!
[1583494571] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.6% (20566384 kB) free!
[1583494573] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 12.9% (2895452 kB) free!
[1583494613] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 6.4% (8378620 kB) free!
[1583494691] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.5% (16516104 kB) free!
[1583494706] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 21.6% (3509512 kB) free.
[1583494733] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 18.0% (23689428 kB) free!
[1583494811] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 9.6% (12685216 kB) free!
[1583495030] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170724 kB) free!
[1583495150] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (207764 kB) free!
[1583495173] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STARTED; Service appears to have started flapping (24.5% change >= 20.0% threshold)
[1583495173] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1497972 kB) free!
[1583495256] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1583495270] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (206624 kB) free!
[1583495333] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (7921976 kB) free!
[1583495411] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.7% (27265848 kB) free.
[1583495856] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1583495859] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.1% (21843108 kB) free.
[1583495901] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 86.7% (21402320 kB) free.
[1583495914] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1583495933] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 10.9% (14356336 kB) free!
[1583495937] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.7% (39267352 kB) free.
[1583495976] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1583496009] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 53.9% (12106452 kB) free.
[1583496011] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.3% (6924500 kB) free!
[1583496084] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 16.9% (8337192 kB) free!
[1583496131] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3382080 kB) free!
[1583496204] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.3% (8015656 kB) free!
[1583496251] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39329616 kB) free.
[1583496324] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 15.7% (7737524 kB) free!
[1583496384] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;1;WARNING - 13.3% (3262852 kB) free!
[1583496444] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;2;WARNING - 13.3% (3269088 kB) free!
[1583496503] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167444 kB) free!
[1583496504] SERVICE ALERT: webserver2;Memory;WARNING;SOFT;3;WARNING - 13.3% (3263776 kB) free!
[1583496533] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.2% (33265832 kB) free.
[1583496564] SERVICE ALERT: webserver2;Memory;WARNING;HARD;4;WARNING - 13.3% (3264000 kB) free!
[1583496623] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (188432 kB) free!
[1583496743] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174184 kB) free!
[1583496830] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 18.4% (9043800 kB) free!
[1583496851] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.3% (13537188 kB) free!
[1583496950] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 17.5% (8636140 kB) free!
[1583496971] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (9525396 kB) free!
[1583497070] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 16.8% (8260000 kB) free!
[1583497091] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (5818912 kB) free!
[1583497133] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 12.9% (17040148 kB) free!
[1583497137] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (247488 kB) free!
[1583497253] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 25.7% (33830244 kB) free.
[1583497257] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (243968 kB) free!
[1583497321] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 16.7% (4114560 kB) free!
[1583497377] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (227332 kB) free!
[1583497441] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 15.4% (3790832 kB) free!
[1583497562] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 14.1% (3470196 kB) free!
[1583497692] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.8% (20797408 kB) free!
[1583497706] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.4% (1032456 kB) free!
[1583497735] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(967) < 172.105.59.173 [06/Mar/2020:17:54:11 +0530] TCP 502 0 0 120.000 "10.147.212.154:3001, 10.147.212.142:3001, backend_3001" "0, 0, 0" "0, 0, 0" "60.000, 60.000, 0.000"
[1583497826] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.1% (2458252 kB) free!
[1583497853] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 13.6% (17970272 kB) free!
[1583497856] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583497918] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583497946] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 24.6% (3992160 kB) free.
[1583497973] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.6% (11289720 kB) free!
[1583498048] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1583498093] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 20.0% (26356896 kB) free!
[1583498100] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583498106] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 38.35, 59.99, 41.84
[1583498124] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3883608 kB) free!
[1583498162] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.2% (2016636 kB) free!
[1583498178] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1583498230] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1583498236] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1583498280] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1583498292] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (377824 kB) free!
[1583498360] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1583498366] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1583498377] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1583498457] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583498471] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 19.2% (4693972 kB) free!
[1583498489] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583498497] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 360 processes
[1583498587] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1583498591] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 18.2% (4466732 kB) free!
[1583498602] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.7% (4338820 kB) free!
[1583498617] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 356 processes
[1583498619] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583498619] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1583498619] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583498693] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (11860488 kB) free!
[1583498712] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.3% (4252164 kB) free!
[1583498717] SERVICE ALERT: smscapp2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1583498722] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.7% (4091392 kB) free!
[1583498749] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1583498749] SERVICE ALERT: smscapp2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1583498749] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1583498762] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 92.6% (22859532 kB) free.
[1583498784] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1583498842] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.0% (3917952 kB) free!
[1583498874] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38760 MB (75% inode=98%):
[1583498875] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2241 processes
[1583498880] SERVICE FLAPPING ALERT: smscapp2;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1583498880] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 58.7% (28911892 kB) free.
[1583498892] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.5% (15151504 kB) free!
[1583498914] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1583498963] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2240 processes
[1583498968] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 25.54, 28.31, 32.76
[1583499034] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1583499108] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 72.0% (17770220 kB) free.
[1583499137] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 32.5% (42817284 kB) free.
[1583499147] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 13.3% (2154564 kB) free!
[1583499154] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1583499217] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1583499220] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 50.1% (24702656 kB) free.
[1583499247] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 26.0% (12836900 kB) free.
[1583499267] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 23.4% (3801836 kB) free.
[1583499323] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1583499352] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 82396 MB (36% inode=99%):
[1583499394] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583499398] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583499490] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583499492] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.5% (36296472 kB) free.
[1583499518] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1583499524] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1583499573] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583499578] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1583499620] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1583499654] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1583499703] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1583499713] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196072 kB) free!
[1583499747] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 58.4% (28768048 kB) free.
[1583499833] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170900 kB) free!
[1583499833] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1583499912] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 7.0% (1705512 kB) free!
[1583499933] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583499953] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (164176 kB) free!
[1583500042] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.0% (1481184 kB) free!
[1583500062] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (1673784 kB) free!
[1583500063] SERVICE ALERT: smscapp2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1583500092] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10355104 kB) free!
[1583500182] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.8% (1659560 kB) free!
[1583500190] SERVICE ALERT: smscapp2;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1583500212] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6197244 kB) free!
[1583500255] SERVICE FLAPPING ALERT: smscapp2;Sensors Status;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1583500255] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1583500302] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1659324 kB) free!
[1583500332] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.8% (2319996 kB) free!
[1583500350] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583500422] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (219088 kB) free!
[1583500480] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1583500542] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (237132 kB) free!
[1583500610] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1583500662] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (241324 kB) free!
[1583500679] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583500679] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583500729] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583500799] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38760 MB (75% inode=98%):
[1583500800] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 2256 processes
[1583500849] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1583500865] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583500932] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.1% (17333232 kB) free!
[1583500995] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1583501017] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1583501067] SERVICE ALERT: webservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.5% (575520 kB) free!
[1583501125] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1583501137] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1583501187] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 13.1% (2122532 kB) free!
[1583501216] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 58.5% (28814160 kB) free.
[1583501257] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1583501307] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.7% (3682984 kB) free.
[1583501503] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 16.0% (3918868 kB) free!
[1583501532] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (416760 kB) free!
[1583501633] SERVICE FLAPPING ALERT: smscapp2;Total Processes;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1583501633] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 2265 processes
[1583501727] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1583501980] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1583501980] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 17.22, 24.69, 30.76
[1583502014] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/03/06 19:08:27 [error] 12257#0: *66137376 upstream timed out (110: Connection timed out) while reading response header from upstream, client: 183.82.105.122, server: webserver2.com, request: "POST /temtapp/index.php/DlrLoad/msisdn HTTP/1.0", upstream: "http://10.147.212.40:80/temtapp/index.php/DlrLoad/msisdn", host: "altsmsc.tesync.net", referrer: "http://altsmsc.tesync.net/temtapp/index.php/DlrLoad/msisdn"
[1583502134] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583502402] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.2% (2138692 kB) free!
[1583502522] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.7% (2062084 kB) free!
[1583502580] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 13.51, 15.94, 23.39
[1583502642] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 13.1% (2125424 kB) free!
[1583502655] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(1333) < 203.212.70.170 [06/Mar/2020:19:20:55 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583502775] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(101) < 49.50.64.122 [06/Mar/2020:19:22:13 +0530] TCP 502 0 0 135.036 "10.147.212.142:3001, 10.147.212.154:3001, 10.147.212.143:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "15.036, 60.000, 60.000, 0.000"
[1583502895] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(49) < 172.105.59.173 [06/Mar/2020:19:24:54 +0530] TCP 502 0 0 75.037 "10.147.212.154:3001, 10.147.212.142:3001, 10.147.212.155:3001, backend_3001" "0, 0, 0, 0" "0, 0, 0, 0" "60.000, 15.037, 0.000, 0.000"
[1583503092] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 368 processes
[1583503212] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 384 processes
[1583503332] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 373 processes
[1583503373] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.4% (9583328 kB) free!
[1583503493] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.2% (9486360 kB) free!
[1583503613] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.0% (9373996 kB) free!
[1583503657] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1583503777] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1583503897] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 359 processes
[1583503903] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169352 kB) free!
[1583503932] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 388 processes
[1583503932] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39241544 kB) free.
[1583504532] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.5% (13880176 kB) free!
[1583504652] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.2% (9548672 kB) free!
[1583504695] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583504772] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (5553572 kB) free!
[1583505013] SERVICE FLAPPING ALERT: sig6-vm2;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1583505097] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1583505373] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.2% (27951772 kB) free.
[1583505507] SERVICE FLAPPING ALERT: webservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1583505658] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 332 processes
[1583505779] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1583505895] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(622) < 203.212.70.170 [06/Mar/2020:20:14:49 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583505973] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (4824556 kB) free!
[1583506016] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(49) < 203.212.70.170 [06/Mar/2020:20:16:07 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583506093] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (904168 kB) free!
[1583506136] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1583506213] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.8% (26117016 kB) free!
[1583506297] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1583506418] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1583506538] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1583506813] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (13224760 kB) free!
[1583507412] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 19.8% (4448320 kB) free!
[1583507413] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36357940 kB) free.
[1583507532] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 18.7% (4197020 kB) free!
[1583507652] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 17.5% (3925048 kB) free!
[1583508013] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8890096 kB) free!
[1583508133] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.5% (4572808 kB) free!
[1583508253] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (582048 kB) free!
[1583508507] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 15.7% (2557016 kB) free!
[1583508627] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 20.8% (3388372 kB) free.
[1583508703] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 54.6% (13389220 kB) free.
[1583508852] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 70.1% (15739264 kB) free.
[1583508853] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.3% (22779404 kB) free!
[1583509013] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4880840 kB) free!
[1583509227] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.3% (3141344 kB) free!
[1583509343] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583509347] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.5% (4138828 kB) free.
[1583509453] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398148 kB) free!
[1583510053] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.3% (21485992 kB) free!
[1583510137] SERVICE ALERT: smscapp2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583510138] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1583510153] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1583510162] SERVICE FLAPPING ALERT: db1-vm2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583510257] SERVICE ALERT: smscapp2;Sensors Status;OK;HARD;1;SENSORS OK
[1583510547] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 16.9% (2751760 kB) free!
[1583510653] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (410472 kB) free!
[1583510667] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 22.0% (3578736 kB) free.
[1583510712] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1583510790] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 79.9% (19583156 kB) free.
[1583510817] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 71.6% (17680744 kB) free.
[1583510842] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1583510854] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 82.5% (40642624 kB) free.
[1583510862] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1583510894] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 40.8% (53812724 kB) free.
[1583510930] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 47.1% (23198012 kB) free.
[1583510956] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 66.1% (16192224 kB) free.
[1583511267] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.7% (1907532 kB) free!
[1583511387] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;2;WARNING - 15.3% (2486516 kB) free!
[1583511423] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196468 kB) free!
[1583511507] SERVICE ALERT: webservercluster2;Memory;WARNING;HARD;3;WARNING - 19.1% (3099136 kB) free!
[1583511533] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (1342472 kB) free!
[1583511543] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187684 kB) free!
[1583511653] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (213044 kB) free!
[1583511663] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175236 kB) free!
[1583511773] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240772 kB) free!
[1583511834] SERVICE FLAPPING ALERT: smscapp2;Total Processes;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1583512107] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 25.4% (4127672 kB) free.
[1583512616] SERVICE FLAPPING ALERT: smscapp2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583512702] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.3% (12201876 kB) free!
[1583512822] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.0% (7848492 kB) free!
[1583512942] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 3.0% (3889672 kB) free!
[1583513021] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1583513138] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1583513141] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1583513257] SERVICE FLAPPING ALERT: smscapp2;Sensors Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583513542] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.5% (25745812 kB) free!
[1583513554] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1583513654] SERVICE ALERT: redis-1;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513654] SERVICE ALERT: redis-2;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513654] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513654] SERVICE ALERT: redis-2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513654] SERVICE ALERT: redis-2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513654] SERVICE ALERT: rq-1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513654] SERVICE ALERT: mrq4;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513657] HOST ALERT: mrq4;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.241)
[1583513659] SERVICE ALERT: rq-1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513659] SERVICE ALERT: redis-1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513659] SERVICE ALERT: redis-1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513671] SERVICE ALERT: rq-1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513671] SERVICE ALERT: rq-1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513671] SERVICE ALERT: rq-1;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513671] SERVICE ALERT: redis-1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513671] SERVICE ALERT: redis-1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513671] SERVICE ALERT: redis-1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513671] SERVICE ALERT: rq-1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513671] SERVICE ALERT: redis-2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513671] SERVICE ALERT: redis-2;HDD_VNS1 status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513671] SERVICE ALERT: redis-2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513671] SERVICE ALERT: redis-2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513684] HOST ALERT: redis-1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1583513684] HOST ALERT: redis-2;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1583513684] HOST ALERT: rq-1;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1583513702] SERVICE ALERT: mrq4;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'mrq4' was out of bounds)
[1583513738] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 367 processes
[1583513744] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'redis-1' was out of bounds)
[1583513753] SERVICE ALERT: rq-2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513756] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583513756] HOST ALERT: rq-2;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.204)
[1583513762] SERVICE ALERT: mrq4;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'mrq4' was out of bounds)
[1583513777] SERVICE ALERT: redis-1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'redis-1' was out of bounds)
[1583513777] SERVICE ALERT: redis-2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'redis-2' was out of bounds)
[1583513777] SERVICE ALERT: redis-2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'redis-2' was out of bounds)
[1583513777] SERVICE ALERT: redis-2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'redis-2' was out of bounds)
[1583513777] SERVICE ALERT: redis-2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'redis-2' was out of bounds)
[1583513777] SERVICE ALERT: rq-1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'rq-1' was out of bounds)
[1583513777] SERVICE ALERT: mrq4;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'mrq4' was out of bounds)
[1583513780] SERVICE ALERT: redis-1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'redis-1' was out of bounds)
[1583513780] SERVICE ALERT: redis-1;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'redis-1' was out of bounds)
[1583513780] SERVICE ALERT: rq-1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'rq-1' was out of bounds)
[1583513794] SERVICE ALERT: rq-1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'rq-1' was out of bounds)
[1583513794] SERVICE ALERT: rq-1;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'rq-1' was out of bounds)
[1583513794] SERVICE ALERT: rq-1;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'rq-1' was out of bounds)
[1583513794] SERVICE ALERT: rq-1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'rq-1' was out of bounds)
[1583513794] SERVICE ALERT: redis-1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'redis-1' was out of bounds)
[1583513794] SERVICE ALERT: redis-1;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'redis-1' was out of bounds)
[1583513794] SERVICE ALERT: redis-1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'redis-1' was out of bounds)
[1583513794] SERVICE ALERT: redis-2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'redis-2' was out of bounds)
[1583513794] SERVICE ALERT: redis-2;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'redis-2' was out of bounds)
[1583513794] SERVICE ALERT: redis-2;HDD_VNS1 status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_VNS1 status' on host 'redis-2' was out of bounds)
[1583513794] SERVICE ALERT: redis-2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'redis-2' was out of bounds)
[1583513797] SERVICE ALERT: rq-1;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'rq-1' was out of bounds)
[1583513814] SERVICE ALERT: rq-2;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'rq-2' was out of bounds)
[1583513840] SERVICE ALERT: mrq4;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'mrq4' was out of bounds)
[1583513840] SERVICE ALERT: rq-2;HDD_VNS1 status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_VNS1 status' on host 'rq-2' was out of bounds)
[1583513840] SERVICE ALERT: rq-2;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'rq-2' was out of bounds)
[1583513840] SERVICE ALERT: rq-2;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'rq-2' was out of bounds)
[1583513858] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 423 processes
[1583513868] SERVICE ALERT: rq-2;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'rq-2' was out of bounds)
[1583513868] SERVICE ALERT: rq-2;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'rq-2' was out of bounds)
[1583513876] SERVICE ALERT: rq-2;rq2-rabbitmqerror;CRITICAL;HARD;1;(Return code of 255 for service 'rq2-rabbitmqerror' on host 'rq-2' was out of bounds)
[1583513876] SERVICE ALERT: rq-2;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'rq-2' was out of bounds)
[1583513885] SERVICE ALERT: mrq4;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'mrq4' was out of bounds)
[1583513908] SERVICE ALERT: redis-1;HDD_VNS1 status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_VNS1 status' on host 'redis-1' was out of bounds)
[1583513919] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 677 processes
[1583513937] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(6377) < 103.255.146.164 [06/Mar/2020:22:28:56 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1583513941] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'mrq4' was out of bounds)
[1583513980] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583514014] SERVICE ALERT: mrq4;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'mrq4' was out of bounds)
[1583514039] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 689 processes
[1583514057] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(8542) < 203.212.70.170 [06/Mar/2020:22:30:57 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1583514111] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;HARD;1;(Return code of 255 for service 'rq1-rabbitmqerror' on host 'rq-1' was out of bounds)
[1583514112] SERVICE ALERT: rq-2;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'rq-2' was out of bounds)
[1583514135] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3974) < 2020/03/06 22:32:14 [error] 12251#0: *66239156 no live upstreams while connecting to upstream, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583514137] SERVICE ALERT: rq-2;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'rq-2' was out of bounds)
[1583514145] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'redis-1' was out of bounds)
[1583514145] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;(Return code of 255 for service 'Memory' on host 'redis-2' was out of bounds)
[1583514159] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 406 processes
[1583514177] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;HARD;3;(7429) < 203.212.70.170 [06/Mar/2020:22:32:57 +0530] TCP 502 0 0 1.002 "10.147.212.142:3001, backend_3001" "0, 0" "0, 0" "1.002, 0.000"
[1583514182] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'mrq4' was out of bounds)
[1583514224] HOST ALERT: mrq4;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.241)
[1583514244] SERVICE ALERT: rq-1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'rq-1' was out of bounds)
[1583514244] SERVICE ALERT: rq-1;HDD_VNS1 status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_VNS1 status' on host 'rq-1' was out of bounds)
[1583514247] SERVICE ALERT: redis-1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'redis-1' was out of bounds)
[1583514247] SERVICE ALERT: redis-2;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'redis-2' was out of bounds)
[1583514250] HOST ALERT: redis-1;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.201)
[1583514250] HOST ALERT: redis-2;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.202)
[1583514250] HOST ALERT: rq-1;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.203)
[1583514255] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(1876) < 2020/03/06 22:34:15 [error] 12246#0: *66241637 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583514316] HOST ALERT: rq-2;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.204)
[1583514375] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;HARD;3;(1815) < 2020/03/06 22:36:15 [error] 12244#0: *66243938 no live upstreams while connecting to upstream, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "backend3", bytes from/to client:0/0, bytes from/to upstream:0/0
[1583514437] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(3661) < 49.50.86.152 [06/Mar/2020:22:37:16 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1583514558] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;SOFT;2;(1115) < 49.50.86.152 [06/Mar/2020:22:39:17 +0530] TCP 502 0 0 0.001 "10.147.212.38:4000, 10.147.212.49:4000, backend3" "0, 0, 0" "0, 0, 0" "0.001, 0.000, 0.000"
[1583514680] SERVICE ALERT: secondary-server;Nginx Access Log Status;CRITICAL;HARD;3;(955) < 103.16.101.72 [06/Mar/2020:22:41:18 +0530] TCP 502 0 0 0.000 "backend3" "0" "0" "0.000"
[1583514736] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 50.39, 21.93, 8.68
[1583514857] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 55.05, 32.88, 14.28
[1583514977] SERVICE ALERT: sig4-vm2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 55.63, 40.32, 19.25
[1583515181] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 279.18, 243.94, 135.65
[1583515259] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583515302] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 279.39, 256.08, 153.77
[1583515384] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1583515423] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 277.83, 263.47, 168.90
[1583515809] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583515903] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1583515935] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 1781 processes
[1583515989] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583516043] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583516119] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1583516173] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1583516250] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1583516303] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1583516855] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1583516909] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 1775 processes
[1583517139] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583517264] SERVICE ALERT: smscapp2;CPU_Procs;OK;HARD;1;CPU OK: 1775 processes
[1583518065] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583518119] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583518195] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1583518249] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 1774 processes
[1583518321] SERVICE ALERT: smscapp2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1583519069] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1583519199] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1583519329] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.