[1578940351] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28342532 kB) free.
[1578940444] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 81.0% (39941440 kB) free.
[1578940756] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1578940805] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(11) < 2020/01/14 00:02:56 [error] 21678#0: *24828588 connect() failed (111: Connection refused) while connecting to upstream, client: 92.42.104.166, server: 0.0.0.0:13001, upstream: "10.147.212.154:13001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578940925] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578940951] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.6% (6102124 kB) free!
[1578941071] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.8% (31316844 kB) free.
[1578941356] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1578941476] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1578941596] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1578941671] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9513212 kB) free!
[1578941791] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (5881188 kB) free!
[1578941911] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.7% (31195408 kB) free.
[1578942196] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1578942316] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1578942361] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177808 kB) free!
[1578942421] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 70.0% (17153752 kB) free.
[1578942511] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9246208 kB) free!
[1578942631] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.3% (5637320 kB) free!
[1578942751] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31092236 kB) free.
[1578942917] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1578943037] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1578943157] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1578943160] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 40.4% (9974244 kB) free.
[1578943351] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8986256 kB) free!
[1578943471] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5397968 kB) free!
[1578943592] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30767060 kB) free.
[1578944191] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7862168 kB) free!
[1578944224] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/14 01:06:20 [error] 18370#0: *38582028 access forbidden by rule while initializing session, client: 103.255.146.164, server: 0.0.0.0:14001
[1578944311] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (8073880 kB) free!
[1578944344] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578944357] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1578944360] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.6% (641556 kB) free!
[1578944431] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27731948 kB) free.
[1578944477] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1578944480] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186376 kB) free!
[1578944597] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1578944600] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169680 kB) free!
[1578945031] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.5% (7286792 kB) free!
[1578945151] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30856300 kB) free.
[1578945197] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1578945318] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1578945751] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9533452 kB) free!
[1578945871] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (5883244 kB) free!
[1578945991] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29456304 kB) free.
[1578946261] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1578946381] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 402 processes
[1578946501] SERVICE ALERT: redis-1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 401 processes
[1578946591] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (7576956 kB) free!
[1578946711] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.3% (30708052 kB) free.
[1578947101] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1578947117] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1578947238] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1578947311] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8852552 kB) free!
[1578947358] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1578947431] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28614884 kB) free.
[1578947958] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1578948031] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (9649464 kB) free!
[1578948151] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.6% (6097572 kB) free!
[1578948272] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27696780 kB) free.
[1578948558] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1578948607] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 42.6% (10447208 kB) free.
[1578948678] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1578948742] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.6% (15307420 kB) free!
[1578948768] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.1% (8412344 kB) free!
[1578948871] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.2% (6878688 kB) free!
[1578948888] SERVICE ALERT: qservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.3% (3106892 kB) free!
[1578948991] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28250668 kB) free.
[1578949008] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (2826284 kB) free!
[1578949591] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (7140464 kB) free!
[1578949608] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 86.1% (42438116 kB) free.
[1578949711] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28421664 kB) free.
[1578949878] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1578949998] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1578950118] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1578950311] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7650540 kB) free!
[1578950431] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.4% (28169104 kB) free.
[1578950457] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.8% (45235544 kB) free.
[1578951031] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8096632 kB) free!
[1578951151] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28523988 kB) free.
[1578951318] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1578951751] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (7998492 kB) free!
[1578951871] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27715944 kB) free.
[1578953071] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.4% (9747668 kB) free!
[1578953191] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.9% (30197056 kB) free.
[1578953791] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.8% (8991792 kB) free!
[1578953911] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29622052 kB) free.
[1578954512] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1578955111] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.6% (9971176 kB) free!
[1578955231] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29646316 kB) free.
[1578955518] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 345 processes
[1578955638] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1578955758] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1578956358] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1578956431] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11371560 kB) free!
[1578956478] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1578956551] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29949964 kB) free.
[1578957078] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1578957199] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1578957751] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11344736 kB) free!
[1578957871] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28674204 kB) free.
[1578958325] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(880) < 2020/01/14 05:00:07 [error] 21683#0: *24843267 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578958398] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 358 processes
[1578958445] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578958518] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1578958750] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(880) < 103.255.146.164 [14/Jan/2020:05:00:07 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1578958809] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177520 kB) free!
[1578958870] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578958929] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192876 kB) free!
[1578959049] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196560 kB) free!
[1578959071] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.7% (18120072 kB) free!
[1578959118] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1578959191] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27392436 kB) free.
[1578959238] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1578959600] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 48.4% (11946264 kB) free.
[1578959838] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1578959958] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1578960078] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1578960991] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.4% (4451816 kB) free!
[1578961112] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27602456 kB) free.
[1578961878] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 353 processes
[1578962478] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1578962542] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (11730624 kB) free!
[1578962598] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 347 processes
[1578962718] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1578962911] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 3.6% (4722808 kB) free!
[1578963031] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26814104 kB) free.
[1578963142] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 31.7% (41805848 kB) free.
[1578963301] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 408 processes
[1578963421] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 406 processes
[1578964108] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (2311724 kB) free!
[1578965119] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1578965719] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1578965839] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1578965959] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1578966559] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1578969845] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2020/01/14 08:07:01 [error] 21681#0: *24851349 connect() failed (111: Connection refused) while connecting to upstream, client: 185.255.8.58, server: 0.0.0.0:13001, upstream: "10.147.212.143:13001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578969965] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578971165] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/14 08:30:37 [error] 21683#0: *24852324 connect() failed (111: Connection refused) while connecting to upstream, client: 185.255.8.58, server: 0.0.0.0:13001, upstream: "10.147.212.142:13001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578971285] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578971359] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1578971459] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.4% (21906328 kB) free.
[1578971479] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1578971503] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.1% (21099764 kB) free.
[1578971542] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.5% (40636524 kB) free.
[1578971599] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1578971638] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.4% (22641984 kB) free.
[1578972107] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176108 kB) free!
[1578972227] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189856 kB) free!
[1578972347] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179952 kB) free!
[1578973231] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5812040 kB) free!
[1578973351] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 18.5% (24322464 kB) free!
[1578973399] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1578973400] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195136 kB) free!
[1578973471] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19866992 kB) free!
[1578973519] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 346 processes
[1578973520] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192604 kB) free!
[1578973640] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197136 kB) free!
[1578975919] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1578976039] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1578976471] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33282440 kB) free.
[1578976519] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.9% (22183656 kB) free.
[1578976562] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 83.2% (20385032 kB) free.
[1578976794] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1578976914] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1578977034] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1578977167] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (183884 kB) free!
[1578977205] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 87.3% (21407216 kB) free.
[1578977809] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182948 kB) free!
[1578977839] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1578977929] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (193384 kB) free!
[1578977959] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1578978049] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1578978049] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177192 kB) free!
[1578978079] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1578978234] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1578978271] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18822056 kB) free!
[1578978391] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38629960 kB) free.
[1578978991] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (16014732 kB) free!
[1578979111] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36138320 kB) free.
[1578979111] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 705 processes
[1578979231] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 704 processes
[1578979279] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1578979351] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 704 processes
[1578979879] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 367 processes
[1578980000] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1578980311] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.8% (15525272 kB) free!
[1578980431] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11689644 kB) free!
[1578980551] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33041360 kB) free.
[1578980599] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1578980719] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1578980868] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.8% (3990348 kB) free!
[1578980988] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.0% (3599556 kB) free!
[1578981108] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 14.4% (3236384 kB) free!
[1578981151] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (10395044 kB) free!
[1578981271] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.3% (32066356 kB) free.
[1578981319] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1578981439] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1578981559] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1578981708] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1536336 kB) free!
[1578981748] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9470392 kB) free!
[1578981868] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (9459660 kB) free!
[1578981871] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (9188256 kB) free!
[1578981926] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.5% (2584260 kB) free!
[1578981989] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.2% (9459340 kB) free!
[1578981991] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.0% (31648200 kB) free.
[1578982046] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (2405040 kB) free!
[1578982167] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2230420 kB) free!
[1578982444] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;1;WARNING - 17.9% (8802080 kB) free!
[1578982564] SERVICE ALERT: smscapp2;Memory;WARNING;SOFT;2;WARNING - 16.7% (8252132 kB) free!
[1578982591] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7375588 kB) free!
[1578982684] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 15.9% (7847476 kB) free!
[1578982711] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.1% (30445348 kB) free.
[1578982760] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1578982951] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 689 processes
[1578983098] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;1;WARNING - 14.8% (3653556 kB) free!
[1578983160] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1578983160] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 89.2% (22008608 kB) free.
[1578983218] SERVICE ALERT: db1-vm2;Memory;WARNING;SOFT;2;WARNING - 13.3% (3290176 kB) free!
[1578983312] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.1% (6695168 kB) free!
[1578983338] SERVICE ALERT: db1-vm2;Memory;WARNING;HARD;3;WARNING - 12.0% (2970060 kB) free!
[1578983360] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1578983432] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.0% (29004320 kB) free.
[1578983447] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 75.4% (16937692 kB) free.
[1578983480] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 367 processes
[1578983514] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 82.7% (20266704 kB) free.
[1578983634] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1578983754] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1578983874] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1578983891] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 62.2% (30651496 kB) free.
[1578983938] SERVICE ALERT: db1-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 5.7% (1411436 kB) free!
[1578984032] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6461400 kB) free!
[1578984116] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194572 kB) free!
[1578984152] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27678736 kB) free.
[1578984236] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189668 kB) free!
[1578984357] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166644 kB) free!
[1578984538] SERVICE ALERT: db1-vm2;Memory;OK;HARD;1;OK - 93.4% (23062580 kB) free.
[1578984752] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 2.7% (3537944 kB) free!
[1578984872] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (380824 kB) free!
[1578984992] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33830548 kB) free.
[1578985075] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.1% (21091824 kB) free.
[1578985280] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1578985400] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1578985457] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.2% (4463216 kB) free!
[1578985568] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168064 kB) free!
[1578985577] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (2178712 kB) free!
[1578985589] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (3988128 kB) free!
[1578985674] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1578985674] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1578985678] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (178720 kB) free!
[1578985688] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.8% (21923992 kB) free.
[1578985697] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2034236 kB) free!
[1578985700] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171096 kB) free!
[1578985727] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.5% (21687132 kB) free.
[1578985762] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.4% (40609592 kB) free.
[1578985790] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.6% (22440164 kB) free.
[1578985859] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 17.4% (4269760 kB) free!
[1578985979] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 16.8% (4118020 kB) free!
[1578986001] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1578986099] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 16.3% (3995932 kB) free!
[1578986121] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1578986241] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1578986330] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187364 kB) free!
[1578986450] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194736 kB) free!
[1578986458] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 17.2% (8484564 kB) free!
[1578986570] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184984 kB) free!
[1578986578] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 16.7% (8221148 kB) free!
[1578986698] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.2% (7968316 kB) free!
[1578986841] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1578986997] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185064 kB) free!
[1578987061] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.3% (21640588 kB) free.
[1578987098] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 84.7% (20754080 kB) free.
[1578987142] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 93.1% (45867968 kB) free.
[1578987169] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.2% (21866688 kB) free.
[1578987441] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1578987485] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/14 13:02:58 [error] 21673#0: *24864471 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:27587/104, bytes from/to upstream:104/27587
[1578987561] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1578987605] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578987663] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (193176 kB) free!
[1578987697] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (166932 kB) free!
[1578987783] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183752 kB) free!
[1578987817] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (182236 kB) free!
[1578987903] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175932 kB) free!
[1578987937] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167316 kB) free!
[1578988097] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170488 kB) free!
[1578988116] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 340 processes
[1578988161] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 347 processes
[1578988217] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163980 kB) free!
[1578988236] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1578988281] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1578988338] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169264 kB) free!
[1578988402] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 347 processes
[1578988639] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.5% (400832 kB) free!
[1578988646] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23934024 kB) free!
[1578988759] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.3% (381568 kB) free!
[1578988766] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 36.9% (48654548 kB) free.
[1578988879] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.4% (393600 kB) free!
[1578989002] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1578989366] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24936708 kB) free!
[1578989406] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/14 13:36:08 [error] 21681#0: *24866031 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:2964/87, bytes from/to upstream:87/2964
[1578989486] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.1% (21184916 kB) free!
[1578989526] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578989602] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1578989606] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.2% (45025580 kB) free.
[1578989651] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.6% (21868064 kB) free.
[1578989722] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1578990206] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22586256 kB) free!
[1578990322] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1578990326] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.8% (45916088 kB) free.
[1578990442] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1578990562] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1578990636] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 338 processes
[1578990756] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1578990926] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (24039092 kB) free!
[1578991046] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.2% (19986052 kB) free!
[1578991162] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 348 processes
[1578991166] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.8% (44509924 kB) free.
[1578991762] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1578991766] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22081696 kB) free!
[1578991886] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.8% (18242656 kB) free!
[1578992006] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.4% (42738768 kB) free.
[1578992526] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(3) < 2020/01/14 14:28:40 [error] 21681#0: *24868742 connect() failed (111: Connection refused) while connecting to upstream, client: 185.255.8.58, server: 0.0.0.0:13001, upstream: "10.147.212.154:13001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578992606] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19911800 kB) free!
[1578992646] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578992727] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.2% (16056340 kB) free!
[1578992847] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40122572 kB) free.
[1578992963] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 376 processes
[1578993057] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 15.6% (3509296 kB) free!
[1578993178] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 14.2% (3179740 kB) free!
[1578993247] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(4) < 2020/01/14 14:39:57 [error] 21681#0: *24869350 connect() failed (111: Connection refused) while connecting to upstream, client: 92.42.104.166, server: 0.0.0.0:13001, upstream: "10.147.212.155:13001", bytes from/to client:0/0, bytes from/to upstream:0/0
[1578993298] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 13.0% (2916052 kB) free!
[1578993367] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1578993447] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16822844 kB) free!
[1578993563] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1578993567] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.1% (41034456 kB) free.
[1578993591] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.7% (22714500 kB) free.
[1578993631] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.4% (21181672 kB) free.
[1578993671] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 88.3% (14358140 kB) free.
[1578993756] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 339 processes
[1578993855] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189024 kB) free!
[1578993876] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1578993898] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 7.4% (1669832 kB) free!
[1578993975] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171892 kB) free!
[1578994074] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1578994096] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (167308 kB) free!
[1578994163] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 354 processes
[1578994167] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.8% (19535324 kB) free!
[1578994194] SERVICE ALERT: sservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1578994237] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185176 kB) free!
[1578994287] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (15636224 kB) free!
[1578994315] SERVICE ALERT: sservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1578994357] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (166460 kB) free!
[1578994407] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33204560 kB) free.
[1578994477] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186720 kB) free!
[1578994498] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 75.3% (16913384 kB) free.
[1578995007] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (7641304 kB) free!
[1578995127] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3451272 kB) free!
[1578995247] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (402376 kB) free!
[1578995364] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.7% (9684200 kB) free!
[1578995484] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (9684992 kB) free!
[1578995604] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.7% (9684232 kB) free!
[1578996447] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20300908 kB) free!
[1578997047] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (398696 kB) free!
[1578997763] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1578997883] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 357 processes
[1578997903] SERVICE ALERT: webservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (3229040 kB) free!
[1578998023] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 28.0% (4556336 kB) free.
[1578998208] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9482104 kB) free!
[1578998247] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1578998247] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.3% (20142144 kB) free!
[1578998312] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.0% (21714236 kB) free.
[1578998328] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.9% (9338140 kB) free!
[1578998348] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 86.5% (21210836 kB) free.
[1578998448] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.7% (9213356 kB) free!
[1578998483] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 361 processes
[1578998603] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1578998723] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1578998848] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (534856 kB) free!
[1578998957] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182412 kB) free!
[1578999077] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189068 kB) free!
[1578999151] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 702 processes
[1578999197] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166084 kB) free!
[1578999271] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 702 processes
[1578999392] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 706 processes
[1578999448] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16137980 kB) free!
[1578999804] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.6% (3743316 kB) free!
[1578999992] SERVICE ALERT: smscapp1;Total Processes;OK;HARD;1;PROCS OK: 698 processes
[1579000048] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.5% (32332972 kB) free.
[1579000315] SERVICE FLAPPING ALERT: sservercluster2;Current Users;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1579000648] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.2% (8170020 kB) free!
[1579000768] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (4427804 kB) free!
[1579000888] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34375704 kB) free.
[1579001123] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1579001243] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1579001488] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.7% (8860580 kB) free!
[1579001608] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.7% (4851364 kB) free!
[1579001728] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35053972 kB) free.
[1579001919] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.3% (4027340 kB) free!
[1579002039] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 10.7% (2643108 kB) free!
[1579002115] SERVICE ALERT: sservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1579002159] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.7% (1653960 kB) free!
[1579002328] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.2% (9524868 kB) free!
[1579002448] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.1% (5339192 kB) free!
[1579002568] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (1445644 kB) free!
[1579002648] SERVICE ALERT: qservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4651924 kB) free!
[1579002814] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 87.7% (21632580 kB) free.
[1579002856] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 93.2% (45939940 kB) free.
[1579002891] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 46.4% (61164704 kB) free.
[1579002929] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.3% (40538280 kB) free.
[1579002965] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.1% (21591208 kB) free.
[1579003369] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;1;WARNING - 18.5% (4527692 kB) free!
[1579003417] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 10.0% (2463500 kB) free!
[1579003489] SERVICE ALERT: primary-server;Memory;WARNING;SOFT;2;WARNING - 17.9% (4391836 kB) free!
[1579003537] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169560 kB) free!
[1579003568] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (169964 kB) free!
[1579003609] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.5% (4277140 kB) free!
[1579003657] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184596 kB) free!
[1579003688] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191136 kB) free!
[1579003799] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4805268 kB) free!
[1579003808] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (199164 kB) free!
[1579003920] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.0% (4645128 kB) free!
[1579004040] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 18.4% (4521324 kB) free!
[1579004098] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 17.4% (3898720 kB) free!
[1579004218] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 16.4% (3678788 kB) free!
[1579004338] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 15.5% (3484204 kB) free!
[1579004736] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (251588 kB) free!
[1579004856] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (213884 kB) free!
[1579004976] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1579004976] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (262264 kB) free!
[1579005299] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.7% (12742840 kB) free!
[1579005368] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/14 18:00:54 [error] 21677#0: *24879131 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.142:3000", bytes from/to client:4928/176, bytes from/to upstream:176/4928
[1579005419] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.4% (42697384 kB) free.
[1579005443] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579005488] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579005538] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 75.2% (16883508 kB) free.
[1579005564] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1579005840] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (2149824 kB) free!
[1579006010] SERVICE ALERT: primary-server;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2117940 kB) free!
[1579006020] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17947732 kB) free!
[1579006140] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.7% (14043812 kB) free!
[1579006163] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1579006260] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 33.5% (44158600 kB) free.
[1579006346] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.9% (9310748 kB) free!
[1579006466] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9081884 kB) free!
[1579006586] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.0% (8862048 kB) free!
[1579006860] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20029080 kB) free!
[1579006980] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 12.0% (15799696 kB) free!
[1579007100] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (11797852 kB) free!
[1579007210] SERVICE ALERT: primary-server;Memory;WARNING;HARD;3;WARNING - 17.4% (4258092 kB) free!
[1579007227] SERVICE FLAPPING ALERT: primary-server;Memory;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1579007227] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 89.3% (21876088 kB) free.
[1579007257] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.0% (20969172 kB) free.
[1579007329] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 49.5% (65260864 kB) free.
[1579007377] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.4% (40109928 kB) free.
[1579007417] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 91.9% (22520216 kB) free.
[1579007460] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 85.9% (21050912 kB) free.
[1579007869] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (163144 kB) free!
[1579007989] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186040 kB) free!
[1579007989] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (236820 kB) free!
[1579008067] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185236 kB) free!
[1579008109] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186016 kB) free!
[1579008109] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (237844 kB) free!
[1579008187] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171628 kB) free!
[1579008229] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (235420 kB) free!
[1579008307] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168808 kB) free!
[1579009088] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/14 19:04:24 [error] 21683#0: *24881847 recv() failed (104: Connection reset by peer) while proxying connection, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:39/0, bytes from/to upstream:0/39
[1579009208] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579009339] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.1% (20990328 kB) free.
[1579009385] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.1% (39947532 kB) free.
[1579009415] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.2% (22103812 kB) free.
[1579009586] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.3% (4563968 kB) free!
[1579009742] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.2% (18715800 kB) free!
[1579009763] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1579009862] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.9% (14401760 kB) free!
[1579009884] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1579009949] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181916 kB) free!
[1579009982] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (10354372 kB) free!
[1579009988] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (214552 kB) free!
[1579010004] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579010016] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (172788 kB) free!
[1579010069] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186496 kB) free!
[1579010108] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (248500 kB) free!
[1579010136] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170152 kB) free!
[1579010189] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (185900 kB) free!
[1579010228] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (226688 kB) free!
[1579010256] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174796 kB) free!
[1579010582] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.4% (20363180 kB) free!
[1579010604] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1579011182] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.8% (35380336 kB) free.
[1579011204] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1579011324] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1579011444] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1579011782] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10676352 kB) free!
[1579011903] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6639984 kB) free!
[1579012023] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.4% (40087536 kB) free.
[1579012604] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.7% (20894740 kB) free.
[1579012623] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.9% (13090132 kB) free!
[1579012644] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 347 processes
[1579012650] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 79.8% (39347548 kB) free.
[1579012691] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 40.2% (52988120 kB) free.
[1579012761] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 91.4% (45066132 kB) free.
[1579012764] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579012796] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 91.5% (22418604 kB) free.
[1579012885] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 356 processes
[1579013207] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191292 kB) free!
[1579013258] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (225948 kB) free!
[1579013327] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189924 kB) free!
[1579013378] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (218208 kB) free!
[1579013400] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (168036 kB) free!
[1579013447] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191156 kB) free!
[1579013484] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1579013498] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (212988 kB) free!
[1579013520] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192900 kB) free!
[1579013604] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1579013640] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186056 kB) free!
[1579013724] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1579014324] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1579014324] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 347 processes
[1579014496] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (9346444 kB) free!
[1579014616] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.0% (5293408 kB) free!
[1579014736] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.5% (24344668 kB) free!
[1579014924] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1579015044] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 350 processes
[1579015337] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (390448 kB) free!
[1579015644] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 349 processes
[1579015765] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 351 processes
[1579015937] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 18.6% (24550652 kB) free!
[1579016250] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.8% (20932656 kB) free.
[1579016291] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.1% (39472088 kB) free.
[1579016322] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 90.8% (22265512 kB) free.
[1579016537] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (1135980 kB) free!
[1579016856] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (196512 kB) free!
[1579016898] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (243788 kB) free!
[1579016925] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (184200 kB) free!
[1579016976] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192648 kB) free!
[1579017018] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (240740 kB) free!
[1579017045] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179236 kB) free!
[1579017096] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187116 kB) free!
[1579017137] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.9% (15700436 kB) free!
[1579017138] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (245132 kB) free!
[1579017165] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (169460 kB) free!
[1579017564] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1579017684] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1579017737] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 24.6% (32368584 kB) free.
[1579017804] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1579018337] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 5.6% (7339380 kB) free!
[1579018404] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1579018457] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 2.6% (3455872 kB) free!
[1579018524] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1579018577] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.7% (37872428 kB) free.
[1579018627] SERVICE FLAPPING ALERT: primary-server;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1579018644] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 348 processes
[1579019177] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.7% (10204568 kB) free!
[1579019245] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1579019297] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.5% (5920472 kB) free!
[1579019417] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (2008796 kB) free!
[1579019877] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 83.9% (20702196 kB) free.
[1579019924] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 43.7% (57550068 kB) free.
[1579019967] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.0% (39438388 kB) free.
[1579020000] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 89.8% (22014864 kB) free.
[1579020488] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189820 kB) free!
[1579020576] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (239184 kB) free!
[1579020607] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175452 kB) free!
[1579020609] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187344 kB) free!
[1579020697] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (236648 kB) free!
[1579020728] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189836 kB) free!
[1579020729] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187396 kB) free!
[1579020817] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (246604 kB) free!
[1579020848] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (184848 kB) free!
[1579021045] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 352 processes
[1579021165] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 353 processes
[1579021286] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1579021808] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2020/01/14 22:38:21 [error] 21677#0: *24890966 recv() failed (104: Connection reset by peer) while proxying connection, client: 92.42.104.166, server: 0.0.0.0:13001, upstream: "10.147.212.143:13001", bytes from/to client:41/0, bytes from/to upstream:0/41
[1579021928] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579022327] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11870612 kB) free!
[1579022447] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 35.2% (46365944 kB) free.
[1579023031] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4859136 kB) free!
[1579023047] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.1% (19942328 kB) free!
[1579023085] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 355 processes
[1579023151] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (4720668 kB) free!
[1579023167] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (15744744 kB) free!
[1579023205] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 355 processes
[1579023271] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.0% (4654004 kB) free!
[1579023288] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (43035976 kB) free.
[1579023325] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1579023887] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23207176 kB) free!
[1579023925] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1579024007] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.7% (19345000 kB) free!
[1579024046] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 345 processes
[1579024127] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16086560 kB) free!
[1579024166] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 345 processes
[1579024727] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.9% (28822984 kB) free.
[1579024766] SERVICE ALERT: sservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 348 processes
[1579024886] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 355 processes
[1579024929] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2020/01/14 23:23:22 [error] 21679#0: *24892784 recv() failed (104: Connection reset by peer) while proxying connection, client: 92.42.104.166, server: 0.0.0.0:13001, upstream: "10.147.212.142:13001", bytes from/to client:41/0, bytes from/to upstream:0/41
[1579025049] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1579025291] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (242828 kB) free!
[1579025411] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (230056 kB) free!
[1579025531] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (236144 kB) free!
[1579025927] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.7% (20706484 kB) free!
[1579026047] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.2% (17373164 kB) free!
[1579026086] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 354 processes
[1579026142] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.3% (389400 kB) free!
[1579026167] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 34.7% (45748572 kB) free.
[1579026206] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1579026262] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (394768 kB) free!
[1579026382] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (1746872 kB) free!