[1572287779] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1572287779] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.77, 2.27, 3.85
[1572287973] SERVICE ALERT: qservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1572288021] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.9% (1930640 kB) free!
[1572288047] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.3% (820828 kB) free!
[1572288061] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.6% (3588148 kB) free!
[1572288141] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.8% (1922540 kB) free!
[1572288167] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.3% (802076 kB) free!
[1572288181] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 14.6% (3586896 kB) free!
[1572288243] SERVICE ALERT: qservercluster2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 551 processes
[1572288261] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.8% (1915664 kB) free!
[1572288287] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.2% (792012 kB) free!
[1572288301] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 14.6% (3583256 kB) free!
[1572288363] SERVICE ALERT: qservercluster2;Total Processes;OK;HARD;1;PROCS OK: 547 processes
[1572288631] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 378 processes
[1572288751] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 2 crit, 0 warn out of 377 processes
[1572288871] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 377 processes
[1572288883] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27654448 kB) free.
[1572289483] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24874452 kB) free!
[1572289603] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.3% (21497548 kB) free!
[1572289723] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18269148 kB) free!
[1572290323] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1572290323] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29388896 kB) free.
[1572290424] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1572290544] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1572290671] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 367 processes
[1572290741] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 40.8% (10063072 kB) free.
[1572290923] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26148256 kB) free!
[1572291043] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.4% (22920188 kB) free!
[1572291163] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19897844 kB) free!
[1572291271] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 374 processes
[1572291341] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197764 kB) free!
[1572291391] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 374 processes
[1572291461] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169464 kB) free!
[1572291511] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 373 processes
[1572291581] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186780 kB) free!
[1572291763] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.6% (28486496 kB) free.
[1572292080] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (231568 kB) free!
[1572292111] SERVICE FLAPPING ALERT: pservercluster1;CPU_Procs;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1572292111] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 372 processes
[1572292199] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 220482 MB (20% inode=99%):
[1572292200] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.1% (4012052 kB) free!
[1572292209] SERVICE ALERT: sigqscluster2;Total Processes;OK;HARD;1;PROCS OK: 405 processes
[1572292282] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8833952 kB) free!
[1572292319] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 220482 MB (20% inode=99%):
[1572292321] SERVICE ALERT: sigqscluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (4003036 kB) free!
[1572292403] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.0% (7399176 kB) free!
[1572292439] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 220482 MB (20% inode=99%):
[1572292523] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 15.0% (7377336 kB) free!
[1572292627] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 19.1% (25127104 kB) free!
[1572292747] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 19.1% (25133824 kB) free!
[1572292867] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.1% (25135492 kB) free!
[1572292963] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (24000088 kB) free!
[1572293084] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.8% (20765948 kB) free!
[1572293204] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17473412 kB) free!
[1572293410] SERVICE ALERT: sigqscluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 904 processes
[1572293530] SERVICE ALERT: sigqscluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 904 processes
[1572293650] SERVICE ALERT: sigqscluster2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 904 processes
[1572293804] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27166004 kB) free.
[1572293911] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 371 processes
[1572294032] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 371 processes
[1572294152] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 371 processes
[1572294404] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23395372 kB) free!
[1572294524] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.4% (20293208 kB) free!
[1572294644] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.5% (29688932 kB) free.
[1572294752] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 374 processes
[1572295181] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 35.5% (8757288 kB) free.
[1572295244] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26268616 kB) free!
[1572295352] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 368 processes
[1572295364] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.4% (22926000 kB) free!
[1572295484] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19736040 kB) free!
[1572295563] SERVICE ALERT: qservercluster2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 553 processes
[1572295684] SERVICE ALERT: qservercluster2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 553 processes
[1572295781] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (188204 kB) free!
[1572295804] SERVICE ALERT: qservercluster2;Total Processes;WARNING;HARD;3;PROCS WARNING: 553 processes
[1572295901] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183536 kB) free!
[1572296021] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1572296021] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191112 kB) free!
[1572296084] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.6% (29788308 kB) free.
[1572296625] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 293 processes
[1572296684] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (26012960 kB) free!
[1572296745] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 300 processes
[1572296804] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.3% (22811388 kB) free!
[1572296924] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.9% (19639904 kB) free!
[1572297004] SERVICE ALERT: qservercluster2;Total Processes;OK;HARD;1;PROCS OK: 548 processes
[1572297524] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.7% (29865160 kB) free.
[1572298124] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25785044 kB) free!
[1572298245] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.1% (22559696 kB) free!
[1572298365] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19248380 kB) free!
[1572298965] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28319144 kB) free.
[1572299179] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1572299565] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (24707896 kB) free!
[1572299685] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.4% (21558580 kB) free!
[1572299805] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (18437560 kB) free!
[1572300405] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.9% (27488544 kB) free.
[1572301005] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (23646568 kB) free!
[1572301125] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.6% (20622152 kB) free!
[1572301245] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30068320 kB) free.
[1572301267] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.2% (26655984 kB) free.
[1572301845] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (25986968 kB) free!
[1572301965] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.3% (22842928 kB) free!
[1572302085] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19739060 kB) free!
[1572302685] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28563572 kB) free.
[1572303285] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25299864 kB) free!
[1572303405] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.8% (22165392 kB) free!
[1572303526] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.5% (19157872 kB) free!
[1572304126] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.3% (28123164 kB) free.
[1572304726] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24915200 kB) free!
[1572304846] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.5% (21757224 kB) free!
[1572304966] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29375932 kB) free.
[1572306166] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (24616612 kB) free!
[1572306286] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.2% (21412324 kB) free!
[1572306406] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.0% (30286212 kB) free.
[1572306752] SERVICE FLAPPING ALERT: pservercluster1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1572307006] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26114700 kB) free!
[1572307126] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.4% (22932812 kB) free!
[1572307246] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.1% (19937432 kB) free!
[1572307421] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1572307846] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28636560 kB) free.
[1572308446] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25484508 kB) free!
[1572308566] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.1% (22549184 kB) free!
[1572308687] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.0% (19809252 kB) free!
[1572309122] SERVICE ALERT: sigqscluster2;Memory;OK;HARD;1;OK - 86.8% (42766620 kB) free.
[1572309667] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.2% (24006364 kB) free!
[1572309787] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 17.6% (23217368 kB) free!
[1572309861] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 84.8% (20773884 kB) free.
[1572309887] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.3% (26760504 kB) free.
[1572309887] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 78.7% (19300796 kB) free.
[1572309901] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 88.0% (21564732 kB) free.
[1572309907] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 17.3% (22824428 kB) free!
[1572310829] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.8% (23152624 kB) free.
[1572311039] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 271648 MB (25% inode=99%):
[1572311091] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 473 processes
[1572311107] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 61.0% (80370496 kB) free.
[1572311151] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 473 processes
[1572311211] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 474 processes
[1572311687] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24740992 kB) free!
[1572311807] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.5% (21692032 kB) free!
[1572311811] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 472 processes
[1572311871] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 473 processes
[1572311927] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.3% (18852108 kB) free!
[1572313127] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.5% (28324856 kB) free.
[1572314327] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25558164 kB) free!
[1572314447] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.3% (22766408 kB) free!
[1572314544] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1572314567] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 15.2% (19998056 kB) free!
[1572314664] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1572315768] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29548552 kB) free.
[1572316368] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25820240 kB) free!
[1572316488] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.1% (22508152 kB) free!
[1572316608] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.6% (19261784 kB) free!
[1572317208] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.4% (29593364 kB) free.
[1572318408] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (24710580 kB) free!
[1572318528] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.5% (21692500 kB) free!
[1572318648] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 14.0% (18514032 kB) free!
[1572319248] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.3% (29339012 kB) free.
[1572320448] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.6% (23248800 kB) free!
[1572320568] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.8% (19461304 kB) free!
[1572320688] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.2% (16049684 kB) free!
[1572321288] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27159284 kB) free.
[1572321634] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170040 kB) free!
[1572321755] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172252 kB) free!
[1572321875] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (181004 kB) free!
[1572321888] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24524552 kB) free!
[1572322009] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.0% (21042140 kB) free!
[1572322129] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17588788 kB) free!
[1572322729] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.7% (28606792 kB) free.
[1572323329] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (25142192 kB) free!
[1572323449] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.3% (21447104 kB) free!
[1572323569] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.3% (17534048 kB) free!
[1572323877] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.7% (22865944 kB) free.
[1572325204] SERVICE ALERT: qservercluster2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 560 processes
[1572325324] SERVICE ALERT: qservercluster2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 563 processes
[1572325369] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.4% (30866548 kB) free.
[1572325444] SERVICE ALERT: qservercluster2;Total Processes;WARNING;HARD;3;PROCS WARNING: 571 processes
[1572325523] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4232992 kB) free!
[1572325672] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.6% (40721632 kB) free.
[1572325969] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.6% (25876856 kB) free!
[1572326089] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.9% (40791708 kB) free.
[1572326152] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1572326272] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1572326392] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1572327256] SERVICE ALERT: sig5-vm2;CPU_Procs;OK;HARD;1;CPU OK: 361 processes
[1572327818] SERVICE ALERT: sigqscluster2;Total Processes;WARNING;HARD;3;PROCS WARNING: 910 processes
[1572327911] SERVICE ALERT: rq-2;CPU_Procs;OK;HARD;1;CPU OK: 348 processes
[1572327964] SERVICE ALERT: rq-1;CPU_Procs;OK;HARD;1;CPU OK: 358 processes
[1572328037] SERVICE ALERT: sigqscluster2;Total Processes;OK;HARD;1;PROCS OK: 912 processes
[1572328352] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 380 processes
[1572328432] SERVICE ALERT: sig3-vm2;CPU_Procs;OK;HARD;1;CPU OK: 454 processes
[1572328472] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 380 processes
[1572328477] SERVICE ALERT: sigqscluster1;Current Users;OK;HARD;1;USERS OK - 8 users currently logged in
[1572328490] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23508840 kB) free!
[1572328495] SERVICE ALERT: rq-2;CPU Load;OK;HARD;1;OK - load average: 18.33, 17.51, 16.86
[1572328592] SERVICE ALERT: pservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 380 processes
[1572328602] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1572328610] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.3% (18898480 kB) free!
[1572328684] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (173256 kB) free!
[1572328730] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34462968 kB) free.
[1572328804] SERVICE ALERT: pservercluster1;Memory;UNKNOWN;SOFT;2;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1572328810] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 383 processes
[1572328833] SERVICE ALERT: qservercluster2;Total Processes;OK;HARD;1;PROCS OK: 590 processes
[1572328850] SERVICE ALERT: qservercluster2;Current Users;OK;HARD;1;USERS OK - 7 users currently logged in
[1572328864] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186304 kB) free!
[1572328926] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1572328944] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1572328944] SERVICE ALERT: pservercluster1;Memory;UNKNOWN;HARD;3;*** WARN level must not be less than CRITICAL when checking FREE memory!
[1572329029] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1572329514] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174492 kB) free!
[1572329525] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (25.3% change >= 20.0% threshold)
[1572329525] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (45.5% change >= 20.0% threshold)
[1572329707] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 0.8% (192816 kB) free!
[1572329930] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26210456 kB) free!
[1572329974] SERVICE ALERT: qservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1572330050] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.2% (21421044 kB) free!
[1572330094] SERVICE ALERT: qservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 330 processes
[1572330170] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 19.9% (26185916 kB) free!
[1572330214] SERVICE ALERT: qservercluster1;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 331 processes
[1572330261] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (4298712 kB) free!
[1572330307] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (184808 kB) free!
[1572330381] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.7% (4101736 kB) free!
[1572330481] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 11.4% (5625588 kB) free!
[1572330501] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.1% (3945168 kB) free!
[1572330601] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 11.4% (5624608 kB) free!
[1572330641] SERVICE ALERT: qservercluster2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 705 processes
[1572330721] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 11.4% (5624520 kB) free!
[1572330762] SERVICE ALERT: qservercluster2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 702 processes
[1572330770] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.1% (34465904 kB) free.
[1572330882] SERVICE ALERT: qservercluster2;Total Processes;WARNING;HARD;3;PROCS WARNING: 700 processes
[1572330908] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 78.7% (19424560 kB) free.
[1572331096] SERVICE ALERT: qservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1572331321] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1572331321] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.9% (3886012 kB) free!
[1572331482] SERVICE FLAPPING ALERT: qservercluster2;Total Processes;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1572331482] SERVICE ALERT: qservercluster2;Total Processes;OK;HARD;1;PROCS OK: 589 processes
[1572331508] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (2001080 kB) free!
[1572331628] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173080 kB) free!
[1572331701] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 6.2% (1508452 kB) free!
[1572331748] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (232920 kB) free!
[1572331970] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25549616 kB) free!
[1572332090] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.8% (20771764 kB) free!
[1572332109] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 83.63, 30.38, 11.69
[1572332211] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.5% (16538468 kB) free!
[1572332229] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 13.00, 20.96, 10.50
[1572332811] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.2% (33211468 kB) free.
[1572333230] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 329 processes
[1572333350] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1572335211] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24805020 kB) free!
[1572335331] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.0% (40921784 kB) free.
[1572335687] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.4% (4258964 kB) free!
[1572335807] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.8% (4119392 kB) free!
[1572335927] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.2% (3965820 kB) free!
[1572337477] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 11 users currently logged in
[1572337597] SERVICE ALERT: sigqscluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 11 users currently logged in
[1572337717] SERVICE ALERT: sigqscluster1;Current Users;WARNING;HARD;3;USERS WARNING - 11 users currently logged in
[1572337727] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 7.3% (1786224 kB) free!
[1572338331] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23410488 kB) free!
[1572338451] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.7% (39136440 kB) free.
[1572339407] SERVICE ALERT: sigqscluster1;Current Users;OK;HARD;1;USERS OK - 12 users currently logged in
[1572340852] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.1% (23819052 kB) free!
[1572340972] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28674532 kB) free.
[1572341800] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1066.07, 759.48, 352.41
[1572341864] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 414 processes
[1572341920] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 1068.13, 861.70, 439.23
[1572341984] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 434 processes
[1572342040] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 904.06, 894.89, 504.06
[1572342105] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 447 processes
[1572342196] SERVICE ALERT: sservercluster1;CPU Load;UNKNOWN;HARD;3;Parameter inconsistency: 5-minute "warning load" is greater than "critical load"
[1572342379] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 1064.58, 829.25, 409.78
[1572342490] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 2.41, 199.62, 311.06
[1572342499] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 144.39, 555.18, 360.21
[1572342501] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.5% (2581432 kB) free!
[1572342619] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 20.33, 371.81, 316.65
[1572342721] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1572342721] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1572342721] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.87, 91.74, 241.84
[1572342882] SERVICE FLAPPING ALERT: qservercluster2;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1572342971] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 94.6% (23342844 kB) free.
[1572343005] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.9% (23009348 kB) free.
[1572343055] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.6% (40726684 kB) free.
[1572343057] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.16, 86.09, 197.55
[1572343111] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.6% (22937424 kB) free.
[1572343159] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.92, 61.83, 177.46
[1572343372] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25434544 kB) free!
[1572343492] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.2% (21314840 kB) free!
[1572343575] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (179036 kB) free!
[1572343612] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.1% (35738964 kB) free.
[1572343695] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.7% (22389284 kB) free.
[1572344308] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186756 kB) free!
[1572344428] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192984 kB) free!
[1572344548] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (198436 kB) free!
[1572344812] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1572345412] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.3% (24145088 kB) free!
[1572345532] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.1% (19847192 kB) free!
[1572345652] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.3% (34605188 kB) free.
[1572347453] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23773124 kB) free!
[1572347573] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 14.9% (19695948 kB) free!
[1572347693] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33810756 kB) free.
[1572348265] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.0% (23436472 kB) free.
[1572348301] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4889412 kB) free!
[1572348421] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (4863828 kB) free!
[1572348541] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 19.7% (4837360 kB) free!
[1572348893] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24286192 kB) free!
[1572349013] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.2% (20008644 kB) free!
[1572349133] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.2% (42481024 kB) free.
[1572350677] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187384 kB) free!
[1572350798] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169304 kB) free!
[1572350918] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186880 kB) free!
[1572351725] SERVICE FLAPPING ALERT: sservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1572352133] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (25953084 kB) free!
[1572352204] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.2% (9457364 kB) free!
[1572352253] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38765940 kB) free.
[1572352324] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (9247760 kB) free!
[1572352445] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.3% (9021824 kB) free!
[1572352741] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2418560 kB) free!
[1572353317] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1212 processes
[1572353437] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1214 processes
[1572353558] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1215 processes
[1572353578] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.0% (20737816 kB) free.
[1572353622] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 92.7% (22733060 kB) free.
[1572353624] SERVICE ALERT: smscapp2;HDD_Root status;CRITICAL;SOFT;1;(Return code of 255 for service 'HDD_Root status' on host 'smscapp2' was out of bounds)
[1572353630] SERVICE ALERT: smscapp2;Total Processes;WARNING;HARD;3;PROCS WARNING: 1218 processes
[1572353677] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 95.2% (46926548 kB) free.
[1572353717] SERVICE ALERT: smscapp2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38864 MB (75% inode=98%):
[1572354185] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (197224 kB) free!
[1572354305] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (179632 kB) free!
[1572354425] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188520 kB) free!
[1572354541] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1572354609] SERVICE FLAPPING ALERT: smscapp2;Total Processes;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1572354609] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 1218 processes
[1572355214] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 1586 processes
[1572355335] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1584 processes
[1572355455] SERVICE ALERT: smscapp2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 1583 processes
[1572358253] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.9% (26288024 kB) free!
[1572358373] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38423288 kB) free.
[1572361373] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24869784 kB) free!
[1572361484] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9687892 kB) free!
[1572361493] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37227032 kB) free.
[1572361604] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.5% (9596516 kB) free!
[1572361724] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.3% (9491900 kB) free!
[1572362270] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (227844 kB) free!
[1572362390] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (232548 kB) free!
[1572362510] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (216868 kB) free!
[1572362825] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 73.2% (18065180 kB) free.
[1572362853] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 95.3% (46953844 kB) free.
[1572362873] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 74.1% (36510016 kB) free.
[1572363425] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189520 kB) free!
[1572363479] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (258348 kB) free!
[1572363546] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192832 kB) free!
[1572363599] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (222204 kB) free!
[1572363666] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194732 kB) free!
[1572363719] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (223348 kB) free!
[1572364493] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25507940 kB) free!
[1572364613] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.1% (37081232 kB) free.
[1572365315] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.5% (4774440 kB) free!
[1572365435] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.3% (4725504 kB) free!
[1572365555] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.0% (4663268 kB) free!
[1572366256] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 1128 processes
[1572367618] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4888788 kB) free!
[1572367738] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.9% (4868020 kB) free!
[1572367858] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.8% (4848472 kB) free!
[1572368813] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.8% (26057784 kB) free!
[1572368933] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.9% (36787096 kB) free.
[1572371657] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1315 processes
[1572371777] SERVICE ALERT: smscapp2;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1315 processes
[1572371897] SERVICE ALERT: smscapp2;Total Processes;WARNING;HARD;3;PROCS WARNING: 1319 processes