[1548356777] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1548356897] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1548357017] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1548357415] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 6.9% (3387028 kB) free!
[1548358015] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.4% (5142856 kB) free!
[1548358533] SERVICE FLAPPING ALERT: qservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1548359029] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1548359150] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 338 processes
[1548359494] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548359614] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548359734] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548360335] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1548360335] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 289 processes
[1548360617] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1548360935] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548361055] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 289 processes
[1548361175] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 288 processes
[1548362975] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548363095] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 292 processes
[1548363215] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 291 processes
[1548363816] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 291 processes
[1548363936] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548363949] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548364056] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 290 processes
[1548364070] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548364656] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 288 processes
[1548368170] SERVICE ALERT: sig2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 40064 MB (17% inode=99%):
[1548368290] SERVICE ALERT: sig2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 40064 MB (17% inode=99%):
[1548368410] SERVICE ALERT: sig2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 40064 MB (17% inode=99%):
[1548376056] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1548392857] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1548392869] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548392871] SERVICE ALERT: db1;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 5519 MB (10% inode=98%):
[1548392977] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 294 processes
[1548392989] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548393097] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 296 processes
[1548393109] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548393415] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 9.8% (4814256 kB) free!
[1548393709] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548394015] SERVICE ALERT: db1;Memory;WARNING;HARD;3;WARNING - 10.4% (5116712 kB) free!
[1548394093] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 4 warn out of 362 processes
[1548394309] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548394429] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548394533] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.98, 10.91, 10.46
[1548394549] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548394653] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.08, 10.90, 10.50
[1548394773] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 11.83, 11.14, 10.63
[1548395215] SERVICE ALERT: db1;Memory;CRITICAL;HARD;3;CRITICAL - 8.8% (4333300 kB) free!
[1548395271] SERVICE ALERT: db1;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 7443 MB (14% inode=99%):
[1548397764] SERVICE ALERT: db2;Memory;WARNING;HARD;3;WARNING - 11.8% (5826748 kB) free!
[1548397897] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1548398497] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1548398617] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 302 processes
[1548398737] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 301 processes
[1548399337] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 304 processes
[1548399349] SERVICE FLAPPING ALERT: db1;CPU_Procs;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1548399349] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1548399457] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 305 processes
[1548399577] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 0 warn out of 344 processes
[1548399610] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1548399627] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Users' on host 'sservercluster1' was out of bounds)
[1548399638] HOST ALERT: sservercluster1;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1548399646] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sservercluster1' was out of bounds)
[1548399671] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sservercluster1' was out of bounds)
[1548399698] SERVICE ALERT: sservercluster1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sservercluster1' was out of bounds)
[1548399716] SERVICE ALERT: sservercluster1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sservercluster1' was out of bounds)
[1548399733] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sservercluster1' was out of bounds)
[1548399750] SERVICE ALERT: sservercluster1;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sservercluster1' was out of bounds)
[1548399753] SERVICE ALERT: sig6;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 356 processes
[1548399768] SERVICE ALERT: sservercluster1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sservercluster1' was out of bounds)
[1548399781] SERVICE ALERT: sservercluster1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sservercluster1' was out of bounds)
[1548399874] SERVICE ALERT: sig6;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 346 processes
[1548400002] SERVICE ALERT: sig6;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1548400004] SERVICE ALERT: sig6;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1548400004] HOST ALERT: sig6;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.58)
[1548400039] SERVICE ALERT: sservercluster1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sservercluster1' was out of bounds)
[1548400125] SERVICE ALERT: sig6;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig6' was out of bounds)
[1548400145] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1548400158] SERVICE ALERT: sig6;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig6' was out of bounds)
[1548400205] HOST ALERT: sservercluster1;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.49)
[1548400265] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1548400271] SERVICE ALERT: sig6;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig6' was out of bounds)
[1548400271] SERVICE ALERT: sig6;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig6' was out of bounds)
[1548400281] SERVICE ALERT: sig6;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig6' was out of bounds)
[1548400327] SERVICE ALERT: sig6;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig6' was out of bounds)
[1548400327] SERVICE ALERT: sig6;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig6' was out of bounds)
[1548400357] SERVICE ALERT: sig6;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sig6' was out of bounds)
[1548400386] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1548400549] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548400571] HOST ALERT: sig6;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1548400669] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548400723] SERVICE ALERT: sig6;Memory;OK;HARD;1;OK - 94.9% (15432572 kB) free.
[1548400759] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.36, 0.37, 0.16
[1548400763] HOST ALERT: sservercluster1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1548400765] SERVICE ALERT: db2;Memory;CRITICAL;HARD;3;CRITICAL - 9.4% (4629740 kB) free!
[1548400777] SERVICE FLAPPING ALERT: db2;CPU_Procs;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1548400777] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1548400789] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548400878] SERVICE ALERT: sig6;Sensors Status;OK;HARD;1;SENSORS OK
[1548400930] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 296 processes
[1548400947] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1548400966] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 95.8% (23486264 kB) free.
[1548400974] SERVICE ALERT: qservercluster2;CPU Load;OK;HARD;1;OK - load average: 4.14, 4.23, 4.98
[1548400991] SERVICE ALERT: sservercluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1548400991] SERVICE ALERT: sig6;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1548400991] SERVICE ALERT: sig6;Total Processes;OK;HARD;1;PROCS OK: 337 processes
[1548401001] SERVICE ALERT: sig6;CPU Load;OK;HARD;1;OK - load average: 0.40, 0.23, 0.11
[1548401018] SERVICE ALERT: sservercluster1;Total Processes;OK;HARD;1;PROCS OK: 294 processes
[1548401036] SERVICE ALERT: sservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1548401047] SERVICE ALERT: sig6;HDD_Root status;OK;HARD;1;DISK OK - free space: / 27119 MB (52% inode=98%):
[1548401047] SERVICE ALERT: sig6;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1548401077] SERVICE ALERT: sig6;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 212988 MB (94% inode=99%):
[1548401088] SERVICE ALERT: sservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 220679 MB (99% inode=99%):
[1548401101] SERVICE ALERT: sservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 41684 MB (81% inode=99%):
[1548401204] SERVICE ALERT: sig6;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 344 processes
[1548401389] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548401525] SERVICE ALERT: sig3;Total Processes;CRITICAL;SOFT;1;(Return code of 255 for service 'Total Processes' on host 'sig3' was out of bounds)
[1548401550] SERVICE ALERT: sig3;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1548401556] SERVICE ALERT: sig3;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1548401556] SERVICE ALERT: sig3;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1548401580] HOST ALERT: sig3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1548401628] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1548401631] HOST ALERT: pservercluster1;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.147.212.37)
[1548401639] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'pservercluster1' was out of bounds)
[1548401647] SERVICE ALERT: sig3;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'sig3' was out of bounds)
[1548401667] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'pservercluster1' was out of bounds)
[1548401667] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'pservercluster1' was out of bounds)
[1548401673] SERVICE ALERT: sig3;Sensors Status;CRITICAL;HARD;1;(Return code of 255 for service 'Sensors Status' on host 'sig3' was out of bounds)
[1548401676] SERVICE ALERT: sig3;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'sig3' was out of bounds)
[1548401676] SERVICE ALERT: sig3;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'sig3' was out of bounds)
[1548401679] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'pservercluster1' was out of bounds)
[1548401701] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Root status' on host 'pservercluster1' was out of bounds)
[1548401751] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'pservercluster1' was out of bounds)
[1548401768] SERVICE ALERT: sig3;CPU_Procs;CRITICAL;HARD;1;(Return code of 255 for service 'CPU_Procs' on host 'sig3' was out of bounds)
[1548401770] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'pservercluster1' was out of bounds)
[1548401804] SERVICE ALERT: sig6;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1548401876] SERVICE ALERT: sig3;CPU Load;CRITICAL;HARD;1;(Return code of 255 for service 'CPU Load' on host 'sig3' was out of bounds)
[1548401977] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1548402010] SERVICE ALERT: sig2;HDD_Home status;CRITICAL;HARD;3;DISK CRITICAL - free space: /home 21983 MB (9% inode=99%):
[1548402056] SERVICE ALERT: sig3;LocalService-NTPD SERVICE;CRITICAL;HARD;1;(Return code of 255 for service 'LocalService-NTPD SERVICE' on host 'sig3' was out of bounds)
[1548402056] SERVICE ALERT: sig3;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'sig3' was out of bounds)
[1548402085] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;HARD;1;(Return code of 255 for service 'HDD_Home status' on host 'pservercluster1' was out of bounds)
[1548402097] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1548402110] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'pservercluster1' was out of bounds)
[1548402111] SERVICE ALERT: sig3;Memory;CRITICAL;HARD;1;(Return code of 255 for service 'Memory' on host 'sig3' was out of bounds)
[1548402146] HOST ALERT: sig3;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.60)
[1548402198] HOST ALERT: pservercluster1;DOWN;HARD;10;CRITICAL - Host Unreachable (10.147.212.37)
[1548402218] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1548402450] HOST ALERT: sig3;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1548402488] SERVICE ALERT: sig3;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1548402490] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 96.9% (23900020 kB) free.
[1548402494] HOST ALERT: pservercluster1;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1548402590] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548402596] SERVICE ALERT: sig3;CPU Load;OK;HARD;1;OK - load average: 0.38, 0.19, 0.08
[1548402710] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 330 processes
[1548402776] SERVICE ALERT: sig3;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 212760 MB (94% inode=99%):
[1548402776] SERVICE ALERT: sig3;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1548402805] SERVICE ALERT: pservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 220675 MB (99% inode=99%):
[1548402830] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 296 processes
[1548402831] SERVICE ALERT: sig3;Memory;OK;HARD;1;OK - 90.2% (14662760 kB) free.
[1548402846] SERVICE ALERT: sig3;Total Processes;OK;HARD;1;PROCS OK: 346 processes
[1548402871] SERVICE ALERT: sig3;Sensors Status;OK;HARD;1;SENSORS OK
[1548402877] SERVICE ALERT: sig3;HDD_Root status;OK;HARD;1;DISK OK - free space: / 37898 MB (74% inode=98%):
[1548402877] SERVICE ALERT: sig3;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1548402948] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1548402959] SERVICE ALERT: pservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1548402988] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1548402988] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 0.00, 0.06, 0.06
[1548402999] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 299 processes
[1548403021] SERVICE ALERT: pservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 41736 MB (81% inode=99%):
[1548403088] SERVICE ALERT: sig3;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 346 processes
[1548403208] SERVICE ALERT: sig3;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 345 processes
[1548403328] SERVICE ALERT: sig3;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 345 processes
[1548403644] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.5% (4533792 kB) free!
[1548403764] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.2% (4468228 kB) free!
[1548403884] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 18.0% (4400832 kB) free!
[1548403928] SERVICE FLAPPING ALERT: sig3;CPU_Procs;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1548403928] SERVICE ALERT: sig3;CPU_Procs;OK;HARD;1;CPU OK: 344 processes
[1548403986] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 4 users currently logged in
[1548406696] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 4 crit, 3 warn out of 340 processes
[1548406911] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548407031] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548407151] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1548407296] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 7 warn out of 339 processes
[1548407484] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 7.5% (1844948 kB) free!
[1548407751] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548407893] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 353 processes
[1548407897] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 9 warn out of 347 processes
[1548408312] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1548408351] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548408432] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1548408471] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 328 processes
[1548408493] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 4 warn out of 362 processes
[1548408552] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1548408617] SERVICE ALERT: sig1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1548409094] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 357 processes
[1548409097] SERVICE FLAPPING ALERT: smscapp2;CPU_Procs;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1548409097] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 10 warn out of 345 processes
[1548409284] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 85.4% (20942412 kB) free.
[1548409287] SERVICE ALERT: webservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 311 processes
[1548409372] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 89.4% (14538612 kB) free.
[1548409384] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 87.6% (21468860 kB) free.
[1548409407] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 310 processes
[1548409441] SERVICE ALERT: sig2;Memory;OK;HARD;1;OK - 93.5% (15190304 kB) free.
[1548409615] SERVICE ALERT: db1;Memory;OK;HARD;1;OK - 74.4% (36653020 kB) free.
[1548409633] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 23.7% (3849772 kB) free.
[1548409671] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 334 processes
[1548409694] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1548409694] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 4 warn out of 358 processes
[1548409697] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 5 warn out of 346 processes
[1548409753] SERVICE ALERT: webservercluster2;Memory;OK;HARD;1;OK - 92.4% (15019144 kB) free.
[1548409765] SERVICE ALERT: db2;Memory;OK;HARD;1;OK - 44.5% (21948592 kB) free.
[1548409792] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 333 processes
[1548409884] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192852 kB) free!
[1548409912] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1548410004] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171548 kB) free!
[1548410124] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1548410124] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191036 kB) free!
[1548410233] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (174168 kB) free!
[1548410294] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 355 processes
[1548410297] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 5 warn out of 346 processes
[1548410353] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (167452 kB) free!
[1548410473] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (155620 kB) free!
[1548410618] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 334 processes
[1548410873] SERVICE ALERT: sig2;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1548411010] SERVICE ALERT: sig2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 110124 MB (49% inode=99%):
[1548411112] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1548411494] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 3 warn out of 354 processes
[1548411573] SERVICE ALERT: qservercluster1;CPU Load;OK;HARD;1;OK - load average: 10.10, 9.73, 9.92
[1548411673] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 56.3% (9152260 kB) free.
[1548411701] SERVICE ALERT: smscapp2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 40762 MB (18% inode=99%):
[1548411772] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (169836 kB) free!
[1548411807] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 310 processes
[1548411892] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (166008 kB) free!
[1548411927] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 309 processes
[1548412012] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (150348 kB) free!
[1548412047] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 309 processes
[1548412301] SERVICE ALERT: smscapp2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 56307 MB (25% inode=99%):
[1548412697] SERVICE ALERT: smscapp2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 6 warn out of 343 processes
[1548412873] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.1% (1154920 kB) free!
[1548412994] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 81.1% (13186212 kB) free.
[1548413018] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 336 processes
[1548413138] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1548413294] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 3 warn out of 354 processes
[1548413320] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47170 MB (20% inode=99%):
[1548413373] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 11.08, 10.91, 10.21
[1548413440] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 47906 MB (21% inode=99%):
[1548413493] SERVICE ALERT: qservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 12.28, 11.38, 10.46
[1548413512] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 331 processes
[1548413613] SERVICE ALERT: qservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 11.95, 11.53, 10.63
[1548413632] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 329 processes
[1548413738] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1548413752] SERVICE ALERT: db1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 328 processes
[1548413858] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1548413978] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1548414186] SERVICE ALERT: db2;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1548414306] SERVICE ALERT: db2;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1548414353] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1548414426] SERVICE ALERT: db2;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1548415097] SERVICE ALERT: smscapp2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 7 warn out of 341 processes
[1548415179] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1548415328] SERVICE FLAPPING ALERT: sig3;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1548415553] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548415673] SERVICE ALERT: db1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1548415793] SERVICE ALERT: db1;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1548416979] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1548417099] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 337 processes
[1548417219] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1548417819] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 338 processes
[1548417939] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 337 processes
[1548418094] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 352 processes
[1548418694] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 4 warn out of 351 processes
[1548419139] SERVICE ALERT: db2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1548419248] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 309 processes
[1548419259] SERVICE ALERT: db2;CPU_Procs;OK;HARD;1;CPU OK: 340 processes
[1548421525] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1548424094] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 3 warn out of 349 processes
[1548424217] SERVICE ALERT: db1;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1548424337] SERVICE ALERT: db1;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1548424457] SERVICE ALERT: db1;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1548424627] SERVICE ALERT: db2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1548424697] SERVICE FLAPPING ALERT: smscapp2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1548425057] SERVICE ALERT: db1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1548425248] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 309 processes
[1548425368] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 309 processes
[1548425489] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 309 processes
[1548425952] SERVICE ALERT: webservercluster2;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1548425993] SERVICE FLAPPING ALERT: db1;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1548426459] SERVICE FLAPPING ALERT: db2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1548426689] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 308 processes
[1548428952] SERVICE ALERT: webservercluster2;Current Users;WARNING;SOFT;1;USERS WARNING - 10 users currently logged in
[1548429072] SERVICE ALERT: webservercluster2;Current Users;CRITICAL;SOFT;2;USERS CRITICAL - 11 users currently logged in
[1548429192] SERVICE ALERT: webservercluster2;Current Users;CRITICAL;HARD;3;USERS CRITICAL - 11 users currently logged in
[1548430694] SERVICE FLAPPING ALERT: pservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1548430840] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 46181 MB (20% inode=99%):
[1548430960] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 46181 MB (20% inode=99%):
[1548431080] SERVICE ALERT: pservercluster2;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46181 MB (20% inode=99%):
[1548439992] SERVICE ALERT: webservercluster2;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in