[1569263653] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1569263653] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 21.1% (5194528 kB) free.
[1569263844] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (199400 kB) free!
[1569263964] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187196 kB) free!
[1569264085] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168828 kB) free!
[1569264253] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (250380 kB) free!
[1569264295] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (256848 kB) free!
[1569264373] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169368 kB) free!
[1569264416] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (241236 kB) free!
[1569264493] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192676 kB) free!
[1569264536] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (239740 kB) free!
[1569264838] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1569264958] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1569266279] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 353 processes
[1569266294] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 43.3% (10696024 kB) free.
[1569266400] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 352 processes
[1569267346] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 212182 MB (19% inode=99%):
[1569267351] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 278 processes
[1569267466] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 202497 MB (18% inode=99%):
[1569267494] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (1993124 kB) free!
[1569267586] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 192714 MB (18% inode=99%):
[1569267588] SERVICE ALERT: smscapp2;Total Processes;OK;HARD;1;PROCS OK: 451 processes
[1569267614] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (174936 kB) free!
[1569267720] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11346976 kB) free!
[1569267734] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (173892 kB) free!
[1569267814] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (234288 kB) free!
[1569267840] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (1026528 kB) free!
[1569267933] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.2% (9466092 kB) free!
[1569267934] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (244480 kB) free!
[1569267960] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (376548 kB) free!
[1569268053] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.2% (9448764 kB) free!
[1569268055] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (226452 kB) free!
[1569268173] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.2% (9448612 kB) free!
[1569268551] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 274 processes
[1569269158] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1569269278] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1569269878] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1569269999] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 401 processes
[1569271607] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 80.66, 28.12, 10.50
[1569271727] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 11.40, 19.02, 9.30
[1569271847] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.30, 12.98, 8.26
[1569272137] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (220808 kB) free!
[1569272257] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (213424 kB) free!
[1569272377] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (233208 kB) free!
[1569272447] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.19, 2.49, 4.70
[1569272977] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.8% (43783660 kB) free.
[1569274736] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 41.4% (10138572 kB) free.
[1569274885] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 47.9% (11732988 kB) free.
[1569278534] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1569283798] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1569283918] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1569285456] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 94.5% (33100148 kB) free.
[1569286935] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 53.9% (13308812 kB) free.
[1569287760] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 66.6% (87842460 kB) free.
[1569287986] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 230032 MB (21% inode=99%):
[1569289136] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187676 kB) free!
[1569289257] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185724 kB) free!
[1569289285] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (182176 kB) free!
[1569289335] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189472 kB) free!
[1569289377] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (168600 kB) free!
[1569289405] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169352 kB) free!
[1569289455] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175128 kB) free!
[1569289525] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183824 kB) free!
[1569289575] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (165324 kB) free!
[1569289799] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 358 processes
[1569289919] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 1 warn out of 358 processes
[1569290040] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 358 processes
[1569290640] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 349 processes
[1569300322] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.3% (22785488 kB) free.
[1569300380] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.7% (22482416 kB) free.
[1569300436] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.2% (39998108 kB) free.
[1569300490] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.6% (22703428 kB) free.
[1569302055] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 582 processes
[1569302115] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 575 processes
[1569302175] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 573 processes
[1569302518] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1569302638] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1569302751] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.97, 5.90, 5.14
[1569302871] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.03, 5.55, 5.10
[1569302991] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.84, 4.92, 4.92
[1569303591] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 10.14, 10.15, 7.49
[1569303711] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 14.09, 12.32, 8.65
[1569303831] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 11.70, 12.48, 9.21
[1569304847] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.08, 11.13, 8.36
[1569304967] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.75, 7.64, 7.41
[1569305087] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.66, 5.32, 6.59
[1569305126] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 11.8% (2906824 kB) free!
[1569305246] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 11.3% (2795788 kB) free!
[1569305366] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 11.0% (2707296 kB) free!
[1569305687] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.96, 1.63, 3.96
[1569305966] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2117180 kB) free!
[1569306468] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 2 warn out of 339 processes
[1569306588] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 339 processes
[1569306638] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.2% (690188 kB) free!
[1569306727] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1569306758] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 4.2% (683664 kB) free!
[1569306848] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 97.4% (22958444 kB) free.
[1569306878] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.2% (684188 kB) free!
[1569308087] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 119.65, 84.52, 39.65
[1569308207] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 62.11, 81.39, 44.39
[1569308327] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 8.86, 54.68, 39.10
[1569308366] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1569308366] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 24.5% (6056584 kB) free.
[1569308678] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (3191640 kB) free!
[1569309350] SERVICE ALERT: primary-server;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 1 warn out of 347 processes
[1569309527] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1569309527] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.85, 1.99, 11.55
[1569310647] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.7% (14253992 kB) free.
[1569310727] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.81, 0.90, 3.86
[1569310778] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.4% (9068640 kB) free!
[1569310898] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.0% (8891456 kB) free!
[1569311018] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.6% (8676820 kB) free!
[1569311366] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 16.5% (4077908 kB) free!
[1569311486] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.1% (3967564 kB) free!
[1569311606] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.6% (3844008 kB) free!
[1569312206] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.3% (825156 kB) free!
[1569312341] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 20553 MB (14% inode=99%):
[1569312461] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 20554 MB (14% inode=99%):
[1569312581] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 20554 MB (14% inode=99%):
[1569312818] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (4235796 kB) free!
[1569313547] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1569313607] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1569313667] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1569313727] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1569314031] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.68, 3.74, 4.55
[1569314186] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.7% (4331908 kB) free!
[1569314206] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.8% (23152992 kB) free.
[1569314267] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.8% (22754636 kB) free.
[1569314352] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 96.1% (47374280 kB) free.
[1569314627] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 58 users currently logged in
[1569314750] SERVICE ALERT: primary-server;CPU_Procs;WARNING;HARD;3;CPU WARNING: 2 warn out of 346 processes
[1569315104] SERVICE ALERT: primary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1569315224] SERVICE ALERT: primary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1569315824] SERVICE ALERT: primary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1569315944] SERVICE ALERT: primary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1569316064] SERVICE ALERT: primary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1569317147] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1569317267] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1569317387] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1569317981] SERVICE ALERT: sig2-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 45074 MB (32% inode=99%):
[1569318418] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174396 kB) free!
[1569318474] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164852 kB) free!
[1569318490] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.1% (3464896 kB) free!
[1569318538] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195776 kB) free!
[1569318594] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194048 kB) free!
[1569318610] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170236 kB) free!
[1569318658] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191112 kB) free!
[1569318714] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (166140 kB) free!
[1569318730] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188460 kB) free!
[1569319064] SERVICE ALERT: primary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1569319427] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1569319487] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1569319547] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1569319583] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.0% (22697988 kB) free.
[1569319607] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 62 users currently logged in
[1569319620] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.4% (22642368 kB) free.
[1569319664] SERVICE ALERT: primary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1569319784] SERVICE ALERT: primary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1569319794] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.4% (22904148 kB) free.
[1569319904] SERVICE ALERT: primary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1569320038] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1569320158] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1569320183] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.9% (1455268 kB) free!
[1569320245] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (225712 kB) free!
[1569320303] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192720 kB) free!
[1569320365] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (237072 kB) free!
[1569320423] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174916 kB) free!
[1569320485] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (240492 kB) free!
[1569320859] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 304 processes
[1569320979] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 307 processes
[1569321487] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569321527] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1569322628] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 11.3% (2767384 kB) free!
[1569322687] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569322748] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 10.7% (2614212 kB) free!
[1569322756] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.0% (9349804 kB) free!
[1569322868] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2455328 kB) free!
[1569322876] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9065032 kB) free!
[1569322996] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 17.8% (8759756 kB) free!
[1569323038] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.53, 6.47, 5.71
[1569323158] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.86, 6.39, 5.79
[1569323278] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.54, 6.41, 5.86
[1569323287] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569323631] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.48, 5.14, 5.13
[1569323751] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.32, 4.99, 5.09
[1569323871] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 4.33, 4.78, 5.00
[1569324314] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.2% (38557632 kB) free.
[1569324361] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 88.4% (21807996 kB) free.
[1569324406] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1569324406] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.5% (22428848 kB) free.
[1569324466] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 96.2% (47388124 kB) free.
[1569324487] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1569324487] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 606 processes
[1569324927] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (260564 kB) free!
[1569324964] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174516 kB) free!
[1569325047] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (261912 kB) free!
[1569325078] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.77, 4.21, 4.97
[1569325084] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (192512 kB) free!
[1569325167] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (263388 kB) free!
[1569325204] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191272 kB) free!
[1569325907] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1569326207] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 63 users currently logged in
[1569326504] SERVICE ALERT: primary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1569327471] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.31, 8.51, 6.24
[1569327581] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 28754 MB (20% inode=99%):
[1569327591] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 5.98, 7.46, 6.13
[1569327701] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 28754 MB (20% inode=99%):
[1569327711] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.84, 6.92, 6.09
[1569327821] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 28756 MB (20% inode=99%):
[1569328078] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 2.27, 6.84, 6.05
[1569328198] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.13, 5.26, 5.56
[1569328318] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 2.07, 4.23, 5.15
[1569328688] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569328738] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.0% (22695964 kB) free.
[1569328816] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1569328816] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.1% (39476384 kB) free.
[1569328918] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1569328918] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.15, 2.77, 3.73
[1569329359] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198324 kB) free!
[1569329479] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (186856 kB) free!
[1569329599] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190976 kB) free!
[1569329870] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 90.5% (22331572 kB) free.
[1569329888] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 604 processes
[1569329988] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 2 warn out of 343 processes
[1569330108] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 340 processes
[1569330228] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 335 processes
[1569330488] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569330708] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 57 users currently logged in
[1569331675] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189996 kB) free!
[1569331688] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569331795] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (187928 kB) free!
[1569331915] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186376 kB) free!
[1569332289] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569332511] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.55, 3.48, 4.79
[1569333111] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.79, 5.43, 5.55
[1569333221] SERVICE ALERT: sig2-vm2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 12347 MB (8% inode=99%):
[1569333231] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.42, 4.27, 5.11
[1569333352] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 1.65, 3.40, 4.69
[1569333413] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1569333489] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569333608] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194720 kB) free!
[1569333728] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185636 kB) free!
[1569333848] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (177372 kB) free!
[1569333951] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 9.43, 7.35, 5.79
[1569334013] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (167712 kB) free!
[1569334071] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.86, 7.06, 5.87
[1569334089] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569334133] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (197108 kB) free!
[1569334191] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 4.22, 6.02, 5.64
[1569334257] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (191604 kB) free!
[1569334791] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1569334791] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.51, 3.98, 4.89
[1569335038] SERVICE ALERT: db2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 2 crit, 1 warn out of 334 processes
[1569335289] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569335621] SERVICE ALERT: sig2-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 36869 MB (26% inode=99%):
[1569336022] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1569336238] SERVICE ALERT: db2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 326 processes
[1569336591] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.32, 7.28, 5.70
[1569336711] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 3.84, 5.93, 5.39
[1569336831] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.29, 4.60, 4.97
[1569339490] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569340179] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 305 processes
[1569340299] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 311 processes
[1569340318] SERVICE FLAPPING ALERT: pservercluster1;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569340690] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569340727] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 56.09, 12.84, 4.90
[1569340847] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 12.17, 12.79, 6.06
[1569340967] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 2.01, 8.73, 5.39
[1569341290] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569341515] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1569341567] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.40, 1.63, 3.09
[1569342490] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569343091] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569344291] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569344891] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569346091] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569346367] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 15.88, 21.99, 10.37
[1569346431] SERVICE FLAPPING ALERT: smscapp2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569346487] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 2.54, 14.89, 9.18
[1569346607] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.83, 10.16, 8.13
[1569346692] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569347207] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1569347207] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.71, 1.91, 4.56
[1569347892] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569348492] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569349692] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes