[1572978602] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (25128800 kB) free!
[1572978723] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40261080 kB) free.
[1572979139] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (192408 kB) free!
[1572979180] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (174180 kB) free!
[1572979261] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (180548 kB) free!
[1572979301] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (185336 kB) free!
[1572979322] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25449584 kB) free!
[1572979340] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (195900 kB) free!
[1572979383] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (179124 kB) free!
[1572979421] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183112 kB) free!
[1572979442] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40269880 kB) free.
[1572980042] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24905028 kB) free!
[1572980162] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40286228 kB) free.
[1572980590] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.0% (1709940 kB) free!
[1572980650] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195208 kB) free!
[1572980710] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 62.3% (15260084 kB) free.
[1572980762] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24514936 kB) free!
[1572980882] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40297044 kB) free.
[1572981482] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24262108 kB) free!
[1572981602] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40297132 kB) free.
[1572982202] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.2% (23967564 kB) free!
[1572982322] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39521904 kB) free.
[1572982400] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 32.2% (7957124 kB) free.
[1572983001] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191092 kB) free!
[1572983027] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 294 processes
[1572983121] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183972 kB) free!
[1572983147] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 302 processes
[1572983241] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187660 kB) free!
[1572983522] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25414420 kB) free!
[1572983642] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40341040 kB) free.
[1572984728] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 744.52, 199.90, 69.69
[1572984843] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.4% (24247988 kB) free!
[1572984848] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 158.62, 183.53, 81.83
[1572984963] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39257760 kB) free.
[1572985051] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(226) < 2019/11/06 01:44:20 [error] 15232#0: *1487966 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1572985171] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572985352] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1572985472] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1572986163] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.9% (23603424 kB) free!
[1572986283] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38850212 kB) free.
[1572986381] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.7% (15456176 kB) free!
[1572987372] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 14.4% (7117576 kB) free!
[1572987474] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (238868 kB) free!
[1572987492] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 85.4% (42093044 kB) free.
[1572987594] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (243176 kB) free!
[1572987714] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (216812 kB) free!
[1572988083] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.4% (25635084 kB) free!
[1572988203] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40436208 kB) free.
[1572988314] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 92.9% (45789168 kB) free.
[1572988472] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1572988593] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1572988641] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 48.0% (11845720 kB) free.
[1572988803] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24549356 kB) free!
[1572988923] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39860220 kB) free.
[1572989241] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (1420020 kB) free!
[1572989361] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.8% (690832 kB) free!
[1572989482] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1572989482] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (401452 kB) free!
[1572989523] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23787268 kB) free!
[1572989643] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.5% (38924296 kB) free.
[1572990250] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(8) < 2019/11/06 03:07:45 [error] 24130#0: *2892324 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572990370] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572991443] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24791536 kB) free!
[1572991564] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.1% (39709764 kB) free.
[1572991752] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 356 processes
[1572991872] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 372 processes
[1572991992] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 364 processes
[1572992170] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(32) < 2019/11/06 03:46:10 [error] 24133#0: *2893501 connect() failed (111: Connection refused) while connecting to upstream, client: 169.38.76.212, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1572992283] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 344 processes
[1572992291] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(6) < 2019/11/06 03:46:36 [error] 24130#0: *2893443 recv() failed (104: Connection reset by peer) while proxying connection, client: 121.241.242.109, server: 0.0.0.0:13000, upstream: "10.147.212.143:13000", bytes from/to client:43/0, bytes from/to upstream:0/43
[1572992404] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 350 processes
[1572992411] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;HARD;3;(3) < 2019/11/06 03:49:57 [error] 24130#0: *2893577 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1572992524] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1572992584] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 11.1% (2719196 kB) free!
[1572992592] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 342 processes
[1572992621] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 11.6% (2838384 kB) free!
[1572992763] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25456084 kB) free!
[1572992883] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40492924 kB) free.
[1572993124] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1572993140] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 11.0% (2703940 kB) free!
[1572993611] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1572994083] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (25974684 kB) free!
[1572994203] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40520540 kB) free.
[1572994803] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24549260 kB) free!
[1572994923] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39256236 kB) free.
[1572996123] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24731728 kB) free!
[1572996140] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (2432220 kB) free!
[1572996184] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2442488 kB) free!
[1572996243] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.0% (39528856 kB) free.
[1572997443] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.3% (25417144 kB) free!
[1572997563] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (40562344 kB) free.
[1572997593] SERVICE ALERT: redis-1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 401 processes
[1572997713] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1572998163] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24803744 kB) free!
[1572998284] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.2% (39865640 kB) free.
[1572998733] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4927332 kB) free!
[1572999221] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2443264 kB) free!
[1572999483] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (25647468 kB) free!
[1572999603] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.8% (40579372 kB) free.
[1573000203] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.7% (25966312 kB) free!
[1573000323] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.5% (23077904 kB) free!
[1573000443] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.8% (39290352 kB) free.
[1573000713] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 407 processes
[1573000781] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 20.2% (26647660 kB) free.
[1573000833] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 407 processes
[1573000882] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573000970] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 80.9% (19972596 kB) free.
[1573001040] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 88.9% (21800800 kB) free.
[1573001043] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24404056 kB) free!
[1573001090] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.0% (40432316 kB) free.
[1573001133] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 87.6% (43184632 kB) free.
[1573001163] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40398096 kB) free.
[1573001177] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 57.7% (14151188 kB) free.
[1573001219] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 93.9% (23015236 kB) free.
[1573001579] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181764 kB) free!
[1573001699] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176068 kB) free!
[1573001763] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.0% (25026004 kB) free!
[1573001820] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196448 kB) free!
[1573001883] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.6% (40290452 kB) free.
[1573002011] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(2) < 2019/11/06 06:30:08 [error] 24131#0: *2922785 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.154:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573002131] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;2;(13) < 2019/11/06 06:30:49 [error] 24130#0: *2922946 connect() failed (111: Connection refused) while connecting to upstream, client: 103.255.146.164, server: 0.0.0.0:3000, upstream: "10.147.212.143:3000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573002251] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573002483] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.1% (25141072 kB) free!
[1573002603] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.5% (40194076 kB) free.
[1573002898] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.7% (9223060 kB) free!
[1573003018] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.7% (9215808 kB) free!
[1573003138] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.9% (6831912 kB) free!
[1573003738] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (4013896 kB) free!
[1573004220] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 15.0% (3713280 kB) free!
[1573004233] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1573004233] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.0% (22956616 kB) free.
[1573004274] SERVICE FLAPPING ALERT: sigqscluster1;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1573004274] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.0% (40423860 kB) free.
[1573004840] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185864 kB) free!
[1573004960] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199964 kB) free!
[1573005080] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194696 kB) free!
[1573006203] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.7% (24591016 kB) free!
[1573006323] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35636880 kB) free.
[1573006678] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.1% (8409276 kB) free!
[1573006799] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.1% (8408356 kB) free!
[1573006919] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.1% (8407436 kB) free!
[1573006923] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.5% (25654092 kB) free!
[1573007043] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 17.2% (22610988 kB) free!
[1573007163] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38322100 kB) free.
[1573007519] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (471300 kB) free!
[1573007763] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.3% (20185840 kB) free!
[1573007828] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.0% (40424588 kB) free.
[1573007867] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.8% (22897084 kB) free.
[1573007883] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38248672 kB) free.
[1573008435] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 13.8% (6797348 kB) free!
[1573008476] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191448 kB) free!
[1573008483] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.2% (21365352 kB) free!
[1573008555] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 13.8% (6798144 kB) free!
[1573008597] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (184864 kB) free!
[1573008603] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38540960 kB) free.
[1573008676] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.8% (6795428 kB) free!
[1573008717] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189092 kB) free!
[1573009203] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.0% (22347288 kB) free!
[1573009276] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (1328112 kB) free!
[1573009323] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38491404 kB) free.
[1573009923] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (22021016 kB) free!
[1573010043] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.2% (38428832 kB) free.
[1573010643] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.7% (22079112 kB) free!
[1573010764] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.1% (38327580 kB) free.
[1573011364] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.4% (20299868 kB) free!
[1573011484] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.5% (37619712 kB) free.
[1573012084] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.2% (17421272 kB) free!
[1573012204] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36637708 kB) free.
[1573012233] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1573012353] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1573012804] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.4% (16306692 kB) free!
[1573012924] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.3% (35945556 kB) free.
[1573013524] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.1% (15972596 kB) free!
[1573013644] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.9% (35458456 kB) free.
[1573014244] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14891632 kB) free!
[1573014364] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.2% (34593292 kB) free.
[1573014491] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.1% (22737004 kB) free.
[1573014538] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.0% (40421204 kB) free.
[1573014964] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (12887492 kB) free!
[1573015084] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.6% (33737324 kB) free.
[1573015093] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 12.2% (1986348 kB) free!
[1573015213] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;2;WARNING - 12.2% (1989448 kB) free!
[1573015333] SERVICE ALERT: webservercluster1;Memory;WARNING;HARD;3;WARNING - 12.2% (1988364 kB) free!
[1573015677] SERVICE ALERT: sig2-vm2;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.48, 7.24, 5.63
[1573015684] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11854740 kB) free!
[1573015797] SERVICE ALERT: sig2-vm2;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.90, 7.47, 5.90
[1573015805] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.1% (33093616 kB) free.
[1573015863] SERVICE ALERT: sig2-vm2;CPU Load;OK;HARD;1;OK - load average: 10.24, 8.17, 6.24
[1573016382] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.2% (24050620 kB) free!
[1573016404] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.6% (11344324 kB) free!
[1573016502] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 16.8% (22105772 kB) free!
[1573016524] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32986520 kB) free.
[1573016622] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 22.9% (30189960 kB) free.
[1573016786] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 17.2% (4227892 kB) free!
[1573016897] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187056 kB) free!
[1573016907] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 16.6% (4065672 kB) free!
[1573017017] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (177392 kB) free!
[1573017027] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 15.9% (3901064 kB) free!
[1573017124] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10311944 kB) free!
[1573017137] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (187784 kB) free!
[1573017222] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 16.8% (22096816 kB) free!
[1573017244] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 22.8% (30088368 kB) free.
[1573017247] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (1798348 kB) free!
[1573017342] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 18.0% (23785068 kB) free!
[1573017367] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.6% (1627316 kB) free!
[1573017462] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 19.5% (25665300 kB) free!
[1573017488] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (1372020 kB) free!
[1573017548] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.1% (2008980 kB) free!
[1573017571] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(166) < 2019/11/06 10:48:19 [error] 15231#0: *1526503 recv() failed (104: Connection reset by peer) while proxying connection, client: 103.16.101.72, server: 0.0.0.0:4000, upstream: "10.147.212.49:4000", bytes from/to client:45/0, bytes from/to upstream:0/45
[1573017668] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.4% (1667280 kB) free!
[1573017691] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573017788] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 4.4% (2154904 kB) free!
[1573017844] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.1% (8099812 kB) free!
[1573017964] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.8% (27473548 kB) free.
[1573018324] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 351 processes
[1573018444] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 341 processes
[1573018564] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.4% (5822796 kB) free!
[1573018684] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27690032 kB) free.
[1573018820] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.8% (20915912 kB) free.
[1573018827] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (1376188 kB) free!
[1573018863] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 91.4% (22404588 kB) free.
[1573018900] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.5% (40178508 kB) free.
[1573018950] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 51.5% (12620252 kB) free.
[1573019262] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 21.5% (28381724 kB) free.
[1573019284] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.5% (5889316 kB) free!
[1573019404] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 20.6% (27119204 kB) free.
[1573019428] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (198272 kB) free!
[1573019520] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;1;WARNING - 13.8% (3092692 kB) free!
[1573019548] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (189952 kB) free!
[1573019640] SERVICE ALERT: sig6-vm2;Memory;WARNING;SOFT;2;WARNING - 12.3% (2770768 kB) free!
[1573019668] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (194884 kB) free!
[1573019760] SERVICE ALERT: sig6-vm2;Memory;WARNING;HARD;3;WARNING - 10.9% (2442108 kB) free!
[1573019862] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;1;WARNING - 18.8% (24785008 kB) free!
[1573019982] SERVICE ALERT: redis-1;Memory;WARNING;SOFT;2;WARNING - 20.0% (26421384 kB) free!
[1573020004] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 4.9% (6479000 kB) free!
[1573020102] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 16.1% (21209508 kB) free!
[1573020124] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 1.7% (2235500 kB) free!
[1573020244] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.9% (23621564 kB) free!
[1573020360] SERVICE ALERT: sig6-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 2.7% (607700 kB) free!
[1573020755] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(1) < Channel error on connection <0.21194.81> (10.147.212.181:48959 -> 10.147.212.160:5672, vhost: '/', user: 'guest'), channel 1:
[1573020845] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (1756328 kB) free!
[1573020875] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573020879] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;1;WARNING - 15.6% (5444632 kB) free!
[1573020999] SERVICE ALERT: sig2-vm2;Memory;WARNING;SOFT;2;WARNING - 14.6% (5086908 kB) free!
[1573021119] SERVICE ALERT: sig2-vm2;Memory;WARNING;HARD;3;WARNING - 13.6% (4728788 kB) free!
[1573021719] SERVICE ALERT: sig2-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.4% (2933388 kB) free!
[1573022160] SERVICE ALERT: sig6-vm2;Memory;OK;HARD;1;OK - 82.7% (18560056 kB) free.
[1573022319] SERVICE ALERT: sig2-vm2;Memory;OK;HARD;1;OK - 92.5% (32227596 kB) free.
[1573022511] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 18.8% (9267104 kB) free!
[1573022631] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.0% (8374716 kB) free!
[1573022636] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.9% (22667488 kB) free.
[1573022645] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 13.9% (18375452 kB) free!
[1573022689] SERVICE FLAPPING ALERT: redis-2;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1573022689] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 42.8% (56428612 kB) free.
[1573022751] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.0% (8371672 kB) free!
[1573022786] SERVICE FLAPPING ALERT: redis-1;Memory;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1573022786] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 85.1% (112180792 kB) free.
[1573022822] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.0% (40392440 kB) free.
[1573023759] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.7% (3852496 kB) free!
[1573023879] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.9% (3658600 kB) free!
[1573023897] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.8% (24739972 kB) free!
[1573023999] SERVICE FLAPPING ALERT: sservercluster1;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1573023999] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 14.1% (3455628 kB) free!
[1573024017] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42900792 kB) free.
[1573024439] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.9% (467908 kB) free!
[1573024559] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (169540 kB) free!
[1573024617] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 16.8% (22092144 kB) free!
[1573024679] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (189740 kB) free!
[1573024737] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.8% (18149532 kB) free!
[1573024857] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.5% (41460260 kB) free.
[1573024914] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.9% (9799440 kB) free!
[1573025034] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 19.2% (9457184 kB) free!
[1573025154] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.5% (9113048 kB) free!
[1573025199] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (1678672 kB) free!
[1573025457] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.4% (18918460 kB) free!
[1573025577] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.3% (14891600 kB) free!
[1573025697] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.0% (38254844 kB) free.
[1573025911] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.3% (214500 kB) free!
[1573026066] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.2% (4451396 kB) free!
[1573026186] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.5% (4282104 kB) free!
[1573026297] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15843480 kB) free!
[1573026306] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 16.8% (4114832 kB) free!
[1573026417] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.4% (38722488 kB) free.
[1573026603] SERVICE ALERT: pservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 16 users currently logged in
[1573026666] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 91.2% (22506268 kB) free.
[1573026709] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 50.7% (12436100 kB) free.
[1573026723] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 15 users currently logged in
[1573026783] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.0% (14632892 kB) free.
[1573026954] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (3976424 kB) free!
[1573027018] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.4% (17609296 kB) free!
[1573027138] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 10.3% (13584336 kB) free!
[1573027258] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.8% (36708308 kB) free.
[1573027507] SERVICE FLAPPING ALERT: pservercluster2;Memory;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1573027507] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (2207820 kB) free!
[1573027858] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.9% (15673592 kB) free!
[1573027978] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28689652 kB) free.
[1573028229] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.0% (4652392 kB) free!
[1573028349] SERVICE ALERT: sservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.8% (4600800 kB) free!
[1573028469] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (4516780 kB) free!
[1573028491] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(1) < 2019/11/06 13:41:54 [error] 15225#0: *1546795 open() "/usr/share/nginx/html/favicon.ico" failed (2: No such file or directory), client: 183.82.105.122, server: 148qservercluster1.com, request: "GET /favicon.ico HTTP/1.1", host: "qservercluster1.com:8091", referrer: "http://qservercluster1.com:8091/"
[1573028553] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 401 processes
[1573028578] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 14.3% (18841008 kB) free!
[1573028611] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573028651] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(60) < 2019/11/06 13:53:31 [error] 24130#0: *3003660 recv() failed (104: Connection reset by peer) while proxying connection, client: 203.212.70.170, server: 0.0.0.0:3001, upstream: "10.147.212.143:3001", bytes from/to client:62020/16682, bytes from/to upstream:16682/62020
[1573028673] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1573028698] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.4% (15060704 kB) free!
[1573028771] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573028818] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.6% (36361060 kB) free.
[1573029211] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(209) < 2019/11/06 14:01:52 [error] 15228#0: *1550610 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.67.175, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:474/1950, bytes from/to upstream:1950/474
[1573029273] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 400 processes
[1573029331] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573029394] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes
[1573029418] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.7% (16690284 kB) free!
[1573029538] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.7% (12829676 kB) free!
[1573029658] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.7% (35146676 kB) free.
[1573030258] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14086756 kB) free!
[1573030378] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 23.6% (31073656 kB) free.
[1573030917] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.3% (4739104 kB) free!
[1573030978] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.0% (15776424 kB) free!
[1573031037] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.3% (4479480 kB) free!
[1573031098] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.9% (11669884 kB) free!
[1573031157] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.3% (4241080 kB) free!
[1573031218] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.0% (32977576 kB) free.
[1573031469] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2131060 kB) free!
[1573031471] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177956 kB) free!
[1573031592] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (173992 kB) free!
[1573031712] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190084 kB) free!
[1573031786] SERVICE FLAPPING ALERT: redis-1;Memory;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1573031819] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11835136 kB) free!
[1573031939] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.2% (8155452 kB) free!
[1573032059] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.7% (23307440 kB) free!
[1573032357] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (2245404 kB) free!
[1573032426] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.0% (7892740 kB) free!
[1573032546] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 11.3% (5548800 kB) free!
[1573032659] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2160080 kB) free!
[1573032666] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 11.3% (5549692 kB) free!
[1573033092] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 18.9% (9333736 kB) free!
[1573033212] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 18.7% (9199552 kB) free!
[1573033266] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.1% (4493976 kB) free!
[1573033332] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 18.4% (9081152 kB) free!
[1573033391] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (166708 kB) free!
[1573033511] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (178616 kB) free!
[1573033631] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (174396 kB) free!
[1573033859] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 11.6% (15279404 kB) free!
[1573035059] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.5% (1958252 kB) free!
[1573035972] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(18) < 2019/11/06 15:46:44 [error] 24131#0: *3025080 connect() failed (111: Connection refused) while connecting to upstream, client: 94.237.68.73, server: 0.0.0.0:13000, upstream: "10.147.212.142:13000", bytes from/to client:0/0, bytes from/to upstream:0/0
[1573036092] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573036259] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 12.3% (16234076 kB) free!
[1573036692] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(17) < 2019/11/06 16:03:37 [error] 24131#0: *3026131 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.155:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573036812] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573036869] SERVICE ALERT: sservercluster2;Memory;WARNING;HARD;3;WARNING - 10.1% (2484796 kB) free!
[1573036929] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 93.3% (23028696 kB) free.
[1573036988] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 90.9% (22287404 kB) free.
[1573037059] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.8% (40288788 kB) free.
[1573037132] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 95.7% (47161824 kB) free.
[1573037172] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 50.2% (12296772 kB) free.
[1573037245] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 90.1% (14641340 kB) free.
[1573037291] SERVICE ALERT: qservercluster2;Memory;OK;HARD;1;OK - 94.8% (46719116 kB) free.
[1573037459] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 2.8% (3720916 kB) free!
[1573037469] SERVICE FLAPPING ALERT: sservercluster2;Memory;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1573037469] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.7% (2140044 kB) free!
[1573037569] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 38.9% (51282412 kB) free.
[1573037622] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 95.0% (23292080 kB) free.
[1573037793] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1573037914] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1573038136] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 13.8% (3402152 kB) free!
[1573038245] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 407 processes
[1573038256] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167796 kB) free!
[1573038365] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 415 processes
[1573038376] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188992 kB) free!
[1573038485] SERVICE ALERT: sservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 465 processes
[1573038612] SERVICE ALERT: primary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(103) < 2019/11/06 16:38:50 [error] 24132#0: *3027461 recv() failed (104: Connection reset by peer) while proxying connection, client: 180.179.244.32, server: 0.0.0.0:13000, upstream: "10.147.212.154:13000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573038733] SERVICE ALERT: primary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573039085] SERVICE ALERT: sservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 336 processes
[1573039981] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 13.6% (17981092 kB) free!
[1573040101] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.0% (42236028 kB) free.
[1573040701] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 15.2% (20014820 kB) free!
[1573040821] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 11.9% (15699516 kB) free!
[1573040941] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 30.7% (40510724 kB) free.
[1573041199] SERVICE FLAPPING ALERT: pservercluster2;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1573041541] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 12.8% (16815032 kB) free!
[1573041661] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 9.6% (12646996 kB) free!
[1573041781] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 29.3% (38659836 kB) free.
[1573041866] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 17.9% (8818300 kB) free!
[1573041931] SERVICE ALERT: secondary-server;Nginx Error Log Status;CRITICAL;SOFT;1;(263) < 2019/11/06 17:33:20 [error] 15229#0: *1571830 recv() failed (104: Connection reset by peer) while proxying connection, client: 49.50.86.152, server: 0.0.0.0:4000, upstream: "10.147.212.61:4000", bytes from/to client:42/0, bytes from/to upstream:0/42
[1573041979] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.0% (4651964 kB) free!
[1573041986] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 17.9% (8814440 kB) free!
[1573042051] SERVICE ALERT: secondary-server;Nginx Error Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1573042099] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 17.7% (4341296 kB) free!
[1573042106] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 17.9% (8805036 kB) free!
[1573042219] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 16.5% (4044468 kB) free!
[1573042381] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 11.3% (14882312 kB) free!
[1573042501] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 8.2% (10745180 kB) free!
[1573042622] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 28.2% (37181824 kB) free.
[1573043221] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 10.7% (14116784 kB) free!
[1573043311] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 82.43, 361.89, 240.87
[1573043341] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 7.5% (9835120 kB) free!
[1573043419] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 5.9% (1446380 kB) free!
[1573043431] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 421.04, 345.81, 246.51
[1573043461] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.4% (36081340 kB) free.
[1573043551] SERVICE ALERT: smscapp2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 98.91, 255.61, 226.22
[1573043861] SERVICE ALERT: rq-1;rq1-rabbitmqerror;CRITICAL;SOFT;1;(2) < webmachine error: path="/api/whoami"
[1573043982] SERVICE ALERT: rq-1;rq1-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573044061] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (11622512 kB) free!
[1573044181] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (7510896 kB) free!
[1573044275] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;CRITICAL;SOFT;1;(2) < webmachine error: path="/api/whoami"
[1573044301] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 27.0% (35650384 kB) free.
[1573044396] SERVICE ALERT: sigqscluster2;sigqscluster2-rabbitmqerror;OK;HARD;1;Log check ok - 0 pattern matches found
[1573044506] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.0% (3941788 kB) free!
[1573044901] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.0% (11927568 kB) free!
[1573045021] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (8022940 kB) free!
[1573045141] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 26.6% (35074204 kB) free.
[1573045741] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10709796 kB) free!
[1573045861] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 5.0% (6627028 kB) free!
[1573045981] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 25.7% (33901532 kB) free.
[1573046551] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 21.71, 21.59, 31.21
[1573046581] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 8.1% (10677376 kB) free!
[1573046702] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.8% (6388980 kB) free!
[1573046822] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.0% (27628572 kB) free.
[1573047151] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 24.41, 24.02, 27.61
[1573047422] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.8% (10294280 kB) free!
[1573047542] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 4.7% (6134576 kB) free!
[1573047662] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.6% (2095188 kB) free!
[1573047829] SERVICE FLAPPING ALERT: sservercluster2;Memory;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573047976] SERVICE FLAPPING ALERT: pservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1573048862] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 16.1% (21171064 kB) free!
[1573049462] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 21.8% (28712604 kB) free.
[1573050062] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;1;CRITICAL - 6.0% (7889364 kB) free!
[1573050182] SERVICE ALERT: redis-2;Memory;CRITICAL;SOFT;2;CRITICAL - 3.1% (4091152 kB) free!
[1573050302] SERVICE ALERT: redis-2;Memory;WARNING;HARD;3;WARNING - 17.4% (22995064 kB) free!
[1573050902] SERVICE ALERT: redis-2;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (1330128 kB) free!
[1573051805] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 92.3% (22787196 kB) free.
[1573051816] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 39.8% (52408120 kB) free.
[1573051819] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 52.4% (12846704 kB) free.
[1573051825] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 82.1% (40469980 kB) free.
[1573053008] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189512 kB) free!
[1573053031] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.8% (7298124 kB) free!
[1573053128] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (170080 kB) free!
[1573053151] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 13.6% (6707292 kB) free!
[1573053248] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (196628 kB) free!
[1573053272] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 13.6% (6697652 kB) free!
[1573053872] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (1910696 kB) free!
[1573054216] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.1% (22535904 kB) free!
[1573054336] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 13.8% (18184476 kB) free!
[1573054456] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42626272 kB) free.
[1573055056] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.8% (23405968 kB) free!
[1573055177] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.0% (19823212 kB) free!
[1573055297] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.3% (42513980 kB) free.
[1573055897] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 17.7% (23360072 kB) free!
[1573056017] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.1% (19952964 kB) free!
[1573056137] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.7% (43121336 kB) free.
[1573056737] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.0% (23690184 kB) free!
[1573056857] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.2% (20032180 kB) free!
[1573056977] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.6% (42995736 kB) free.
[1573057577] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.5% (24353072 kB) free!
[1573057697] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 15.9% (20930504 kB) free!
[1573057817] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42895880 kB) free.
[1573057819] SERVICE FLAPPING ALERT: sservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1573058417] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.9% (24886224 kB) free!
[1573058537] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.2% (21374524 kB) free!
[1573058657] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.5% (42799328 kB) free.
[1573059257] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25288596 kB) free!
[1573059378] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.6% (21914864 kB) free!
[1573059498] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.4% (42745164 kB) free.
[1573059799] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.9% (4865440 kB) free!
[1573059919] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.8% (4854300 kB) free!
[1573060039] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 19.8% (4842884 kB) free!
[1573061740] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (225656 kB) free!
[1573061860] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.4% (221280 kB) free!
[1573061898] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 18.6% (24566112 kB) free!
[1573061913] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1573061980] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (232780 kB) free!
[1573062018] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.2% (21310700 kB) free!
[1573062019] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.6% (4804980 kB) free!
[1573062033] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1573062138] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 31.8% (41920492 kB) free.
[1573062139] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.5% (4778692 kB) free!
[1573062259] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.4% (4749244 kB) free!
[1573062580] SERVICE ALERT: smscapp2;Memory;WARNING;HARD;3;WARNING - 10.4% (5114972 kB) free!
[1573062738] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;1;WARNING - 19.2% (25343440 kB) free!
[1573062858] SERVICE ALERT: redis-2;Memory;WARNING;SOFT;2;WARNING - 16.8% (22180116 kB) free!
[1573062978] SERVICE ALERT: redis-2;Memory;OK;HARD;1;OK - 32.4% (42669804 kB) free.
[1573063180] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4927860 kB) free!
[1573064778] SERVICE FLAPPING ALERT: redis-2;Memory;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)