[1596739110] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596739168] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (187388 kB) free!
[1596739190] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177768 kB) free!
[1596739240] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596739288] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (198800 kB) free!
[1596739310] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (199004 kB) free!
[1596739350] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183024 kB) free!
[1596739370] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596739409] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (192672 kB) free!
[1596739430] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172032 kB) free!
[1596739470] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175180 kB) free!
[1596739590] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 41.7% (10283260 kB) free.
[1596739971] SERVICE ALERT: pservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596740010] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 79.6% (19503040 kB) free.
[1596740030] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 83.3% (20420008 kB) free.
[1596740191] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (165640 kB) free!
[1596740311] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (194008 kB) free!
[1596740431] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 85.3% (21053088 kB) free.
[1596742637] SERVICE ALERT: webserver2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186624 kB) free!
[1596742697] SERVICE ALERT: webserver2;Memory;OK;HARD;1;OK - 70.6% (17295832 kB) free.
[1596743144] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.0% (4412200 kB) free!
[1596743538] SERVICE FLAPPING ALERT: mrq4;CPU Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1596743538] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 65.29, 54.93, 46.93
[1596743959] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (181208 kB) free!
[1596744014] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (803136 kB) free!
[1596744079] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 85.4% (20940288 kB) free.
[1596744134] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (804116 kB) free!
[1596744138] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 33.81, 39.44, 42.82
[1596744244] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.4% (214384 kB) free!
[1596744254] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (806292 kB) free!
[1596744364] SERVICE ALERT: smscapp2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (239296 kB) free!
[1596744485] SERVICE ALERT: smscapp2;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (234032 kB) free!
[1596744499] SERVICE FLAPPING ALERT: sservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1596744679] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 8.8% (2166692 kB) free!
[1596744799] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (164352 kB) free!
[1596744920] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 90.2% (22114828 kB) free.
[1596745519] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.8% (923340 kB) free!
[1596745639] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183460 kB) free!
[1596745759] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (186744 kB) free!
[1596745938] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 41.79, 48.16, 46.01
[1596746285] SERVICE ALERT: smscapp2;Memory;OK;HARD;1;OK - 68.4% (33699248 kB) free.
[1596746302] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(7) < 103.255.146.164 [07/Aug/2020:02:07:18 +0530] TCP 502 0 0 0.000 "10.147.212.142:3000" "0" "0" "0.000"
[1596746423] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596746431] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (190728 kB) free!
[1596746538] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 37.51, 39.78, 43.06
[1596746551] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (171088 kB) free!
[1596746610] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (164036 kB) free!
[1596746630] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (185136 kB) free!
[1596746671] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (176748 kB) free!
[1596746732] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 79.3% (19433868 kB) free.
[1596746751] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 80.5% (19732168 kB) free.
[1596746959] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 44.7% (10957344 kB) free.
[1596747271] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 73.1% (18028612 kB) free.
[1596747559] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 2.0% (493044 kB) free!
[1596747679] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (183404 kB) free!
[1596747800] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 69.7% (17085040 kB) free.
[1596747818] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.6% (286392 kB) free!
[1596747938] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (285316 kB) free!
[1596748058] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.6% (283780 kB) free!
[1596748338] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 44.09, 47.95, 45.17
[1596748399] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 15.1% (3688956 kB) free!
[1596748519] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 93.1% (22825292 kB) free.
[1596748658] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 76.5% (37694376 kB) free.
[1596748938] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 25.40, 32.45, 38.45
[1596750138] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 56.75, 49.51, 43.71
[1596750738] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 46.08, 43.08, 42.37
[1596750919] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (194188 kB) free!
[1596751040] SERVICE ALERT: sservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (200352 kB) free!
[1596751160] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 47.9% (11740476 kB) free.
[1596751338] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 56.38, 47.89, 44.88
[1596751938] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 51.71, 46.24, 44.53
[1596753138] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 38.29, 47.79, 46.43
[1596753499] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16082.32, 10167.58, 4411.01
[1596753619] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16184.56, 12163.70, 5840.47
[1596753738] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 25.50, 33.29, 39.96
[1596753739] SERVICE ALERT: sservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16198.45, 13499.37, 7096.61
[1596755690] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 1 warn out of 320 processes
[1596755810] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 320 processes
[1596755927] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756047] SERVICE ALERT: pservercluster2;CPU Load;UNKNOWN;SOFT;2;Error opening /usr/bin/uptime
[1596756144] SERVICE ALERT: pservercluster2;HDD_Root status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596756177] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756181] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756206] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756226] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756243] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756265] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756274] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756275] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756279] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756311] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756336] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756340] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756345] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756350] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756356] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756373] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756395] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756395] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756404] SERVICE ALERT: pservercluster2;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756405] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756409] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756441] SERVICE ALERT: pservercluster2;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756466] SERVICE ALERT: pservercluster2;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756470] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756475] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756480] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756486] SERVICE ALERT: pservercluster2;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756503] SERVICE ALERT: pservercluster2;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756525] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756525] SERVICE ALERT: pservercluster2;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756535] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756539] SERVICE ALERT: pservercluster2;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756600] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756605] SERVICE ALERT: pservercluster2;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756610] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756655] SERVICE ALERT: pservercluster2;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1596756738] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 59.18, 46.86, 42.97
[1596756739] SERVICE ALERT: sservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.86, 4.35, 678.79
[1596757338] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 34.79, 43.15, 43.12
[1596757339] SERVICE ALERT: sservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.70, 1.76, 356.28
[1596758804] SERVICE ALERT: pservercluster2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 38286 MB (74% inode=98%):
[1596758866] SERVICE ALERT: pservercluster2;zombie_procs;OK;HARD;1;PROCS OK: 1 process with STATE = Z
[1596758891] SERVICE ALERT: pservercluster2;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=1.24% system=3.61% iowait=0.00% idle=95.15%
[1596758903] SERVICE ALERT: pservercluster2;Current Users;OK;HARD;1;USERS OK - 2 users currently logged in
[1596758925] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 332 processes
[1596758935] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 332 processes
[1596758939] SERVICE ALERT: pservercluster2;Open-Files;OK;HARD;1;OK: 10880 open files (0% of max 2424156)
[1596759000] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 85.2% (20882536 kB) free.
[1596759005] SERVICE ALERT: pservercluster2;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 176666 MB (78% inode=99%):
[1596759010] SERVICE ALERT: pservercluster2;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1596759055] SERVICE ALERT: pservercluster2;Sensors Status;OK;HARD;1;SENSORS OK
[1596759138] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 63.98, 48.25, 44.39
[1596759177] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.90, 10.93, 12.38
[1596759441] SERVICE ALERT: pservercluster2;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596759738] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 38.60, 43.52, 43.15
[1596759777] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 13619.43, 4993.71, 1817.12
[1596759897] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 15852.93, 8702.01, 3561.20
[1596760017] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16153.66, 11183.12, 5093.72
[1596760135] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;1;PROCS CRITICAL: 950 processes
[1596760255] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;SOFT;2;PROCS CRITICAL: 1028 processes
[1596760375] SERVICE ALERT: pservercluster2;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 654 processes
[1596760938] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 48.39, 49.20, 46.51
[1596760957] SERVICE ALERT: pservercluster1;Messages Log Status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596760975] SERVICE FLAPPING ALERT: pservercluster2;Total Processes;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1596760975] SERVICE ALERT: pservercluster2;Total Processes;OK;HARD;1;PROCS OK: 332 processes
[1596761025] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761054] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 11.7% (15468500 kB) free!
[1596761076] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761082] SERVICE ALERT: pservercluster1;Messages Log Status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761150] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761173] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761178] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761186] SERVICE ALERT: pservercluster1;Sensors Status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596761196] SERVICE ALERT: pservercluster1;Memory;UNKNOWN;SOFT;2;NRPE: Call to fork() failed
[1596761202] SERVICE ALERT: pservercluster1;Messages Log Status;UNKNOWN;HARD;3;NRPE: Call to fork() failed
[1596761221] SERVICE ALERT: pservercluster1;HDD_Home status;UNKNOWN;SOFT;1;NRPE: Call to fork() failed
[1596761273] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761275] SERVICE ALERT: pservercluster1;Total Processes;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761293] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761294] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761303] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761306] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761310] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761321] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761342] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761346] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761396] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761398] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761414] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761418] SERVICE ALERT: pservercluster1;zombie_procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761426] SERVICE ALERT: pservercluster1;Sensors Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761428] SERVICE ALERT: pservercluster1;CPU Load;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761430] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761435] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761466] SERVICE ALERT: pservercluster1;HDD_Home status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761467] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761518] SERVICE ALERT: pservercluster1;Current Users;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761521] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761539] SERVICE ALERT: pservercluster1;Open-Files;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761555] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761560] SERVICE ALERT: pservercluster1;CPU-STATISTICS;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761592] SERVICE ALERT: pservercluster1;HDD_Root status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761646] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761654] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (456208 kB) free!
[1596761680] SERVICE ALERT: pservercluster1;CPU_Procs;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596761807] SERVICE ALERT: pservercluster1;Messages Log Status;CRITICAL;HARD;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1596762138] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 33.31, 37.39, 43.19
[1596762475] SERVICE ALERT: pservercluster1;Total Processes;OK;HARD;1;PROCS OK: 315 processes
[1596762521] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 84.8% (20931000 kB) free.
[1596762618] SERVICE ALERT: pservercluster1;zombie_procs;OK;HARD;1;PROCS OK: 0 processes with STATE = Z
[1596762626] SERVICE ALERT: pservercluster1;Sensors Status;OK;HARD;1;SENSORS OK
[1596762628] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 28.73, 102.84, 54.74
[1596762666] SERVICE ALERT: pservercluster1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 197706 MB (89% inode=99%):
[1596762718] SERVICE ALERT: pservercluster1;Current Users;OK;HARD;1;USERS OK - 1 users currently logged in
[1596762738] SERVICE ALERT: mrq4;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 61.42, 53.16, 48.67
[1596762739] SERVICE ALERT: pservercluster1;Open-Files;OK;HARD;1;OK: 8288 open files (0% of max 2440681)
[1596762766] SERVICE ALERT: pservercluster1;CPU-STATISTICS;OK;HARD;1;CPU STATISTICS OK: user=1.13% system=3.94% iowait=0.00% idle=94.93%
[1596762792] SERVICE ALERT: pservercluster1;HDD_Root status;OK;HARD;1;DISK OK - free space: / 39505 MB (77% inode=99%):
[1596762846] SERVICE ALERT: pservercluster1;LocalService-NTPD SERVICE;OK;HARD;1;PROCS OK: 1 process with command name 'ntpd'
[1596762881] SERVICE ALERT: pservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 313 processes
[1596763008] SERVICE ALERT: pservercluster1;Messages Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596763017] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1596763017] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 1.02, 4.82, 679.23
[1596763454] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 60.8% (80090244 kB) free.
[1596763617] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.97, 1.51, 356.42
[1596772375] SERVICE FLAPPING ALERT: pservercluster2;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596773810] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 321 processes
[1596773930] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1596774138] SERVICE FLAPPING ALERT: mrq4;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596775017] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596779158] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 19.7% (4828744 kB) free!
[1596779278] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.7% (4818432 kB) free!
[1596779398] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 21.1% (5180880 kB) free.
[1596779720] SERVICE ALERT: primary-server;Memory;OK;HARD;1;OK - 95.5% (23396868 kB) free.
[1596779799] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.5% (38669736 kB) free.
[1596781320] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 3 crit, 0 warn out of 335 processes
[1596781441] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 3 crit, 0 warn out of 341 processes
[1596781550] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (177204 kB) free!
[1596781561] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 3 crit, 0 warn out of 334 processes
[1596781670] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (172684 kB) free!
[1596781730] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 330 processes
[1596781790] SERVICE ALERT: sservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188096 kB) free!
[1596781850] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 318 processes
[1596782398] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;1;WARNING - 20.0% (4897568 kB) free!
[1596782519] SERVICE ALERT: secondary-server;Memory Status;WARNING;SOFT;2;WARNING - 19.9% (4874112 kB) free!
[1596782639] SERVICE ALERT: secondary-server;Memory Status;WARNING;HARD;3;WARNING - 19.8% (4859796 kB) free!
[1596782990] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 34.2% (8374460 kB) free.
[1596783590] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (176420 kB) free!
[1596783711] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 57.4% (14081104 kB) free.
[1596784310] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 7.3% (1781188 kB) free!
[1596784430] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 6.1% (1493568 kB) free!
[1596784550] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 79.2% (19420888 kB) free.
[1596784561] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 329 processes
[1596785039] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 28.1% (6886036 kB) free.
[1596785161] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 2 crit, 0 warn out of 332 processes
[1596785282] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 332 processes
[1596785350] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 313 processes
[1596785402] SERVICE ALERT: sservercluster2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 339 processes
[1596785470] SERVICE ALERT: secondary-server;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 315 processes
[1596785590] SERVICE ALERT: secondary-server;CPU_Procs;OK;HARD;1;CPU OK: 312 processes
[1596785639] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;1;CRITICAL - 0.7% (174940 kB) free!
[1596785750] SERVICE ALERT: sservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.8% (2403616 kB) free!
[1596785759] SERVICE ALERT: secondary-server;Memory Status;CRITICAL;SOFT;2;CRITICAL - 0.8% (188992 kB) free!
[1596785870] SERVICE ALERT: sservercluster2;Memory;OK;HARD;1;OK - 24.0% (5872284 kB) free.
[1596785879] SERVICE ALERT: secondary-server;Memory Status;OK;HARD;1;OK - 87.9% (21537088 kB) free.
[1596786002] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1596786002] SERVICE ALERT: sservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 326 processes
[1596788545] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 20.0% (6552168 kB) free!
[1596788665] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 19.8% (6468724 kB) free!
[1596788785] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 19.6% (6407488 kB) free!
[1596789385] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 21.4% (7012676 kB) free.
[1596789650] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 320 processes
[1596789770] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 316 processes
[1596790585] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;1;WARNING - 19.2% (6290912 kB) free!
[1596790705] SERVICE ALERT: mrq4;Memory;WARNING;SOFT;2;WARNING - 19.0% (6228300 kB) free!
[1596790825] SERVICE ALERT: mrq4;Memory;WARNING;HARD;3;WARNING - 18.8% (6160340 kB) free!
[1596793058] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 19.1% (9430436 kB) free!
[1596793178] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.9% (9312400 kB) free!
[1596793298] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.6% (9165504 kB) free!
[1596793608] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 16.2% (7979684 kB) free!
[1596793728] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 15.2% (7467208 kB) free!
[1596793849] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 16.6% (8189500 kB) free!
[1596794422] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(40) < 182.18.184.244 [07/Aug/2020:15:24:03 +0530] TCP 502 0 0 0.000 "10.147.212.49:5501" "0" "0" "0.000"
[1596794542] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596795649] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 3.9% (1937036 kB) free!
[1596795770] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 4 warn out of 328 processes
[1596795890] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 2 warn out of 329 processes
[1596796010] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 320 processes
[1596797210] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 328 processes
[1596797330] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 327 processes
[1596797403] SERVICE FLAPPING ALERT: sservercluster2;CPU_Procs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1596798098] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.9% (4888184 kB) free!
[1596798142] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(16) < 182.18.184.244 [07/Aug/2020:16:25:44 +0530] TCP 502 0 0 0.000 "10.147.212.49:5501" "0" "0" "0.000"
[1596798262] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596799017] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 16120.40, 10650.82, 4723.86
[1596799130] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 320 processes
[1596799137] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 16187.39, 12486.19, 6115.07
[1596799250] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 321 processes
[1596799257] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 16196.87, 13714.42, 7337.65
[1596799850] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 324 processes
[1596799971] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 321 processes
[1596801163] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 75.5% (24730344 kB) free.
[1596801243] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 78.4% (38651248 kB) free.
[1596801356] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 78.4% (38647376 kB) free.
[1596803570] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 324 processes
[1596803690] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 317 processes
[1596804657] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.00, 3.61, 633.76
[1596809090] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 2 warn out of 321 processes
[1596809210] SERVICE ALERT: webservercluster1;CPU_Procs;WARNING;SOFT;2;CPU WARNING: 1 warn out of 322 processes
[1596809330] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 320 processes
[1596810262] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(5) < 103.255.146.164 [07/Aug/2020:19:54:09 +0530] TCP 502 0 0 0.001 "10.147.212.142:3000" "0" "0" "0.001"
[1596810383] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596813178] SERVICE ALERT: mrq4;Memory;CRITICAL;SOFT;1;CRITICAL - 5.8% (1893232 kB) free!
[1596813298] SERVICE ALERT: mrq4;Memory;CRITICAL;SOFT;2;CRITICAL - 5.7% (1861968 kB) free!
[1596813418] SERVICE ALERT: mrq4;Memory;CRITICAL;HARD;3;CRITICAL - 5.6% (1831484 kB) free!
[1596814618] SERVICE ALERT: mrq4;Memory;OK;HARD;1;OK - 44.8% (14680792 kB) free.
[1596815055] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (9775280 kB) free!
[1596815175] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 18.6% (9149468 kB) free!
[1596815296] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 18.8% (9242688 kB) free!
[1596818178] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1517 processes
[1596818299] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1517 processes
[1596818419] SERVICE ALERT: sigqscluster1;Total Processes;WARNING;HARD;3;PROCS WARNING: 1510 processes
[1596820221] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;1;WARNING - 19.7% (9722380 kB) free!
[1596820341] SERVICE ALERT: qservercluster2;Memory;WARNING;SOFT;2;WARNING - 19.7% (9693868 kB) free!
[1596820461] SERVICE ALERT: qservercluster2;Memory;WARNING;HARD;3;WARNING - 19.6% (9675272 kB) free!
[1596821783] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(240) < 203.212.70.170 [07/Aug/2020:23:05:58 +0530] TCP 502 0 0 0.000 "backend_3001" "0" "0" "0.000"
[1596821903] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596822503] SERVICE ALERT: primary-server;Nginx Access Log Status;CRITICAL;SOFT;1;(637) < 172.105.59.173 [07/Aug/2020:23:12:24 +0530] TCP 502 0 0 0.000 "backend_13001" "0" "0" "0.000"
[1596822624] SERVICE ALERT: primary-server;Nginx Access Log Status;OK;HARD;1;Log check ok - 0 pattern matches found
[1596824538] SERVICE ALERT: mrq4;CPU Load;WARNING;HARD;3;WARNING - load average: 45.01, 38.37, 41.48