[1569350892] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569351492] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569351675] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 585 processes
[1569351735] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 581 processes
[1569352092] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569352558] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1569352678] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1569353278] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1569353293] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569353398] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1569353515] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 36.5% (9013216 kB) free.
[1569353641] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 366 processes
[1569353750] SERVICE ALERT: sig5-vm2;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 2 warn out of 279 processes
[1569353761] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 360 processes
[1569353857] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;1;CRITICAL - 9.5% (3323476 kB) free!
[1569353893] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569353977] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.6% (224896 kB) free!
[1569353986] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;1;DISK WARNING - free space: /vns1 173004 MB (16% inode=99%):
[1569354097] SERVICE ALERT: sig5-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (235400 kB) free!
[1569354106] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;SOFT;2;DISK WARNING - free space: /vns1 163314 MB (15% inode=99%):
[1569354116] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (171884 kB) free!
[1569354226] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 153601 MB (14% inode=99%):
[1569354236] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176944 kB) free!
[1569354356] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171736 kB) free!
[1569354360] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (1086292 kB) free!
[1569354480] SERVICE ALERT: redis-1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.3% (404176 kB) free!
[1569354600] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (393000 kB) free!
[1569354826] SERVICE ALERT: redis-1;HDD_VNS1 status;CRITICAL;HARD;3;DISK CRITICAL - free space: /vns1 110711 MB (10% inode=99%):
[1569354950] SERVICE ALERT: sig5-vm2;CPU_Procs;WARNING;HARD;3;CPU WARNING: 1 warn out of 275 processes
[1569355093] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569355693] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569356807] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;1;CRITICAL - load average: 94.22, 46.77, 18.86
[1569356894] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569356927] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;SOFT;2;CRITICAL - load average: 97.82, 63.88, 28.52
[1569356999] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1569357047] SERVICE ALERT: pservercluster2;CPU Load;CRITICAL;HARD;3;CRITICAL - load average: 98.28, 75.32, 37.00
[1569357119] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 403 processes
[1569357494] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569357657] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 10.3% (2517568 kB) free!
[1569358094] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 606 processes
[1569358247] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.81, 2.63, 12.98
[1569358675] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (232412 kB) free!
[1569358795] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (239084 kB) free!
[1569358857] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (2443884 kB) free!
[1569358915] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (234700 kB) free!
[1569359447] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 0.94, 0.95, 4.25
[1569359515] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 87.6% (43153612 kB) free.
[1569359894] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569360495] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 603 processes
[1569361257] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 48.5% (11886536 kB) free.
[1569361319] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 403 processes
[1569361439] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1569361449] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 48.2% (11812388 kB) free.
[1569362895] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569364096] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569364696] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569365896] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569366496] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569367096] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569367300] SERVICE ALERT: webservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 310 processes
[1569367420] SERVICE ALERT: webservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 303 processes
[1569368297] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569369497] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569370097] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569370247] SERVICE FLAPPING ALERT: pservercluster2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1569370697] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569371497] SERVICE ALERT: sig5-vm2;Memory;OK;HARD;1;OK - 94.4% (33076888 kB) free.
[1569371898] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569372000] SERVICE ALERT: redis-1;Memory;WARNING;HARD;3;WARNING - 13.6% (17900408 kB) free!
[1569372226] SERVICE ALERT: redis-1;HDD_VNS1 status;WARNING;HARD;3;DISK WARNING - free space: /vns1 125015 MB (11% inode=99%):
[1569372498] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569372600] SERVICE ALERT: redis-1;Memory;CRITICAL;HARD;3;CRITICAL - 0.3% (372772 kB) free!
[1569373698] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569374157] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 64.7% (15962568 kB) free.
[1569374236] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 510 processes
[1569374296] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 515 processes
[1569374400] SERVICE ALERT: redis-1;Memory;OK;HARD;1;OK - 72.1% (95082232 kB) free.
[1569374757] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (191036 kB) free!
[1569374877] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (195612 kB) free!
[1569374898] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 606 processes
[1569374997] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172344 kB) free!
[1569375059] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (183536 kB) free!
[1569375180] SERVICE ALERT: pservercluster2;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (163796 kB) free!
[1569375249] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 19.8% (4849824 kB) free!
[1569375301] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (174656 kB) free!
[1569375370] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 19.7% (4827288 kB) free!
[1569375490] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 19.6% (4811636 kB) free!
[1569375499] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 598 processes
[1569376099] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 601 processes
[1569376222] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9873560 kB) free!
[1569376343] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 19.9% (9789200 kB) free!
[1569376463] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 19.8% (9763920 kB) free!
[1569377299] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569378496] SERVICE ALERT: webserver2;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 522 processes
[1569378499] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569378556] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 518 processes
[1569379099] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569380299] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 604 processes
[1569380899] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569381839] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1569381959] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 402 processes
[1569382099] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569383299] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569383899] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 607 processes
[1569385099] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569385699] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 604 processes
[1569386290] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.6% (2103084 kB) free!
[1569386299] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569386516] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 95.2% (23495504 kB) free.
[1569386638] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 94.4% (23135944 kB) free.
[1569386729] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 81.3% (40059812 kB) free.
[1569386779] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 92.5% (22679288 kB) free.
[1569387499] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 604 processes
[1569388431] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.95, 10.00, 7.63
[1569388551] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 10.11, 10.26, 8.02
[1569388671] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 11.15, 10.43, 8.35
[1569389899] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569390728] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;1;WARNING - 14.3% (3519548 kB) free!
[1569390848] SERVICE ALERT: pservercluster1;Memory;WARNING;SOFT;2;WARNING - 14.0% (3454776 kB) free!
[1569390968] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 9.7% (2397700 kB) free!
[1569391099] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569391699] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569392299] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569392513] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;1;USERS WARNING - 7 users currently logged in
[1569392633] SERVICE ALERT: sservercluster1;Current Users;WARNING;SOFT;2;USERS WARNING - 7 users currently logged in
[1569392753] SERVICE ALERT: sservercluster1;Current Users;WARNING;HARD;3;USERS WARNING - 7 users currently logged in
[1569392768] SERVICE ALERT: pservercluster1;Memory;WARNING;HARD;3;WARNING - 17.1% (4221468 kB) free!
[1569392856] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.7% (270700 kB) free!
[1569392976] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (167016 kB) free!
[1569393096] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (148764 kB) free!
[1569393228] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;1;CPU CRITICAL: 1 crit, 0 warn out of 331 processes
[1569393336] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 14.5% (7168068 kB) free!
[1569393349] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;SOFT;2;CPU CRITICAL: 1 crit, 0 warn out of 330 processes
[1569393368] SERVICE FLAPPING ALERT: pservercluster1;Memory;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1569393368] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (188848 kB) free!
[1569393398] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.3% (14196688 kB) free.
[1569393408] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1569393456] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 33.3% (16402856 kB) free.
[1569393468] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 62 users currently logged in
[1569393469] SERVICE ALERT: webservercluster1;CPU_Procs;CRITICAL;HARD;3;CPU CRITICAL: 1 crit, 0 warn out of 335 processes
[1569393482] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 79.3% (19573748 kB) free.
[1569393528] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 62 users currently logged in
[1569393588] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 62 users currently logged in
[1569393718] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 6.70, 6.49, 5.25
[1569393822] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;SOFT;1;DISK WARNING - free space: / 20924 MB (15% inode=99%):
[1569393838] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 8.27, 7.16, 5.65
[1569393942] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;SOFT;2;DISK WARNING - free space: / 20924 MB (15% inode=99%):
[1569393958] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 6.40, 6.95, 5.76
[1569394006] SERVICE ALERT: webservercluster1;Memory;WARNING;SOFT;1;WARNING - 10.1% (1645016 kB) free!
[1569394057] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.5% (244264 kB) free!
[1569394062] SERVICE ALERT: sig2-vm2;HDD_Root status;WARNING;HARD;3;DISK WARNING - free space: / 20926 MB (15% inode=99%):
[1569394086] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (189744 kB) free!
[1569394099] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569394126] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.5% (405556 kB) free!
[1569394177] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (255268 kB) free!
[1569394206] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191476 kB) free!
[1569394246] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (172148 kB) free!
[1569394297] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.4% (213076 kB) free!
[1569394326] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (171676 kB) free!
[1569394699] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 604 processes
[1569394714] SERVICE FLAPPING ALERT: webservercluster1;Memory;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1569394714] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 81.9% (13312148 kB) free.
[1569394776] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 77.6% (19159504 kB) free.
[1569394836] SERVICE ALERT: sigqscluster1;Memory;OK;HARD;1;OK - 80.7% (39788228 kB) free.
[1569395158] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 4.24, 4.37, 4.73
[1569395299] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569395326] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (174548 kB) free!
[1569395383] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.4% (1339744 kB) free!
[1569395446] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (169184 kB) free!
[1569395503] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 2.7% (672828 kB) free!
[1569395515] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9190804 kB) free!
[1569395566] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (180592 kB) free!
[1569395623] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (200380 kB) free!
[1569395635] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 17.7% (8726232 kB) free!
[1569395755] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 16.8% (8300252 kB) free!
[1569395956] SERVICE ALERT: webserver2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 631 processes
[1569395987] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1569396016] SERVICE ALERT: webserver2;CPU_Procs;OK;HARD;1;CPU OK: 637 processes
[1569396107] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1569396227] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1569396499] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569396588] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1569396827] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1569396955] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 6.8% (3332820 kB) free!
[1569396961] SERVICE ALERT: pservercluster2;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 365 processes
[1569396963] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 85.0% (13824296 kB) free.
[1569397004] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 76.5% (18889940 kB) free.
[1569397062] SERVICE ALERT: sig2-vm2;HDD_Root status;CRITICAL;HARD;3;DISK CRITICAL - free space: / 12460 MB (9% inode=99%):
[1569397081] SERVICE ALERT: pservercluster2;CPU_Procs;OK;HARD;1;CPU OK: 365 processes
[1569397091] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 96.0% (47300292 kB) free.
[1569397427] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1569397547] SERVICE ALERT: secondary-server;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1569397563] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (150156 kB) free!
[1569397616] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (195632 kB) free!
[1569397667] SERVICE ALERT: secondary-server;Current Users;WARNING;HARD;3;USERS WARNING - 6 users currently logged in
[1569397683] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (190668 kB) free!
[1569397736] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (167920 kB) free!
[1569397803] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (172304 kB) free!
[1569397856] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (172652 kB) free!
[1569397868] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 86.0% (13979456 kB) free.
[1569398388] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 65 users currently logged in
[1569398475] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (169476 kB) free!
[1569398595] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (167544 kB) free!
[1569398688] SERVICE FLAPPING ALERT: webserver2;Current Users;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1569398688] SERVICE ALERT: webserver2;Current Users;CRITICAL;HARD;4;USERS CRITICAL - 66 users currently logged in
[1569398715] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (169252 kB) free!
[1569399245] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.8% (4359740 kB) free!
[1569399365] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 17.4% (4258744 kB) free!
[1569399485] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 17.1% (4185628 kB) free!
[1569400071] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 3.09, 3.79, 4.88
[1569400508] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.3% (14190152 kB) free.
[1569400574] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 73.9% (18226120 kB) free.
[1569400631] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 93.3% (22860580 kB) free.
[1569400699] SERVICE ALERT: smscapp1;Total Processes;WARNING;HARD;3;PROCS WARNING: 600 processes
[1569400788] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 63 users currently logged in
[1569401116] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (153796 kB) free!
[1569401153] SERVICE ALERT: sservercluster1;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1569401184] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 5.7% (1410128 kB) free!
[1569401236] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.0% (163420 kB) free!
[1569401304] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 3.2% (799836 kB) free!
[1569401356] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (148756 kB) free!
[1569401408] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.8% (14272432 kB) free.
[1569401424] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197516 kB) free!
[1569401502] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 76.6% (18894612 kB) free.
[1569401867] SERVICE FLAPPING ALERT: secondary-server;Current Users;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1569401867] SERVICE ALERT: secondary-server;Current Users;OK;HARD;1;USERS OK - 5 users currently logged in
[1569401899] SERVICE ALERT: smscapp1;Total Processes;CRITICAL;HARD;3;PROCS CRITICAL: 602 processes
[1569402012] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (169672 kB) free!
[1569402105] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 4.7% (1149412 kB) free!
[1569402133] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (171596 kB) free!
[1569402225] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.5% (379156 kB) free!
[1569402253] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (165764 kB) free!
[1569402345] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (183360 kB) free!
[1569402888] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 54 users currently logged in
[1569404182] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;1;WARNING - 18.4% (4521940 kB) free!
[1569404262] SERVICE ALERT: sig2-vm2;HDD_Root status;OK;HARD;1;DISK OK - free space: / 113126 MB (81% inode=99%):
[1569404302] SERVICE ALERT: sservercluster1;Memory;WARNING;SOFT;2;WARNING - 18.1% (4425152 kB) free!
[1569404422] SERVICE ALERT: sservercluster1;Memory;WARNING;HARD;3;WARNING - 17.7% (4340352 kB) free!
[1569404626] SERVICE ALERT: redis-1;HDD_VNS1 status;OK;HARD;1;DISK OK - free space: /vns1 513334 MB (48% inode=99%):
[1569404688] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1569404729] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 68.8% (16973108 kB) free.
[1569404748] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1569404770] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.9% (14279616 kB) free.
[1569404808] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1569404868] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1569405168] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 59 users currently logged in
[1569405334] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 3.7% (922720 kB) free!
[1569405374] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.9% (153040 kB) free!
[1569405454] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.3% (313156 kB) free!
[1569405494] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.2% (195420 kB) free!
[1569405574] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (197076 kB) free!
[1569405614] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (178172 kB) free!
[1569406668] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;1;USERS WARNING - 61 users currently logged in
[1569406728] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;2;USERS WARNING - 61 users currently logged in
[1569406788] SERVICE ALERT: webserver2;Current Users;WARNING;SOFT;3;USERS WARNING - 61 users currently logged in
[1569406848] SERVICE ALERT: webserver2;Current Users;WARNING;HARD;4;USERS WARNING - 61 users currently logged in
[1569407158] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 7.47, 7.31, 6.16
[1569407271] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 15.98, 10.40, 6.37
[1569407278] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 7.73, 7.39, 6.32
[1569407391] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 9.57, 9.91, 6.68
[1569407398] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 8.35, 7.65, 6.54
[1569407422] SERVICE ALERT: sservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 8.9% (2179060 kB) free!
[1569407511] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 13.29, 10.92, 7.43
[1569409667] SERVICE FLAPPING ALERT: secondary-server;Current Users;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1569411348] SERVICE FLAPPING ALERT: webserver2;Current Users;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1569411614] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 26.3% (4278044 kB) free.
[1569412214] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.2% (189372 kB) free!
[1569412334] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (174672 kB) free!
[1569412443] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 75.9% (18735080 kB) free.
[1569412454] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.1% (177232 kB) free!
[1569412481] SERVICE ALERT: sservercluster1;Memory;OK;HARD;1;OK - 94.7% (23206800 kB) free.
[1569412519] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 87.1% (14151544 kB) free.
[1569412699] SERVICE FLAPPING ALERT: smscapp1;Total Processes;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1569413045] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (170228 kB) free!
[1569413047] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;1;WARNING - 15.6% (3678108 kB) free!
[1569413124] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (159688 kB) free!
[1569413165] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (176364 kB) free!
[1569413167] SERVICE ALERT: sig3-vm2;Memory;WARNING;SOFT;2;WARNING - 14.3% (3362048 kB) free!
[1569413246] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (173340 kB) free!
[1569413274] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 3.84, 14.29, 8.41
[1569413286] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190992 kB) free!
[1569413287] SERVICE ALERT: sig3-vm2;Memory;WARNING;HARD;3;WARNING - 13.0% (3069392 kB) free!
[1569413369] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.9% (149732 kB) free!
[1569413394] SERVICE ALERT: sservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 1.58, 9.93, 7.52
[1569413447] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 1.65, 6.09, 5.49
[1569413514] SERVICE ALERT: sservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 0.79, 6.90, 6.70
[1569413567] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.61, 4.55, 5.00
[1569413887] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;HARD;3;CRITICAL - 8.1% (1913628 kB) free!
[1569413969] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 68.0% (11046192 kB) free.
[1569414167] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;1;WARNING - load average: 0.99, 8.45, 9.19
[1569414287] SERVICE ALERT: pservercluster2;CPU Load;WARNING;SOFT;2;WARNING - load average: 0.97, 5.98, 8.20
[1569414407] SERVICE ALERT: pservercluster2;CPU Load;WARNING;HARD;3;WARNING - load average: 0.69, 4.25, 7.30
[1569414495] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;1;WARNING - 18.6% (9150460 kB) free!
[1569414569] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.0% (156472 kB) free!
[1569414615] SERVICE ALERT: smscapp1;Memory;WARNING;SOFT;2;WARNING - 18.4% (9062068 kB) free!
[1569414648] SERVICE ALERT: webserver2;Current Users;OK;HARD;1;USERS OK - 59 users currently logged in
[1569414689] SERVICE ALERT: webservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 1.1% (173276 kB) free!
[1569414711] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 2.65, 2.76, 4.30
[1569414714] SERVICE ALERT: sservercluster1;CPU Load;OK;HARD;1;OK - load average: 1.03, 1.90, 4.85
[1569414735] SERVICE ALERT: smscapp1;Memory;WARNING;HARD;3;WARNING - 18.2% (8949900 kB) free!
[1569414809] SERVICE ALERT: webservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 1.0% (157524 kB) free!
[1569415007] SERVICE ALERT: pservercluster2;CPU Load;OK;HARD;1;OK - load average: 1.08, 1.42, 4.32
[1569415088] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 97.3% (22922644 kB) free.
[1569415559] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 404 processes
[1569415679] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 404 processes
[1569415698] SERVICE ALERT: sig3-vm2;Memory;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1569415824] SERVICE ALERT: sig3-vm2;Memory;OK;HARD;1;OK - 96.7% (22793428 kB) free.
[1569415911] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;1;WARNING - load average: 8.34, 10.34, 7.71
[1569416031] SERVICE ALERT: smscapp2;CPU Load;WARNING;SOFT;2;WARNING - load average: 6.06, 8.59, 7.38
[1569416151] SERVICE ALERT: smscapp2;CPU Load;WARNING;HARD;3;WARNING - load average: 5.96, 7.65, 7.17
[1569416535] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 9.2% (4537560 kB) free!
[1569418034] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;1;WARNING - 17.5% (4281876 kB) free!
[1569418069] SERVICE ALERT: webservercluster1;CPU_Procs;OK;HARD;1;CPU OK: 333 processes
[1569418155] SERVICE ALERT: pservercluster2;Memory;WARNING;SOFT;2;WARNING - 16.8% (4114620 kB) free!
[1569418275] SERVICE ALERT: pservercluster2;Memory;WARNING;HARD;3;WARNING - 15.5% (3806364 kB) free!
[1569418798] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 3.49, 3.56, 4.37
[1569419475] SERVICE ALERT: pservercluster2;Memory;CRITICAL;HARD;3;CRITICAL - 8.3% (2028796 kB) free!
[1569419840] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 62.5% (15411804 kB) free.
[1569419870] SERVICE ALERT: pservercluster2;Memory;OK;HARD;1;OK - 92.6% (22702148 kB) free.
[1569419901] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 39.9% (19642792 kB) free.
[1569419935] SERVICE ALERT: webservercluster1;Memory;OK;HARD;1;OK - 95.5% (15530056 kB) free.
[1569420443] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.7% (175784 kB) free!
[1569420563] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.8% (191912 kB) free!
[1569420643] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;1;WARNING - 20.0% (9835036 kB) free!
[1569420684] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.8% (190528 kB) free!
[1569420763] SERVICE ALERT: sigqscluster1;Memory;WARNING;SOFT;2;WARNING - 20.0% (9834628 kB) free!
[1569420883] SERVICE ALERT: sigqscluster1;Memory;WARNING;HARD;3;WARNING - 20.0% (9835608 kB) free!
[1569420968] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 45477 MB (20% inode=99%):
[1569421088] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 45337 MB (20% inode=99%):
[1569421208] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 45203 MB (20% inode=99%):
[1569421551] SERVICE ALERT: smscapp2;CPU Load;OK;HARD;1;OK - load average: 1.47, 2.50, 4.34
[1569421808] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 49148 MB (21% inode=99%):
[1569423608] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;1;DISK WARNING - free space: /home 47060 MB (20% inode=99%):
[1569423728] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;SOFT;2;DISK WARNING - free space: /home 47041 MB (20% inode=99%):
[1569423848] SERVICE ALERT: smscapp1;HDD_Home status;WARNING;HARD;3;DISK WARNING - free space: /home 46974 MB (20% inode=99%):
[1569425346] SERVICE FLAPPING ALERT: webservercluster1;Memory;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1569427329] SERVICE ALERT: pservercluster1;Memory;OK;HARD;1;OK - 69.7% (17207064 kB) free.
[1569427798] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;1;WARNING - load average: 4.93, 12.61, 11.48
[1569427918] SERVICE ALERT: pservercluster1;CPU Load;WARNING;SOFT;2;WARNING - load average: 4.44, 9.90, 10.62
[1569427933] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;1;CRITICAL - 0.8% (186164 kB) free!
[1569428038] SERVICE ALERT: pservercluster1;CPU Load;WARNING;HARD;3;WARNING - load average: 4.48, 8.14, 9.89
[1569428053] SERVICE ALERT: pservercluster1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.7% (175688 kB) free!
[1569428173] SERVICE ALERT: pservercluster1;Memory;CRITICAL;HARD;3;CRITICAL - 0.7% (175052 kB) free!
[1569429838] SERVICE ALERT: pservercluster1;CPU Load;OK;HARD;1;OK - load average: 2.02, 3.41, 5.00
[1569430104] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;1;CRITICAL - 1.1% (547688 kB) free!
[1569430224] SERVICE ALERT: smscapp1;Memory;CRITICAL;SOFT;2;CRITICAL - 0.5% (233052 kB) free!
[1569430344] SERVICE ALERT: smscapp1;Memory;CRITICAL;HARD;3;CRITICAL - 0.5% (243332 kB) free!
[1569431048] SERVICE ALERT: smscapp1;HDD_Home status;OK;HARD;1;DISK OK - free space: /home 111050 MB (49% inode=99%):
[1569431544] SERVICE ALERT: smscapp1;Memory;OK;HARD;1;OK - 88.7% (43693832 kB) free.
[1569434683] SERVICE ALERT: sigqscluster1;Memory;CRITICAL;HARD;3;CRITICAL - 10.0% (4913720 kB) free!
[1569435480] SERVICE ALERT: redis-1;CPU_Procs;WARNING;SOFT;1;CPU WARNING: 1 warn out of 402 processes
[1569435600] SERVICE ALERT: redis-1;CPU_Procs;OK;HARD;1;CPU OK: 400 processes